<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Int_Reciprocal_Sqrt_2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Complete the square in the quadratic expression q which is given by \[ q= {@q1@}, \] and use this to evaluate the following integral \[ \int {@f1@}d{@v1@} \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Completing the square for \(q\) gives 
\[q={@n1^2@}{@v1*v1@} + {@2*n0*n1@}{@v1@} + {@n0^2+c0^2@}={@n1^2@}{@v1*v1@}+ {@2*n1*n0@}{@v1@}+{@n0^2@}+{@c0^2@}=({@n1@}{@v1@}+{@n0@})^2 + {@c0^2@}\]</p>
<p>Therefore \[\int{@f1@}d{@v1@} = \int\frac{1}{@sqrt((n1*v1+n0)^2 + c0^2)@}d{@v1@}=\int\frac{1}{@sqrt(c0^2*((n1*v1/c0+n0/c0)^2 +1))@}d{@v1@}\]</p>
<p> </p>
<p>Let \({@n1*v1/c0@}+{@n0/c0@} = \sinh(u)\)</p>
<p>giving \[d{@v1@} = {@c0/n1*cosh(u)@}du\]</p>
<p>Therefore \[\int{@f1@}d{@v1@} = \int\frac{1}{@c0*sqrt(sinh(u)^2+1)@}{@c0/n1*cosh(u)@}du\]</p>
<p>and since \(\cosh^2(u)-\sinh^2(u) = 1\)</p>
<p>\[\begin{array}{rcl}\displaystyle \int\frac{1}{@sqrt(sinh(u)^2+1)@}{@cosh(u)/n1@}du & \displaystyle =&\displaystyle\int\frac{1}{@cosh(u)@}{@cosh(u)/n1@}du\\  \int \frac{du}{@n1@} & \displaystyle =& \displaystyle \frac{u}{@n1@} + c \\ & \displaystyle =&\displaystyle {@correctanswer@} \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n0:rand(5)+1;
 n1:rand(4)+2;
 c0:rand(3)+2;
 v1:x;
 q1:expand((n1*v1+n0)^2 + c0^2);
 f1:1/sqrt(q1);
 correctanswer:expand(int(f1,v1))+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \int {@f1@} d{@v1@}\) gives {@correctanswer@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>integrate</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>3.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions> v1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(f1,v1)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2107951149</deployedseed>
    <deployedseed>836782808</deployedseed>
    <deployedseed>150041981</deployedseed>
    <deployedseed>86879394</deployedseed>
    <deployedseed>1786876046</deployedseed>
    <deployedseed>1166400803</deployedseed>
    <deployedseed>74545899</deployedseed>
    <deployedseed>790409795</deployedseed>
    <deployedseed>82965218</deployedseed>
    <deployedseed>236156191</deployedseed>
    <deployedseed>1808925781</deployedseed>
    <deployedseed>633916508</deployedseed>
    <deployedseed>478725008</deployedseed>
    <deployedseed>1084170933</deployedseed>
    <deployedseed>68939336</deployedseed>
    <deployedseed>1865600840</deployedseed>
    <deployedseed>1038990469</deployedseed>
    <deployedseed>893906692</deployedseed>
    <deployedseed>1737650300</deployedseed>
    <deployedseed>1385164720</deployedseed>
    <deployedseed>1070120215</deployedseed>
    <deployedseed>743119025</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>diff(f1,v1)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer-c</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
