<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Int_x_Substitution_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the following integral, using a substitution if necessary: \[ \int {@k@}{@v@}^{@n1@}({@v@}^{@k@}-{@n3@})^\frac{@n2@}{@m@} d{@v@} \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Using a substitution, let \[u={@v@}^{@k@}-{@n3@}\]</p>
<p>Differentiating \(u\) with respect to {@v@} gives \[du={@k@}{@v@}^{@n1@}d{@v@}\]</p>
<p>Therefore</p>
<p>\[\begin{array}{rcl}\displaystyle \int {@k@}{@v@}^{@n1@}({@v@}^{@k@}-{@n3@})^\frac{@n2@}{@m@}\ d{@v@} &amp;\displaystyle = &amp;\displaystyle\int u^{@p@} du \\ &amp;\displaystyle = &amp;\displaystyle {@q@}u^{@(1/q)@}+c  \\ &amp;\displaystyle = &amp;\displaystyle {@q@}({@v@}^{@k@}-{@n3@})^{@(1/q)@}+c \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(5)+2;
  n2:rand(8)+1;
  n3:rand(9)+1;
  v:x;
  k:n1+1;
  m:2*n2+1;
  f:k*v^n1*(v^k-n3)^(n2/m);
  p:n2/m; 
  q:1/(p+1);
  correctanswer:int(f,v)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \int {@k@}{@v@}^{@n1@}({@v@}^{@k@}-{@n3@})^\frac{@n2@}{@m@} d{@v@} = {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions> v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(f,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>475904532</deployedseed>
    <deployedseed>442591240</deployedseed>
    <deployedseed>1544795150</deployedseed>
    <deployedseed>1494760574</deployedseed>
    <deployedseed>794482547</deployedseed>
    <deployedseed>956724280</deployedseed>
    <deployedseed>1683917961</deployedseed>
    <deployedseed>2011672903</deployedseed>
    <deployedseed>329520694</deployedseed>
    <deployedseed>1458655084</deployedseed>
    <deployedseed>1233637054</deployedseed>
    <deployedseed>532774871</deployedseed>
    <deployedseed>2075261859</deployedseed>
    <deployedseed>888835940</deployedseed>
    <deployedseed>1736607010</deployedseed>
    <deployedseed>125885998</deployedseed>
    <deployedseed>1155381478</deployedseed>
    <deployedseed>738280311</deployedseed>
    <deployedseed>2021720405</deployedseed>
    <deployedseed>1361592244</deployedseed>
    <deployedseed>547086608</deployedseed>
    <deployedseed>1979260678</deployedseed>
    <deployedseed>1105704938</deployedseed>
    <deployedseed>1301288034</deployedseed>
    <deployedseed>166194026</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>diff(f,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer-c</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
