<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Vectors_add columns</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the folowing expression: <br /> \[ \left( ({@n3@}-{@n4@}){@c1@} - {@n2@} {@c2@}\right) + {@n1@} {@c3@}.\] <br /> [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\[\left( ({@n3@}-{@n4@}){@c1@} - {@n2@} {@c2@}\right) + {@n1@} {@c3@} = \left({@n3-n4@}{@c1@} - {@n2@} {@c2@}\right) + {@n1@} {@c3@} \] \[ =\left( {@(n3-n4)*c1@} - {@n2*c2@}\right) + {@n1*c3@} = {@correctanswer@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>c1:matrix([rand_with_step(-9,9,1)],[rand_with_step(-9,9,1)],[rand_with_step(-9,9,1)]);
   c2:matrix([rand_with_step(-9,9,1)],[rand_with_step(-9,9,1)],[rand_with_step(-9,9,1)]);
   c3:matrix([rand_with_step(-9,9,1)],[rand_with_step(-9,9,1)],[rand_with_step(-9,9,1)]);
   n1:rand(10)+1;
   n2:rand(10)+1;
   n3:rand(10)+1;
   n4:rand(10)+1;
   correctanswer:ev(((n3-n4)*c1-n2*c2)+n1*c3);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \left( ({@n3@}-{@n4@}){@c1@} - {@n2@} {@c2@}\right) + {@n1@} {@c3@}\)  \( = {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>correctanswer</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>517923401</deployedseed>
    <deployedseed>1290219456</deployedseed>
    <deployedseed>826799751</deployedseed>
    <deployedseed>735023879</deployedseed>
    <deployedseed>2095154866</deployedseed>
    <deployedseed>1692189631</deployedseed>
    <deployedseed>58269562</deployedseed>
    <deployedseed>15370684</deployedseed>
    <deployedseed>320606333</deployedseed>
    <deployedseed>260092687</deployedseed>
    <deployedseed>2125742612</deployedseed>
    <deployedseed>923906566</deployedseed>
    <deployedseed>970866154</deployedseed>
    <deployedseed>1360376489</deployedseed>
    <deployedseed>1629656914</deployedseed>
    <deployedseed>1067081007</deployedseed>
    <deployedseed>1303531381</deployedseed>
    <deployedseed>1095142198</deployedseed>
    <deployedseed>420101519</deployedseed>
    <deployedseed>531018846</deployedseed>
    <deployedseed>2019012806</deployedseed>
    <deployedseed>1484937780</deployedseed>
    <deployedseed>41232927</deployedseed>
    <deployedseed>765553240</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>c1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
