<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Vector_Independent_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Given are the vectors: \[ {@v1@},\quad {@v2@}. \] <br /> Find a third vector which is linearly independent from the two given vectors. [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The two given vectors \( {@v1@},{@v2@} \) and the third vector should form a matrix whose rowspace has dimension 3 (or whose rank is dimension 3). So for example, \[{@am@}\] has rank 3 since after reduction to Echelon form, we have three non-zero rows.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v1:[1,rand_with_step(-9,9,1),rand_with_step(-9,9,1)];
 v2:[0,rand_with_step(-9,9,1),1];
 v3:[0,1,0];
 am:matrix(v1,v2,v3);
 correctanswer:matrix([0,1,0]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@v1@}, {@v2@}\), linearly independent, possibly \( {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>correctanswer</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>am[3,1]:ans1[1,1];
  am[3,2]:ans1[1,2];
  am[3,3]:ans1[1,3];
  rkanswer:rank(am);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rkanswer</sans>
        <tans>3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The rank of the matrix \({@am@}\) is less then 3, hence the three vectors are linear independent.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1230467281</deployedseed>
    <deployedseed>1794441020</deployedseed>
    <deployedseed>823385897</deployedseed>
    <deployedseed>1182590914</deployedseed>
    <deployedseed>2106488680</deployedseed>
    <deployedseed>724187873</deployedseed>
    <deployedseed>1925519458</deployedseed>
    <deployedseed>595246740</deployedseed>
    <deployedseed>1125137901</deployedseed>
    <deployedseed>930345148</deployedseed>
    <deployedseed>1788961441</deployedseed>
    <deployedseed>247522967</deployedseed>
    <deployedseed>1641492822</deployedseed>
    <deployedseed>1723934299</deployedseed>
    <deployedseed>109639187</deployedseed>
    <deployedseed>1141041621</deployedseed>
    <deployedseed>552896168</deployedseed>
    <deployedseed>382384202</deployedseed>
    <deployedseed>944563508</deployedseed>
    <deployedseed>505917431</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>matrix(v1)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
