<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>14.2.2.4 Finding r.m.s. of trig functions</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Calculate the r.m.s. value of the following, to 4 decimal places.</p><br>

<ol class="HELM_parts">
    <li>
        <p>\(f(t)={@f1@}\) across \([{@a1@},{@b1@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
    </li>
    <li>
        <p>\(f(t)={@f2@}\) across \([{@a2@},{@b2@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
    </li>

    <li>
        <p>\(f(t)={@f7@}\) across \([{@a7@},{@b7@}]\)<br><span style="font-size: 0.9375rem;">r.m.s. value \(=\) [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</span></p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Remember the Key Point, r.m.s. value \(\displaystyle = \sqrt{\frac{1}{b-a} \int_a^b [f(t)]^2 \ \mathrm{d}t}\).</p><br>

<ol class="HELM_parts">
    <li>Using the Key Point with \(a={@a1@}\), \(b={@b1@}\) and \(f(t)={@f1@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b1@}-{@a1@}} \int_{@a1@}^{@b1@} \left[{@f1@}\right]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{{@1/(b1-a1)@}\left[{@integrate(f1^2,t)@}\right]_{@a1@}^{@b1@}}\\
        &amp;= \sqrt{{@1/(b1-a1)@}\left[{@subst(t=b1,integrate(f1^2,t))@}-{@subst(t=a1,integrate(f1^2,t))@}\right]}\\ &amp;\approx {@ta1@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a2@}\), \(b={@b2@}\) and \(f(t)={@f2@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b2@}-{@a2@}} \int_{@a2@}^{@b2@} \left[{@f2@}\right]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{{@1/(b2-a2)@}\left[{@integrate(f2^2,t)@}\right]_{@a2@}^{@b2@}}\\
        &amp;= \sqrt{{@1/(b2-a2)@}\left[{@subst(t=b2,integrate(f2^2,t))@}-{@subst(t=a2,integrate(f2^2,t))@}\right]}\\ &amp;\approx {@ta2@} \end{aligned}\]</li>



    <li>Using the Key Point with \(a={@a7@}\), \(b={@b7@}\) and \(f(t)={@f7@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b7@}-{@a7@}} \int_{@a7@}^{@b7@} \left[{@f7@}\right]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{{@1/(b7-a7)@}\left[{@integrate(f7^2,t)@}\right]_{@a7@}^{@b7@}}\\
        &amp;= \sqrt{{@1/(b7-a7)@}\left[{@subst(t=b7,integrate(f7^2,t))@}-\left({@subst(t=a7,integrate(f7^2,t))@}\right)\right]}\\ &amp;\approx {@ta7@} \end{aligned}\]</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>omega: 1/2;
f1: sin(t);
a1: 0;
b1: pi/2;
ta1: sqrt(1/(b1-a1)*integrate(f1^2,t,a1,b1));
ta1: decimalplaces(float(ta1),4);

f2: sin(t);
a2: 0;
b2: pi;
ta2: sqrt(1/(b2-a2)*integrate(f2^2,t,a2,b2));
ta2: decimalplaces(float(ta2),4);

f3: sin(omega*t);
a3: 0;
b3: pi;
ta3: sqrt(1/(b3-a3)*integrate(f3^2,t,a3,b3));
ta3: decimalplaces(float(ta3),4);

f4: cos(t);
a4: 0;
b4: pi/2;
ta4: sqrt(1/(b4-a4)*integrate(f4^2,t,a4,b4));
ta4: decimalplaces(float(ta4),4);

f5: cos(t);
a5: 0;
b5: pi;
ta5: sqrt(1/(b5-a5)*integrate(f5^2,t,a5,b5));
ta5: decimalplaces(float(ta5),4);

f6: cos(omega*t);
a6: 0;
b6: pi;
ta6: sqrt(1/(b6-a6)*integrate(f6^2,t,a6,b6));
ta6: decimalplaces(float(ta6),4);

f7: sin(omega*t)+ cos(omega*t);
a7: 0;
b7: 1;
ta7: sqrt(1/(b7-a7)*integrate(f7^2,t,a7,b7));
ta7: decimalplaces(float(ta7),4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>({@f1@} , [{@a1@},{@b1@}]), ({@f2@} , [{@a2@},{@b2@}]), ({@f3@} , [{@a3@},{@b3@}]), ({@f4@} , [{@a4@},{@b4@}]), ({@f5@} , [{@a5@},{@b5@}])</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans7</name>
      <type>numerical</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt7-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
