<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>14.2.2.2 Finding r.m.s of functions</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Calculate the r.m.s. value of the given functions across the specified interval, to 4 decimal places.</p><br>

<ol class="HELM_parts">
    <li>
        <p>\(f(t)={@f1@}\) across \([{@a1@},{@b1@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
    </li>
    <li>
        <p>\(f(x)={@f2@}\) across \([{@a2@},{@b2@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
    </li>
    <li>
        <p>\(f(t)={@f3@}\) across \([{@a3@},{@b3@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
    </li>
    <li>
        <p>\(f(t)={@f4@}\) across \([{@a4@},{@b4@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>
    </li>
    <li>
        <p>\(f(z)={@f5@}\) across \([{@a5@},{@b5@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Remember the Key Point, r.m.s. value \(\displaystyle = \sqrt{\frac{1}{b-a} \int_a^b [f(t)]^2 \ \mathrm{d}t}\).</p><br>

<ol class="HELM_parts">
    <li>Using the Key Point with \(a={@a1@}\), \(b={@b1@}\) and \(f(t)={@f1@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b1@}-{@a1@}} \int_{@a1@}^{@b1@} [{@f1@}]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{\frac{1}{{@b1-a1@}}\left[{@integrate(f1^2,t)@}\right]_{@a1@}^{@b1@}}\\
        &amp;= \sqrt{\frac{1}{{@b1-a1@}}\left[{@subst(t=b1,integrate(f1^2,t))@}-{@subst(t=a1,integrate(f1^2,t))@}\right]}\\ &amp;\approx {@ta1@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a2@}\), \(b={@b2@}\) and \(f(x)={@f2@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b2@}-({@a2@})} \int_{@a2@}^{@b2@} [{@f2@}]^2 \ \mathrm{d}x}\\ &amp;= \sqrt{\frac{1}{{@b2-a2@}}\left[{@integrate(f2^2,x)@}\right]_{@a2@}^{@b2@}}\\
        &amp;= \sqrt{\frac{1}{{@b2-a2@}}\left[{@subst(x=b2,integrate(f2^2,x))@}-\left({@subst(x=a2,integrate(f2^2,x))@}\right)\right]}\\ &amp;\approx {@ta2@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a3@}\), \(b={@b3@}\) and \(f(t)={@f3@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b3@}-{@a3@}} \int_{@a3@}^{@b3@} [{@f3@}]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{\frac{1}{{@b3-a3@}}\left[{@integrate(f3^2,t)@}\right]_{@a3@}^{@b3@}}\\
        &amp;= \sqrt{\frac{1}{{@b3-a3@}}\left[{@subst(t=b3,integrate(f3^2,t))@}-{@subst(t=a3,integrate(f3^2,t))@}\right]}\\ &amp;\approx {@ta3@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a4@}\), \(b={@b4@}\) and \(f(t)={@f4@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b4@}-{@a4@}} \int_{@a4@}^{@b4@} [{@f4@}]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{\frac{1}{{@b4-a4@}}\left[{@integrate(f4^2,t)@}\right]_{@a4@}^{@b4@}}\\
        &amp;= \sqrt{\frac{1}{{@b4-a4@}}\left[{@subst(t=b4,integrate(f4^2,t))@}-{@subst(t=a4,integrate(f4^2,t))@}\right]}\\ &amp;\approx {@ta4@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a5@}\), \(b={@b5@}\) and \(f(z)={@f5@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b5@}-{@a5@}} \int_{@a5@}^{@b5@} [{@f5@}]^2 \ \mathrm{d}z}\\ &amp;= \sqrt{\frac{1}{{@b5-a5@}}\left[{@integrate(f5^2,z)@}\right]_{@a5@}^{@b5@}}\\
        &amp;= \sqrt{\frac{1}{{@b5-a5@}}\left[{@subst(z=b5,integrate(f5^2,z))@}-{@subst(z=a5,integrate(f5^2,z))@}\right]}\\ &amp;\approx {@ta5@} \end{aligned}\]</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>f1: rand([1,2,3])+t;
a1: 0;
b1: 2;
ta1: sqrt(1/(b1-a1)*integrate(f1^2,t,a1,b1));
ta1: decimalplaces(float(ta1),4);

f2: 2*x+rand([1,-1]);
a2: -1;
b2: 1;
ta2: sqrt(1/(b2-a2)*integrate(f2^2,x,a2,b2));
ta2: decimalplaces(float(ta2),4);

f3: t^2;
a3: 0;
b3: rand([1,2]);
ta3: sqrt(1/(b3-a3)*integrate(f3^2,t,a3,b3));
ta3: decimalplaces(float(ta3),4);

f4: t^2;
a4: 0;
b4: b3+1;
ta4: sqrt(1/(b4-a4)*integrate(f4^2,t,a4,b4));
ta4: decimalplaces(float(ta4),4);

f5: z^2+z;
a5: rand([1,2,3]);
b5: a5+2;
ta5: sqrt(1/(b5-a5)*integrate(f5^2,z,a5,b5));
ta5: decimalplaces(float(ta5),4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ ({@f1@} , [{@a1@},{@b1@}], \, m= {@ta1@}), \quad ({@f2@} , [{@a2@},{@b2@}], \, m = {@ta2@}) \]
\[ ({@f3@} , [{@a3@},{@b3@}], \, m = {@ta3@}), \quad ({@f4@} , [{@a4@},{@b4@}], \, m = {@ta4@}), \quad ({@f5@} , [{@a5@},{@b5@}], \, m = {@ta5@}) \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>764103329</deployedseed>
    <deployedseed>139501117</deployedseed>
    <deployedseed>1386598139</deployedseed>
    <deployedseed>306429570</deployedseed>
    <deployedseed>2055448053</deployedseed>
    <deployedseed>1599354366</deployedseed>
    <deployedseed>918689949</deployedseed>
    <deployedseed>1111757641</deployedseed>
    <deployedseed>1573707059</deployedseed>
    <deployedseed>289576159</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta1,4)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta2,4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ta3,4)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispdp(ta4,4)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>dispdp(ta5,4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ev(ta1+1,simp),4)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ev(ta2+1,simp),4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ev(ta3+1,simp),4)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispdp(ev(ta4+1,simp),4)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>dispdp(ev(ta5+1,simp),4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
