<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.2.1.4 Functions as Limits of Integration Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the following integrals. [Hint: Note how the same curves can define different regions.]</p>
<ol class="HELM_parts">
<li><p>\(I_1 = \displaystyle \int_{x=0}^1\displaystyle\int_{y=3x}^{x^2+2}xy\, \mathrm{d}y \, \mathrm{d}x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>
<li><p>\(I_2 = \displaystyle \int_{x=1}^2\displaystyle\int_{y=x^2+2}^{3x}xy\, \mathrm{d}y \, \mathrm{d}x = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> For each part, it helps to sketch the curves, noting their points of intersection, and identify the region of integration. You should notice that, for the first integral, \(x^2 + 2\) is the upper curve, while in the second, it is the lower curve. </p>
<p> (a) We have \(I_1 = \displaystyle \int_{x=0}^1\displaystyle\int_{y=3x}^{x^2+2}xy\, \mathrm{d}y \, \mathrm{d}x\) </p>
<p> Considering the inner integral, we have </p>
<p> \(\qquad \displaystyle \int_{y=3x}^{x^2+2} xy \, \mathrm{d}y  = \left[x \dfrac{y^2}{2}\right]_{y=3x}^{x^2+2} = x \dfrac{(x^2+2)^2}{2} - x \dfrac{(3x)^2}{2}   = \dfrac{x^5 - 5x^3 + 4x}{2} \)</p>

<p> So our double integral becomes </p>
<p> \(\begin{align}
I_1 &= \displaystyle \int_{x=0}^1 \dfrac{x^5 - 5x^3 + 4x}{2} \, \mathrm{d}x = \dfrac{1}{2} \displaystyle\int_{x=0}^1 (x^5 - 5x^3 + 4x)\, \mathrm{d}x \\
&= \dfrac{1}{2} \left[\dfrac{x^6}{6} - \dfrac{5x^4}{4} + 2x^2 \right]_{x=0}^{1} \\
&= \dfrac{1}{2} \left( \dfrac{1^6}{6} - \dfrac{5(1)^4}{4} + 2(1)^2 \right) - 0 \\
&= \dfrac{1}{2} \times \dfrac{11}{12} = \dfrac{11}{24}. \end{align} \) </p>

<br>
<p> (b) We have \(I_2 = \displaystyle \int_{x=1}^2\displaystyle\int_{y=x^2+2}^{3x}xy\, \mathrm{d}y \, \mathrm{d}x \) </p>
<p> Considering the inner integral, we have </p>
<p> \(\qquad \displaystyle \int_{x^2+2}^{3x} xy \, \mathrm{d}y = \left[x \dfrac{y^2}{2}\right]_{x^2+2}^{3x} = x \dfrac{(3x)^2}{2} - x \dfrac{(x^2+2)^2}{2} = \dfrac{5x^3 - x^5 - 4x}{2} \) </p>

<p> So our double integral becomes </p>
<p> \(\begin{align}
I_2 &= \displaystyle \int_{x=1}^{2} \dfrac{5x^3 - x^5 - 4x}{2} \, \mathrm{d}x = \dfrac{1}{2} \displaystyle \int_{x=1}^{2} (5x^3 - x^5 - 4x) \, \mathrm{d}x \\
&=  \dfrac{1}{2} \left[\dfrac{5x^4}{4} - \dfrac{x^6}{6} - 2x^2\right]_{1}^{2} \\
&= \dfrac{1}{2} \left(\dfrac{5(2)^4}{4} - \dfrac{(2)^6}{6} - 2(2)^2 - \left(\dfrac{5(1)^4}{4} - \dfrac{(1)^6}{6} - 2(1)^2\right)\right) \\
&= \dfrac{1}{2} \times \dfrac{9}{4} \\
&= {@ta2@}.\end{align} \) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
      <text>ta1:11/24;
ta2:9/8;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\(I_1 = {@ta1@}, I_2 = {@ta2@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1, simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
