<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.2.1.5 Functions as Limits of Integration Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the following integral. Do not convert your answer to a decimal. [Hint: use integration by parts for the outer integral.] </p>
<p>\(\displaystyle I = \int_{x={@aa@}}^{@bb@}\displaystyle\int_{y=1}^{x^2}\frac{x}{y}\, \mathrm{d}y \, \mathrm{d}x\) </p>
<p>\(I = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> We have \(\displaystyle I = \int_{x={@aa@}}^{@bb@}\displaystyle\int_{y=1}^{x^2}\frac{x}{y}\, \mathrm{d}y \, \mathrm{d}x\) </p>
<p> Evaluating the inner integral, we get </p>
<p> \(\displaystyle\int_{y=1}^{x^2}\dfrac{x}{y}\, \mathrm{d}y \) </p>
<p> \(\displaystyle= \displaystyle\left[x \ln(y)\right]_{y=1}^{x^2} \) </p>
<p> \(= x\times \ln(x^2) - x\times\ln(1)\) </p>
<p> \(= 2x\ln(x)\) </p>

<p> Substituting this into the outer integral and evaluating, we get: </p>
<p>\(I = \displaystyle \int_{x={@aa@}}^{@bb@} 2x\ln(x) \, \mathrm{d}x = 2\int_{x={@aa@}}^{@bb@} x\ln(x) \, \mathrm{d}x \) </p>

<p> So computing this integral using integration by parts, we get </p>
<p> \(\begin{align} 
I &= 2 \int_{x={@aa@}}^{@bb@} x\ln(x) \, \mathrm{d}x = 2 \times \left(\left[\dfrac{x^2}{2}\ln(x)\right]_{x={@aa@}}^{@bb@} - \displaystyle\int_{x={@aa@}}^{@bb@} \dfrac{1}{x} \times \dfrac{x^2}{2} \, \mathrm{d}x \right) \\
&= 2 \times \left(\left[\dfrac{x^2}{2}\ln(x)\right]_{x={@aa@}}^{@bb@} - \dfrac{1}{2}\displaystyle\int_{x={@aa@}}^{@bb@} x \, \mathrm{d}x \right) \\
&=  2 \times \left(\left[\dfrac{x^2}{2}\ln(x)\right]_{x={@aa@}}^{@bb@} - \dfrac{1}{2} \left[\dfrac{x^2}{2}\right]_{x={@aa@}}^{@bb@} \right) \\
&= 2 \times \left({@bb^2/2*ln(bb)@} - {@aa^2/2 * ln(aa)@} - \dfrac{1}{2} \times \left({@bb^2/2@} - {@aa^2/2@}\right)\right) \\
&= {@ta1@}
\end{align}\) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand([1,2]);
bb: rand_with_prohib(2,5,[aa]);

F: x^2*log(x) - x^2/2;

raw_ans: ev(F, x=bb) - ev(F, x=aa);

ta1: fullratsimp(raw_ans);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>\(I = \displaystyle\int_{x={@aa@}}^{@bb@}\displaystyle\int_{y=1}^{x^2}\frac{x}{y}\, \mathrm{d}y \, \mathrm{d}x= {@ta1@} \)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2087659824</deployedseed>
    <deployedseed>1910465174</deployedseed>
    <deployedseed>1520949153</deployedseed>
    <deployedseed>1723402493</deployedseed>
    <deployedseed>1319394220</deployedseed>
    <deployedseed>711672310</deployedseed>
    <deployedseed>1916012798</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(factor(ta1),simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
