<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.2.1.2 Functions as Limits of Integration Task</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the following double integral over a non-rectangular region. </p>
<p> \(\displaystyle \displaystyle\int_{x=0}^{@aa@}\displaystyle\int_{y=x}^{@aa@}\left(x^2+y^2\right)\; \mathrm{d}y \, \mathrm{d}x\)</p>
<p>It would be helpful to sketch in a paper the region of the xy-plane determined by the limits of integration. </p>

<ol class="HELM_parts">
<li><p>Evaluate the inner integral: Inner Integral = [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p> </li>
<li><p>Evaluate the outer integral: Outer Integral = [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p> </li> </ol>
<p>[Note that the above task is simply one of integrating a function over a region - there is no reference to a volume here.]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>You are asked to evaluate the following double integral over a non-rectangular region. </p>
<p> \(\displaystyle \displaystyle\int_{x=0}^{@aa@}\displaystyle\int_{y=x}^{@aa@}\left(x^2+y^2\right)\; \mathrm{d}y \, \mathrm{d}x\)</p>
<p>The sketch should look something like this: </p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [-0.5, {#aa+1#}, {#aa+1#}, -0.5],
            keepaspectratio: true,
            axis: true,
            showCopyright: false,
			showNavigation: false,
			grid: false
		});
		
		// Hide tick labels and ticks on x-axis but label it
		board.defaultAxes.x.setAttribute({
		  ticks: {
			drawLabels: false, // hide numbers
			majorHeight: 0,
			minorHeight: 0
		  },
		  name: 'x',
		  withLabel: true,
		  label: {
			position: 'rt',
			offset: [10, -10]
		  }
		});


		// Hide tick labels and ticks on y-axis but label it
		board.defaultAxes.y.setAttribute({
		  ticks: {
			drawLabels: false,
			majorHeight: 0,
			minorHeight: 0
		  },
		  name: 'y',
		  withLabel: true,
		  label: {
			position: 'rt',
			offset: [-15, 0]
		  }
		});
		
		// Label at x = aa (vertical line position)
		board.create('text', [{#aa#}, -0.05, 'x = {#aa#}'], {
			anchorX: 'left',
			anchorY: 'top',
                        fixed: true,
		});

		// Label at y = 1 (horizontal line position)
		board.create('text', [-0.05, {#aa#}, 'y = {#aa#}'], {
			anchorX: 'right',
			anchorY: 'bottom',
                        fixed: true,
		});

		// Label at x = y (diagonal line position)
		board.create('text', [{#aa+0.05#}, {#aa+0.05#}, 'x = y'], {
			anchorX: 'left',
			anchorY: 'bottom',
                        fixed: true,
		});
		
		
		// line at x=1
		const p1 = [{#aa#}, -1];
		const p2 = [{#aa#}, 6];

		const verticalLine = board.create('segment', [p1, p2], {
			dash: 2,
			color: 'gray',
			fixed: true
		});
		
		// points
		var p = board.create('point', [0,0], {fixed: true, withLabel:false, visible: false});
		var q = board.create('point', [0,{#aa#}], {fixed: true, withLabel:false, visible: false});
		var r = board.create('point', [{#aa#},{#aa#}], {fixed: true, withLabel:false, visible: false});
		var triangle = board.create('polygon', [p, q, r],{name:"Triangle",withLabel:false, color: 'grey'});
[[/jsxgraph]]
<p> Where the shaded area is the area of integration. </p>

<p> (a) In the triangle, \(x\) varies between \(x = 0\) and \(x = {@aa@}\). For every value of \(x\), \(y\) varies between \(y = x\) and \(y = {@aa@}\). </p>
<p>The inner integral is given by </p>
\[\begin{align}
\text{Inner}\hspace{0.1cm} \text{Integral} = \int_{y=x}^{@aa@} \left(x^2 + y^2\right)d\, \mathrm{d}y
&= \left[ x^2 y + \dfrac{y^3}{3} \right]_x^{@aa@} \\
&= x^2 \times {@aa@} + \dfrac{{@aa@}^3}{3} - (x^2 \times x + \dfrac{x^3}{3}) \\
&= {@aa*x^2@} + {@aa^3/3@} - \dfrac{4x^3}{3}
\end{align}\]

<p> (b) The inner integral is placed in the outer integral to give </p>
\begin{align}
\text{Outer}\hspace{0.1cm} \text{Integral} = \int_0^{@aa@}
\left({@aa*x^2@} + \dfrac{{@aa^3@}}{3} - \dfrac{4x^3}{3}\right)\, \mathrm{d}x
&= \left[ \dfrac{{@aa*x^3@}}{3} + \dfrac{{@aa^3@}x}{3} - \dfrac{x^4}{3}\right]_0^{@aa@} \\
&= (\dfrac{{@aa@} \times {@aa^3@}}{3} + \dfrac{{@aa^3@} \times {@aa@}}{3} - \dfrac{{@aa@}^4}{3}) - 0 \\
&= {@aa^4/3@}
\end{align}]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(5) + 1;
exp: x^2 + y^2;
ta1: aa*x^2 + aa^3/3 - (4*x^3)/3;
ta2: int(ta1, x, 0, aa);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\[I = \displaystyle\int_{x=0}^{@aa@}\displaystyle\int_{y=x}^{@aa@}\left(x^2+y^2\right)\; \mathrm{d}y \, \mathrm{d}x\]
\[\text{Inner Integral} = {@ta1@}\]
\[ \text{Outer Integral} =  {@ta2@}\]</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>[y,NOCONST]</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1688210557</deployedseed>
    <deployedseed>643296616</deployedseed>
    <deployedseed>2090237492</deployedseed>
    <deployedseed>847008440</deployedseed>
    <deployedseed>1580175891</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1, simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>jsxgraph_solution</text>
</tag>
    </tags>
  </question>
</quiz>
