<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.5.2.6 Inequalities with var</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<h4 class="HELM_exercise">Exercise</h4>
<p>In the following questions solve the given inequality algebraically:</p>
<ol class="HELM_parts">
<li>\({@seq4[1][1]@}\)   [[input:ans17]] [[validation:ans17]][[feedback:prt17]]
</li><li>\({@seq4[2][1]@}\)   [[input:ans18]] [[validation:ans18]][[feedback:prt18]]
</li><li>\({@seq4[3][1]@}\)  [[input:ans19]] [[validation:ans19]][[feedback:prt19]]
</li><li>\({@seq4[4][1]@}\)    [[input:ans20]] [[validation:ans20]][[feedback:prt20]]
</li><li>\({@seq4[5][1]@}\)  [[input:ans21]] [[validation:ans21]][[feedback:prt21]]
</li><li>\({@seq4[6][1]@}\)    [[input:ans22]] [[validation:ans22]][[feedback:prt22]]
</li><li>\({@seq4[7][1]@}\)   [[input:ans23]] [[validation:ans23]][[feedback:prt23]]
</li><li>\({@seq4[8][1]@}\)   [[input:ans24]] [[validation:ans24]][[feedback:prt24]]
</li><li>\({@seq4[9][1]@}\)   [[input:ans25]] [[validation:ans25]][[feedback:prt25]]
</li></ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[var:x;
rnd3:rand_with_step(1,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp17:rnd3*var + rndd3>rnddd3;
ta17:var>(rnddd3-rndd3)/rnd3;
rnd3:rand_with_step(1,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp18:rnd3*var + rndd3>rnddd3;
ta18:var>(rnddd3-rndd3)/rnd3;
rnd3:rand_with_step(1,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp19:rnd3*var + rndd3>=rnddd3;
ta19:var>=(rnddd3-rndd3)/rnd3;
rnd3:rand_with_step(1,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp20:rnd3*var + rndd3<rnddd3;
ta20:var<(rnddd3-rndd3)/rnd3;
rnd3:rand_with_step(1,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp21:rnd3*var + rndd3<rnddd3;
ta21:var<(rnddd3-rndd3)/rnd3;
rnd3:rand_with_step(1,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp22:rnd3*var + rndd3<=rnddd3;
ta22:var<=(rnddd3-rndd3)/rnd3;

rnd3:rand_with_step(-5,-1,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp23:rnd3*var + rndd3<=rnddd3;
ta23:var>=(rnddd3-rndd3)/rnd3;
rnd3:rand_with_step(-5,-1,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
exp24:rnd3*var + rndd3>rnddd3;
ta24:var<(rnddd3-rndd3)/rnd3;

rnd3:rand_with_step(5,9,1);
rndd3:rand_with_step(-5,10,1);
rnddd3:rand_with_step(-5,30,1);
rand4:rand_with_step(1,4,1);
exp25:rnd3*var + rndd3<=rand4*var+rnddd3;
ta25:var<=(rnddd3-rndd3)/(rnd3-rand4);

seq3:[[exp17,ta17],[exp18,ta18],[exp19,ta19],[exp20,ta20],[exp21,ta21],[exp22,ta22],[exp23,ta23],[exp24,ta24],[exp25,ta25]];
seq4:rand_selection(seq3,9);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@seq4[1][1]@}\) ans: \({@seq4[1][2]@}\); \({@seq4[2][1]@}\) ans: \({@seq4[2][2]@}\); \({@seq4[3][1]@}\) ans: \({@seq4[3][2]@}\); \({@seq4[4][1]@}\) ans: \({@seq4[4][2]@}\); \({@seq4[5][1]@}\) ans: \({@seq4[5][2]@}\);
 \({@seq4[6][1]@}\) ans: \({@seq4[6][2]@}\); \({@seq4[7][1]@}\) ans: \({@seq4[7][2]@}\); \({@seq4[8][1]@}\) ans: \({@seq4[8][2]@}\); \({@seq4[9][1]@}\) ans: \({@seq4[9][2]@}\);</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans17</name>
      <type>algebraic</type>
      <tans>seq4[1][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans18</name>
      <type>algebraic</type>
      <tans>seq4[2][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans19</name>
      <type>algebraic</type>
      <tans>seq4[3][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans20</name>
      <type>algebraic</type>
      <tans>seq4[4][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans21</name>
      <type>algebraic</type>
      <tans>seq4[5][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans22</name>
      <type>algebraic</type>
      <tans>seq4[6][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans23</name>
      <type>algebraic</type>
      <tans>seq4[7][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans24</name>
      <type>algebraic</type>
      <tans>seq4[8][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans25</name>
      <type>algebraic</type>
      <tans>seq4[9][2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt17</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans17</sans>
        <tans>seq4[1][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt17-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt17-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[1][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt18</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans18</sans>
        <tans>seq4[2][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt18-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt18-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[2][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt19</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans19</sans>
        <tans>seq4[3][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt19-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt19-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[3][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt20</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans20</sans>
        <tans>seq4[4][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt20-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt20-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[4][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt21</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans21</sans>
        <tans>seq4[5][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt21-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt21-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[5][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt22</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans22</sans>
        <tans>seq4[6][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt22-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt22-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[6][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt23</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans23</sans>
        <tans>seq4[7][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt23-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt23-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[7][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt24</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans24</sans>
        <tans>seq4[8][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt24-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt24-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[8][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt25</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans25</sans>
        <tans>seq4[9][2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt25-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt25-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The correct answer is \({@seq4[9][2]@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1811779389</deployedseed>
    <deployedseed>948441209</deployedseed>
    <deployedseed>1215328418</deployedseed>
    <deployedseed>1775569777</deployedseed>
    <deployedseed>963795157</deployedseed>
    <deployedseed>449932046</deployedseed>
    <deployedseed>1297095282</deployedseed>
    <deployedseed>2033705189</deployedseed>
    <deployedseed>1401593298</deployedseed>
    <deployedseed>941248890</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans17</name>
        <value>seq4[1][2]</value>
      </testinput>
      <testinput>
        <name>ans18</name>
        <value>seq4[2][2]</value>
      </testinput>
      <testinput>
        <name>ans19</name>
        <value>seq4[3][2]</value>
      </testinput>
      <testinput>
        <name>ans20</name>
        <value>seq4[4][2]</value>
      </testinput>
      <testinput>
        <name>ans21</name>
        <value>seq4[5][2]</value>
      </testinput>
      <testinput>
        <name>ans22</name>
        <value>seq4[6][2]</value>
      </testinput>
      <testinput>
        <name>ans23</name>
        <value>seq4[7][2]</value>
      </testinput>
      <testinput>
        <name>ans24</name>
        <value>seq4[8][2]</value>
      </testinput>
      <testinput>
        <name>ans25</name>
        <value>seq4[9][2]</value>
      </testinput>
      <expected>
        <name>prt17</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt17-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt18</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt18-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt19</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt19-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt20</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt20-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt21</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt21-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt22</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt22-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt23</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt23-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt24</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt24-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt25</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt25-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
