<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.6.2.3_Find the partial fractions (NOT IN QUIZ)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise<br></p>
<p>Find the partial fractions of:</p>
<ol class="HELM_parts">
<li><p></p><h4>\({@WholeExptFact1@}= \)[[input:ans1]] </h4><p></p></li>[[validation:ans1]][[feedback:prt1]]
<br>
<li><p></p><h4>\({@wholeExp2@}= \)[[input:ans2]]</h4><p></p></li> [[validation:ans2]][[feedback:prt2]]
<br>
<li><p></p><h4>\({@WholeExptFact3@}= \)[[input:ans3]]</h4><p></p></li> [[validation:ans3]][[feedback:prt3]]

</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>var:x;
var2:A;
var3:B;
rand1:rand_with_step(1,2,1);
rand2:rand_with_step(-5,-1,1);
rand3:rand_with_step(1,3,1);
rand4:rand_with_step(0,2,1);
exp1:rand1*var+(rand1*rand2+1);
exp2:rand3*var+(rand3*rand4-1);
expDenom1:expand(exp1*exp2);
ta11:(exp1)*(exp2);
expA:rand_with_prohib(-9,9,[0]);
expB:rand_with_prohib(-9,9,[0]);
expNum1:expand(expA*exp2+expB*exp1);
wholeExp1: expNum1/expDenom1;
WholeExptFact1: expNum1/ta11;
ta1:expA/exp1+expB/exp2;

rand1:rand_with_step(1,3,1);
rand2:rand_with_step(-3,-1,1);
rand3:rand_with_step(1,3,1);
rand4:rand_with_step(0,2,1);
exp1:rand1*var+(rand1*rand2+1);
exp2:rand3*var+(rand3*rand4-1);
expDenom1:expand(exp1*exp2);
ta11:(exp1)*(exp2);
expA:rand_with_prohib(-9,9,[0]);
expB:rand_with_prohib(-9,9,[0]);
expNum1:expand(expA*exp2+expB*exp1);
wholeExp2: expNum1/expDenom1;
WholeExptFact2: expNum1/ta11;
ta2:expA/exp1+expB/exp2;

rand1:rand_with_step(1,8,1);
rand2:rand_with_step(-5,-1,1);
rand3:rand_with_step(1,3,1);
rand4:rand_with_step(1,5,1);
exp1:rand1*var+(rand1*rand2+1);
exp2:rand3*var+(rand3*rand4+1);
expDenom1:expand(exp1*exp2);
ta11:(exp1)*(exp2);
expA:rand_with_prohib(-9,9,[0]);
expB:rand_with_prohib(-9,9,[0]);
expNum1:expand(expA*exp2+expB*exp1);
wholeExp3: expNum1/expDenom1;
WholeExptFact3: expNum1/ta11;
ta3:expA/exp1+expB/exp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@WholeExptFact1@}\), \({@wholeExp2@}\), \({@WholeExptFact3@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your answer is incorrect. The correct partial fraction is \({@ta1@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your answer is incorrect. The correct partial fraction is \({@ta2@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your answer is incorrect. The correct partial fraction is \({@ta3@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>268363270</deployedseed>
    <deployedseed>1836852489</deployedseed>
    <deployedseed>291840650</deployedseed>
    <deployedseed>913944909</deployedseed>
    <deployedseed>2060580556</deployedseed>
    <deployedseed>25571352</deployedseed>
    <deployedseed>1120158455</deployedseed>
    <deployedseed>1328382954</deployedseed>
    <deployedseed>1028097588</deployedseed>
    <deployedseed>1589969008</deployedseed>
    <deployedseed>546859691</deployedseed>
    <deployedseed>1753069844</deployedseed>
    <deployedseed>271081459</deployedseed>
    <deployedseed>325235837</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
