<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.6.4.3 Express each fraction as sum of partial fractions (NOT IN QUIZ)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise<br></p>
<p>Express each of the following as the sum of its partial fractions.</p>
<ol class="HELM_parts">
<li>\(\dfrac{{@expNum1@}}{{@expDenom1@}}=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</li>
<br>
<li>\(\dfrac{{@expNum2@}}{{@expDenom2@}}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</li>

</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">
    <li>
        <p> First, by multiplying both sides by \(({@exp11@})({@exp12@})\) we obtain \({@ta11@}\). By letting \({@var@}={@ta1X@}\) we find that the correct value of \({@var4@}\) is \({@ta12@}\). By equating the coefficients of \({@var@}^2\) we find \({@var2@}={@ta13@}\).
            If we substitute any other value for \({@var@}\) (or equate coefficients of \({@var@}\)) we find \({@var3@}={@ta14@}\). Finally, the partial fractions are \(\displaystyle {@ta15@}\).</p>
    </li>
    <li>
        <p>First, by multiplying both sides by \(({@exp21@})({@exp22@})\) we obtain \({@ta21@}\). By letting \({@var@}={@ta2X@}\) we find that the correct value of \({@var4@}\) is \({@ta22@}\). By equating the coefficients of \({@var@}^2\) we find \({@var2@}={@ta23@}\).
            If we substitute any other value for \({@var@}\) (or equate coefficients of \({@var@}\)) we find \({@var3@}={@ta24@}\). Finally, the partial fractions are \(\displaystyle {@ta25@}\).</p>
    </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>var:x;
var2:A;
var3:B;
var4:C;
expA:rand_with_prohib(-4,4,[0]);
expB:rand_with_prohib(-4,4,[0]);
expC:rand_with_prohib(-4,4,[0]);
rand1:rand_with_step(1,3,1);
rand2:rand_with_step(-2,-1,1);
rand3:rand_with_step(-2,-1,1);
exp11:rand1*var^2+(rand1*rand2+1)*var+(rand1*rand3-1);
rand4:rand_with_step(1,4,1);
rand5:rand_with_prohib(-2,2,[0]);
exp12:rand4*var-(rand3*rand4-1);
expDenom1:(exp11)*(exp12);

ta11:(var2*var+var3)*exp12+var4*exp11;
expNum1:expand((expA*var+expB)*exp12+expC*exp11);



ta1X:rand3;

ta12:expC;
ta13:expA;
ta14:expB;

ta15:(expA*var+expB)/exp11+expC/exp12;

expA:rand_with_prohib(-4,4,[0]);
expB:rand_with_prohib(-4,4,[0]);
expC:rand_with_prohib(-4,4,[0]);
rand1:rand_with_step(1,3,1);
rand2:rand_with_step(-2,-1,1);
rand3:rand_with_step(-2,-1,1);
exp21:rand1*var^2+(rand1*rand2+1)*var+(rand1*rand3-1);
rand4:rand_with_step(1,4,1);
rand5:rand_with_prohib(-2,2,[0]);
exp22:rand4*var-(rand3*rand4-1);
expDenom2:(exp21)*(exp22);

ta21:(var2*var+var3)*exp22+var4*exp21;
expNum2:expand((expA*var+expB)*exp22+expC*exp21);



ta2X:rand3;

ta22:expC;
ta23:expA;
ta24:expB;

ta25:(expA*var+expB)/exp21+expC/exp22;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\dfrac{{@expNum1@}}{{@expDenom1@}}\] 
\[ \dfrac{{@expNum2@}}{{@expDenom2@}}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta15</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta25</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta15</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The sum of the partial fractions is \({@ta1@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta25</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>The sum of the partial fractions is \({@ta2@}\).</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>150902875</deployedseed>
    <deployedseed>700809957</deployedseed>
    <deployedseed>446701185</deployedseed>
    <deployedseed>1117844166</deployedseed>
    <deployedseed>1217393901</deployedseed>
    <deployedseed>12236181</deployedseed>
    <deployedseed>72173548</deployedseed>
    <deployedseed>2070522697</deployedseed>
    <deployedseed>1370934146</deployedseed>
    <deployedseed>1863205092</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta15</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta25</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta15+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta25+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
