<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.6.3.3 Exercises_Express in partial fractions (short)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise<br></p>
<p>Express the following in partial fractions.</p>
<p>\(\displaystyle{{@WholeExpFrac3@}=}\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the expression \[\displaystyle{{@WholeExpFrac3@}}.\] The repeated linear factor is \({@ta31@}\). Therefore the partial fraction decomposition is \[ {@WholeExpFrac3@} = \frac{A}{{@exp11@}} + \frac{B}{({@exp11@})^2} = \frac{A\cdot ({@exp11@})
    + B}{ ({@exp11@})^2}. \] </p>

<p>Equating the top line yields \[ {@expNum11@} = A\cdot ({@exp11@}) + B.\] Then by substituting values of \(x\) we can solve for \(A, \,B\) and obtain \( A = {@expA@}\) and \(B = {@expB@}.\)</p><p>Therefore \[{@WholeExpFrac3@} = {@ta32@}.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>var:x;
var1:s;
var2:A;
var3:B;

rand1:rand_with_step(1,4,1);
rand2:rand_with_step(-4,-1,1);
exp1:rand1*var+(rand1*rand2-1);
expDenom1:expand(exp1^2);
ta31:exp1^2;
expA:rand_with_prohib(-4,4,[0]);
expB:rand_with_prohib(-4,4,[0]);
expNum1:expand(expA*exp1+expB);
wholeExp3: expNum1/expDenom1;
WholeExpFrac3: expNum1/exp1^2;
ta32:expA/exp1+expB/exp1^2;

exp11: exp1;
expNum11: expNum1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[{@WholeExpFrac3@} = {@ta32@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta32</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta32</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>628667984</deployedseed>
    <deployedseed>709613715</deployedseed>
    <deployedseed>1592171129</deployedseed>
    <deployedseed>1378982455</deployedseed>
    <deployedseed>156250520</deployedseed>
    <deployedseed>2076675791</deployedseed>
    <deployedseed>1850372578</deployedseed>
    <deployedseed>882804877</deployedseed>
    <deployedseed>866864305</deployedseed>
    <deployedseed>456874969</deployedseed>
    <deployedseed>1413417506</deployedseed>
    <deployedseed>652230326</deployedseed>
    <deployedseed>1565786609</deployedseed>
    <deployedseed>261171551</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans3</name>
        <value>ta32</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans3</name>
        <value>ev(ta32+1,simp)</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
