<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>23.5.2.4 ex3b sine series</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
Obtain the half-range Fourier Sine series for the functions
\[f(t) = e^{2t} \hspace{1cm} 0 \leq t \leq 1\]
<p></p>
Series is
<p>\( \displaystyle{ \sum_{n=1}^\infty} \) [[input:ans1]]</p>
[[validation:ans1]][[feedback:prt1]]&nbsp;]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[We extend as an odd function, so 
\[ f(t) = - e^{-2t} \hspace{1cm} -1 \leq t \leq 0 \]
Therefore we define 
\[ F(t) = 
\begin{cases} 
- e^{-2t} \hspace{1cm} -1 \leq t \leq 0 \\
 e^{2t} \hspace{1cm} 0 \leq t \leq 1
\end{cases}
\]
As \( F(t)\) is an odd function, \(a_n = 0\). Then we have
\[
\begin{align}
b_n &amp;= \int_{-1}^1 F(t) \sin n\pi t \hspace{0.1cm} dt = 2\int_{0}^1 e^{2t} \sin n\pi t \hspace{0.1cm} dt  \\
&amp;= 2\left\{ \left[ - \dfrac{e^{2t}}{n\pi} \cos n\pi t \right]_0^1 + \dfrac{2}{n\pi} \int_0^1 e^{2t} \cos n\pi t \hspace{0.1cm} dt \right\} \\
&amp;= 2\left\{ - \dfrac{e^2}{n\pi}\cos n\pi + \dfrac{1}{n\pi} + \dfrac{2}{n\pi} \left( \left[ \dfrac{e^{2t}}{n\pi} \sin n\pi t \right]_0^1 - \dfrac{2}{n\pi} \int_0^1 e^{2t} \sin n\pi t \hspace{0.1cm} dt  \right) \right\} \\
&amp;= 2\left\{\dfrac{1}{n\pi}  - \dfrac{e^{2}}{n\pi} \cos n\pi  - \dfrac{4}{n^2\pi^2} \int_0^1 e^{2t} \sin n\pi t \hspace{0.1cm} dt \right\} \\
&amp;= 2 \left( \dfrac{1}{n\pi} - \dfrac{e^2}{n\pi} \cos n\pi \right) - \dfrac{4}{n^2\pi^2}b_n
\end{align}
\]
Rearranging now for \(b_n\) gives
\[b_n = \dfrac{2n\pi}{4 + n^2\pi^2} \left( 1 - e^2 \cos n\pi \right) \sin n\pi t \]
<p></p>
Finally the half range Fourier sine series is given by
\[ F(t) = \sum_{n=1}^\infty \dfrac{2n\pi}{4 + n^2\pi^2} \left( 1 - e^2 \cos n\pi \right) \sin n\pi t\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021120900</text>
    </stackversion>
    <questionvariables>
      <text>ta1: (2*n*pi)/(4+n^2*pi^2)*(1 - e^2*cos(n*pi))*sin(n*pi*t);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>50</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
