<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>5.0.19 (Find cos of angle - pythagoras)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the exact value of \(\cos \theta\) if \(\sin \theta ={@frac@}\) and \(0 < \theta < 90^\circ\).</p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>In this case we have \(\cos^2 \theta + {@frac@}^2 = 1\). Hence
\[
\cos^2 \theta = {@frac2@} \implies \cos \theta = {@ta1@}.
\]
Since we require that \(0 < \theta < 90^\circ\), we take the positive value.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1+ rand(10);
bb: 10+ rand(15);
frac: aa/bb;
frac2: 1-frac^2;
ta1: sqrt(frac2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\cos \theta ={@frac@}\), \(\sin \theta = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1546684171</deployedseed>
    <deployedseed>683834402</deployedseed>
    <deployedseed>1699305333</deployedseed>
    <deployedseed>887174928</deployedseed>
    <deployedseed>1370186040</deployedseed>
    <deployedseed>1434946482</deployedseed>
    <deployedseed>564201291</deployedseed>
    <deployedseed>1793907181</deployedseed>
    <deployedseed>1474369587</deployedseed>
    <deployedseed>1371668701</deployedseed>
    <deployedseed>784709168</deployedseed>
    <deployedseed>66188666</deployedseed>
    <deployedseed>1553165883</deployedseed>
    <deployedseed>868775365</deployedseed>
    <deployedseed>1306394381</deployedseed>
    <deployedseed>655342384</deployedseed>
    <deployedseed>1946113741</deployedseed>
    <deployedseed>824339958</deployedseed>
    <deployedseed>1472358535</deployedseed>
    <deployedseed>3341683</deployedseed>
    <deployedseed>195340027</deployedseed>
    <deployedseed>473314466</deployedseed>
    <deployedseed>653485410</deployedseed>
    <deployedseed>1751976309</deployedseed>
    <deployedseed>417085840</deployedseed>
    <deployedseed>2069415204</deployedseed>
    <deployedseed>601914126</deployedseed>
    <deployedseed>234218454</deployedseed>
    <deployedseed>541415129</deployedseed>
    <deployedseed>146355327</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
