<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>5.1.20 (Cosine rule find angle)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Suppose that we have a triangle \(ABC\) and suppose that length \(BC = {@BC@}\)cm, length \(AC = {@AC@}\) and \(AB = {@AB@}\), find the size of angle \(C\). Give your answer to the nearest degree.</p>
<ul style="list-style: none;">
<li>AB = [[input:ans1]] [[validation:ans1]] degrees</li>
</ul>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Here we use the cosine rule. Recall that 
\[ 
a^2 = b^2 + c^2 - 2bc \cos A \qquad \mbox{or} \qquad b^2 = a^2 + c^2 - 2ac \cos B \qquad \mbox{or} \qquad c^2 = a^2 + b^2 - 2ab \cos C.
 \]
In order to find the angle \(C\) we take the final equation and rearrange it to get 
\[
\cos C = \frac{a^2 + b^2 - c^2}{2 a b}.
\]
Performing inverse cosine on each side then gives the size of angle \(C\). 
So in this case we have that length \(BC =a = {@BC@}\), length \(AC = b = {@AC@}\) and length \(AB = c = {@AB@}\). Thus inputting into the cosine rule we have 
\[
\cos C = \frac{{@BC@}^2 + {@AC@}^2 - {@AB@}^2}{2 \cdot{@BC@}\cdot{@AC@}} = {@coscc@},
\]
Hence \(C = {@ta1@}\) to the nearest degree.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025042500</text>
    </stackversion>
    <questionvariables>
      <text>cosd(d):=cos(d*%pi/180);
acosd(x):=acos(x)*180/%pi;

AC: 10 + rand(10);
BC: 5 + rand(5);
AB: 7+ rand(7);

aa: 1;
bb: 1;
coscc: (BC^2 + AC^2 - AB^2)/(2*BC*AC)
cc: acosd(coscc); 

ta1: round(cc);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\cos C = \frac{{@BC@}^2 + {@AC@}^2 - {@AB@}^2}{2\cdot {@BC@}\cdot {@AC@}} = {@coscc@} \implies C= {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>974990506</deployedseed>
    <deployedseed>1176765704</deployedseed>
    <deployedseed>516404731</deployedseed>
    <deployedseed>2049040616</deployedseed>
    <deployedseed>115969035</deployedseed>
    <deployedseed>276034134</deployedseed>
    <deployedseed>550779728</deployedseed>
    <deployedseed>1524168441</deployedseed>
    <deployedseed>840013376</deployedseed>
    <deployedseed>380531752</deployedseed>
    <deployedseed>527643122</deployedseed>
    <deployedseed>1156084562</deployedseed>
    <deployedseed>1048748723</deployedseed>
    <deployedseed>1248440342</deployedseed>
    <deployedseed>1313061765</deployedseed>
    <deployedseed>1471966511</deployedseed>
    <deployedseed>1851512524</deployedseed>
    <deployedseed>1214669018</deployedseed>
    <deployedseed>804797490</deployedseed>
    <deployedseed>1508652251</deployedseed>
    <deployedseed>1867073767</deployedseed>
    <deployedseed>97562304</deployedseed>
    <deployedseed>1395956557</deployedseed>
    <deployedseed>2023539619</deployedseed>
    <deployedseed>1320273304</deployedseed>
    <deployedseed>1968959315</deployedseed>
    <deployedseed>931869647</deployedseed>
    <deployedseed>1970669676</deployedseed>
    <deployedseed>1451519876</deployedseed>
    <deployedseed>1565481742</deployedseed>
    <deployedseed>1989901664</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
