<?xml version='1.0' encoding='utf-8'?>
<quiz><question type="stack">
    <name>
      <text>text_4_complex-De Moivre's</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Given a complex number \(\displaystyle z={@q@}\) &lt;span class="nolink"&gt;determine&lt;br /&gt;&lt;/span&gt;&lt;/p&gt;
&lt;p&gt;\(|z^{@n@}|=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]&lt;/p&gt;
&lt;p&gt;and \(\arg(z^{@n@})=\)[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;p&gt;It makes sense that the index laws should still apply.  This is called De Moivre's theorem.&lt;/p&gt;
&lt;p&gt;\[ {@q^n@} ={@a^n@} e^{@b*n*%i*%pi@}.\]&lt;/p&gt;
&lt;p&gt;Recall that&lt;/p&gt;
&lt;p&gt;\[ e^{i\theta} = \cos(\theta)+i\sin(\theta).\]&lt;/p&gt;
&lt;p&gt;Working with the principle argument \(0\leq \theta \leq 2\pi\) gives us&lt;/p&gt;
&lt;p&gt;\[ {@q^n@} = {@a^n@} e^{@b*n*%i*%pi@} = {@a^n@} e^{@ev(b*n,simp)*%i*%pi@} = {@a^n@} e^{@p*%i*%pi@}.\]&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <stackversion>
      <text />
    </stackversion>
    <questionvariables>
      <text>a : ev(2+rand(15),simp);
b : ev((-1)^rand(2)*((1+rand(10)))/(2+rand(15)),simp);
n : ev(3+rand(20),simp);
q : a*%e^(b*%i*%pi);
p : ev(mod(b*n,2),simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text />
    </specificfeedback>
    <questionnote>
      <text>{@q^n = a^n*(cos(p*%i*%pi)+%i*sin(p*%i*%pi))@}</text>
    </questionnote>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;Your answer is partially correct.&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>a^n</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>p*%pi</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>a^n</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>[cos(ans2),sin(ans2)]</sans>
        <tans>[cos(b*%pi*n),sin(b*%pi*n)]</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>a^n</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value />
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore />
        <expectedpenalty />
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value />
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>b*n*%pi</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore />
        <expectedpenalty />
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value />
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>p*%pi</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore />
        <expectedpenalty />
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>


  </quiz>