<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>11.7.2c first imp derivative</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Compute the derivative, \(\dfrac{dy}{dx}\) given \({@f@}={@g@}\)&nbsp;</p><p>
\( \dfrac{dy}{dx} = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the expression \({@f@}={@g@}.\)<br></p>
<p>We can deal with each term separately:</p>
<p>\(\dfrac{d}{dx} {@n*x^a*y^b@}= {@n*a@}y^{@b@}x^{{@a-1@}}+ {@n*b@}x^{@a@}y^{@b-1@}\dfrac{dy}{dx}\) (using product rule).</p>
<p>\(\dfrac{d}{dx}x^{@c@} = {@c@}x^{@c-1@}.\)<br></p>
<p>\(\dfrac{d}{dx}{@y^d@} = {@d@}y^{{@d-1@}}\dfrac{dy}{dx}.\)</p>
<p>\(\dfrac{d}{dx}{@b@} = 0.\)</p>
<p>This leaves the equation: \[{@n*a@}y^{@b@}x^{@a-1@}+ {@n*b@}x^{@a@}y^{@b-1@}\dfrac{dy}{dx} + {@c@}x^{@c-1@} = {@d@}y^{{@d-1@}}\dfrac{dy}{dx}.\]</p>
<p>By rearranging, and factoring out the term \(\dfrac{dy}{dx}\), we find that:</p>
<p>\[\dfrac{dy}{dx} = \displaystyle{{@ta1@}}= \displaystyle {@TA1@}.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>a : rand(3)+2
b : rand(5)+3
n : rand(3)+2
d : rand(5)+3
c : rand(5)+3

f : n*x^a*y^b + x^c 
g : y^d +b
ta1 : (n*a*y^b*x^(a-1)+c*x^(c-1))/(d*y^(d-1)-n*b*x^a*y^(b-1));
TA1: factor(ta1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@f@}={@g@}, \quad \dfrac{dy}{dx} ={@ta1@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><span class="correct">Correct answer, well done.</span></p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span class="partially">Your answer is partially correct.</span></p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><span class="incorrect">Incorrect answer.</span></p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>948960700</deployedseed>
    <deployedseed>1564411834</deployedseed>
    <deployedseed>1991426982</deployedseed>
    <deployedseed>545816618</deployedseed>
    <deployedseed>791697730</deployedseed>
    <deployedseed>248176090</deployedseed>
    <deployedseed>883314161</deployedseed>
    <deployedseed>493958584</deployedseed>
    <deployedseed>735421716</deployedseed>
    <deployedseed>1013045468</deployedseed>
    <deployedseed>1239919950</deployedseed>
    <deployedseed>1549401743</deployedseed>
    <deployedseed>1203408383</deployedseed>
    <deployedseed>1928980751</deployedseed>
    <deployedseed>194652863</deployedseed>
    <deployedseed>1994674515</deployedseed>
    <deployedseed>1594444087</deployedseed>
    <deployedseed>1090729792</deployedseed>
    <deployedseed>12762129</deployedseed>
    <deployedseed>450230672</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
