<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 36  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.a</text>
    </name>
    <questiontext format="html">
      <text>Differentiate \[@p@\] with respect to @v@.

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left[ @p@\right]  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left[@pp1@\right] +\frac{d}{d@v@}\left(@pp2@\right)  \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \frac{d}{d@v@}[@ppp1@]+\frac{d}{d@v@}\left(@ppp2@\right)  \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot [@diff(ppp1,v)@]+@diff(ppp2,v)@ \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:rand(4)+2;
rdm2:2+rand(5);
ppp1:v;
ppp2:rdm2;
pp1:rdm1*ppp1;
pp2:ppp2;
p:pp1+pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q2:ppp1+ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>You appear to have forgotten the coefficients @rdm1@ and @rdm2@.

Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 37  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.b</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Differentiate \[@p@\] with respect to @v@. [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In order to solve this problem we will have to use the linearity rules for differentition. calc_diff_linearity_rule \[\begin{array}{rcl}\displaystyle \frac{d}{d@v@}\left[ @p@\right] &amp;\displaystyle =&amp;\displaystyle \frac{d}{d@v@}\left[@pp1@\right] +\frac{d}{d@v@}\left( @-pp2@\right) \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\cdot \frac{d}{d@v@}[@ppp1@]+@-rdm2@\cdot \frac{d}{d@v@}\left( @ppp2@\right) \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\cdot [@diff(ppp1,v)@]+@-rdm2@\cdot \left[ @diff(ppp2,v)@\right] \\ &amp;\displaystyle =&amp;\displaystyle @diff(p,v)@ \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:rand(4)+2;
rdm2:2+rand(3);
ppp1:v;
ppp2:v^2;
pp1:rdm1*ppp1;
pp2:rdm2*ppp2;
p:pp1-pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<p><span class="correct">Correct answer, well done.</span></p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span class="partially">Your answer is partially correct.</span></p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><span class="incorrect">Incorrect answer.</span></p>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q1:pp1+pp2;
q2:ppp1-ppp2;
q3:ppp1+ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q1,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have lost a minus sign. Remember: \[\frac{d}{dx}(f-g)=\frac{df}{dx}-\frac{dg}{dx}\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have forgotten the coefficients @rdm1@ and @rdm2@. Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q3,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have forgotten the coefficients @rdm1@ and @rdm2@ and combined the functions together using an incorrect sign. Remember: \[\frac{d}{dx}(a\cdot f-b\cdot g)=a\frac{df}{dx}-b\frac{dg}{dx}\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(pp1+pp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1-ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 38  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.d</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Differentiate \[@p@\] with respect to @v@. [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In order to solve this problem we will have to use the linearity rules for differentition. calc_diff_linearity_rule \[\begin{array}{rcl}\displaystyle \frac{d}{d@v@}\left[ @p@\right] &amp;\displaystyle =&amp;\displaystyle \frac{d}{d@v@}\left[@pp1@\right] +\frac{d}{d@v@}\left( @pp2@\right) \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\cdot \frac{d}{d@v@}[@ppp1@]+@-rdm2@\cdot \frac{d}{d@v@}\left( @ppp2@\right) \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\cdot [@diff(ppp1,v)@]+@-rdm2@\cdot \left[ @diff(ppp2,v)@\right] \\ &amp;\displaystyle =&amp;\displaystyle @diff(p,v)@ \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:rand(4)+2;
rdm2:-5+rand(3);
ppp1:sin(rdm1*v);
ppp2:v^(rdm2);
pp1:rdm1*ppp1;
pp2:rdm2*ppp2;
p:pp1-pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<p><span class="correct">Correct answer, well done.</span></p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span class="partially">Your answer is partially correct.</span></p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><span class="incorrect">Incorrect answer.</span></p>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q1:pp1+pp2;
q2:ppp1+ppp2;
q3:ppp1-ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q1,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have lost a minus sign. Remember: \[\frac{d}{dx}(f-g)=\frac{df}{dx}-\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have forgotten the coefficients @rdm1@ and @rdm2@. Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q3,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have forgotten the coefficients @rdm1@ and @rdm2@ and combined the functions together using an incorrect sign. Remember: \[\frac{d}{dx}(a\cdot f-b\cdot g)=a\frac{df}{dx}-b\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(pp1+pp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1-ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 39  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.e</text>
    </name>
    <questiontext format="html">
      <text>Differentiate \[@p@\] with respect to @v@.

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left[ @p@\right]  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left[@pp1@\right] +\frac{d}{d@v@}\left(@pp2@\right)  \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \frac{d}{d@v@}[@ppp1@]+\frac{d}{d@v@}\left(@ppp2@\right)  \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot [@diff(ppp1,v)@]+@diff(ppp2,v)@ \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:rand(4)+2;
rdm2:2+rand(5);
ppp1:exp(v);
ppp2:exp(-rdm2*v);
pp1:rdm1*ppp1;
pp2:ppp2;
p:pp1+pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q1:pp1-pp2;
q2:ppp1-ppp2;
q3:ppp1+ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>You appear to have forgotten the coefficients @rdm1@ and @rdm2@.

Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q1,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text>The derivative of $e^{-x}$ is $-e^{-x}$, not $e^{-x}$.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q3,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text>You appear to have forgotten the coefficients @rdm1@ and @rdm2@, and used an incorrect sign.

Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(pp1-pp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1-ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 40  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.f</text>
    </name>
    <questiontext format="html">
      <text>Differentiate \[@p@\] with respect to @v@.

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left[ @p@\right]  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left[@pp1@\right] +\frac{d}{d@v@}\left[@pp2@\right]+\frac{d}{d@v@}\left[@pp3@\right] \\
   &\displaystyle =&\displaystyle  \frac{d}{d@v@}[@ppp1@]+\frac{d}{d@v@}\left[@ppp2@\right] +@rdm3@\cdot \frac{d}{d@v@}\left[@ppp3@\right] \\
   &\displaystyle =&\displaystyle  @diff(ppp1,v)@+[@diff(ppp2,v)@]+@rdm3@\cdot \left[@diff(ppp3,v)@\right] \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm2:-(2+rand(5));
rdm3:-2-rand(4);
ppp1:1/v;
ppp2:rdm2;
ppp3:log(v);
pp1:ppp1;
pp2:ppp2;
pp3:rdm3*ppp3;
p:pp1+pp2+pp3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 41  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.g</text>
    </name>
    <questiontext format="html">
      <text>Differentiate \[@p@\] with respect to @v@.

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left[ @p@\right]  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left[@pp1@\right] +\frac{d}{d@v@}\left[@pp2@\right]+\frac{d}{d@v@}\left[@pp3@\right] \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \frac{d}{d@v@}[@ppp1@]+@rdm2@\cdot \frac{d}{d@v@}\left[@ppp2@\right] +@rdm3@\cdot \frac{d}{d@v@}\left[@ppp3@\right] \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot [@diff(ppp1,v)@]+@rdm2@\cdot [@diff(ppp2,v)@] +@rdm3@\cdot \left[@diff(ppp3,v)@\right] \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:rand(4)+2;
rdm2:-(2+rand(5));
rdm3:-2-rand(2);
rdm4:rdm1+1;
rdm5:5+rand(4);
rdm6:3+rand(4);
ppp1:v^rdm4;
ppp2:tan(rdm5*v);
ppp3:exp(rdm6*v);
pp1:rdm1*ppp1;
pp2:rdm2*ppp2;
pp3:rdm3*ppp3;
p:pp1+pp2+pp3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 42  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-2.a</text>
    </name>
    <questiontext format="html">
      <text>Find the following derivative:
\[\frac{d}{d@v@}\left( @p@\right) .\]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left( @p@\right)  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @pp1@\right) +\frac{d}{d@v@}\left[ @pp2@\right]  \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \frac{d}{d@v@}(@ppp1@)+@1/rdm2@\cdot \frac{d}{d@v@}[@ppp2@] \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \left( @diff(ppp1,v)@\right) +@1/rdm2@\cdot \left[ @diff(ppp2,v)@\right]  \\
   &\displaystyle =&\displaystyle  @rdm1*diff(ppp1,v)@+@1/rdm2*diff(ppp2,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:z;
rdm1:3+rand(6);
rdm2:4+rand(6);
ppp1:v^(1/rdm1);
ppp2:v^rdm2;
pp1:ppp1*rdm1;
pp2:ppp2/rdm2;
p:pp1+pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q2:ppp1+ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>You appear to have forgotten the coefficients @rdm1@ and @rdm2@.

Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 43  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-2.b</text>
    </name>
    <questiontext format="html">
      <text>Find the following derivative:
\[\frac{d}{d@v@}\left( @p@\right) .\]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left( @p@\right)  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @pp1@\right) +\frac{d}{d@v@}\left[ @pp2@\right]  \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \frac{d}{d@v@}(@ppp1@)+@rdm2@\cdot \frac{d}{d@v@}[@ppp2@] \\
   &\displaystyle =&\displaystyle  @rdm1@\cdot \left( @diff(ppp1,v)@\right) +@rdm2@\cdot \left[ @diff(ppp2,v)@\right]  \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:z;
rdm1:2+rand(4);
rdm2:-2-rand(6);
ppp1:cos(v/2/rdm1);
ppp2:exp(-v/2/rdm1);
pp1:ppp1*rdm1;
pp2:ppp2*rdm2;
p:pp1+pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q2:ppp1+ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>You appear to have forgotten the coefficients @rdm1@ and @rdm2@.

Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 44  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-2.c</text>
    </name>
    <questiontext format="html">
      <text>Find the following derivative:
\[\frac{d}{d@v@}\left( @p@\right) .\]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left( @p@\right)  &\displaystyle =&\displaystyle  @rdm1/rdm2@\cdot \frac{d}{d@v@}(@ppp1@) \\
   &\displaystyle =&\displaystyle  @rdm1/rdm2@\cdot \left( @diff(ppp1,v)@\right) \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:z;
rdm1:2+2*rand(4);
rdm2:3+2*rand(4);
ppp1:exp(rdm1*v/rdm2);
pp1:ppp1*rdm1/rdm2;
p:pp1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 45  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-2.d</text>
    </name>
    <questiontext format="html">
      <text>Find the following derivative:
\[\frac{d}{d@v@}\left( @p@\right) .\]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left( @p@\right)  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @pp1@\right) +\frac{d}{d@v@}\left[ @-pp2@\right]  \\
   &\displaystyle =&\displaystyle  @rdm1/rdm2^2@\cdot \frac{d}{d@v@}(@ppp1@)+@-rdm2/rdm1^2@\cdot \frac{d}{d@v@}[@ppp2@] \\
   &\displaystyle =&\displaystyle  @rdm1/rdm2^2@\cdot \left( @diff(ppp1,v)@\right) +@-rdm2/rdm1^2@\cdot \left[ @diff(ppp2,v)@\right]  \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:z;
rdm1:2+2*rand(3);
rdm2:3+2*rand(3);
ppp1:tan(v*rdm2/rdm1);
ppp2:cos(4*rdm1*v);
pp1:ppp1*rdm1/rdm2^2;
pp2:-rdm2*ppp2/rdm1^2;
p:pp1+pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q2:ppp1+ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>You appear to have forgotten the coefficients @rdm1@ and @rdm2@.

Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 46  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-2.e</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the following derivative: \[\frac{d}{d@v@}\left( @p@\right) .\] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In order to solve this problem we will have to use the linearity rules for differentition. calc_diff_linearity_rule \[\begin{array}{rcl}\displaystyle \frac{d}{d@v@}\left( @p@\right) &amp;\displaystyle =&amp;\displaystyle \frac{d}{d@v@}\left( @pp1@\right) +\frac{d}{d@v@}\left[ @pp2@\right] \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\cdot \frac{d}{d@v@}(@ppp1@)+@rdm2@\cdot \frac{d}{d@v@}[@ppp2@] \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\cdot \left( @diff(ppp1,v)@\right) +@rdm2@\cdot \left[ @diff(ppp2,v)@\right] \\ &amp;\displaystyle =&amp;\displaystyle @rdm1*diff(ppp1,v)@+@rdm2*diff(ppp2,v)@ \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:z;
rdm1:1/(rand([2,4,5,7,8]));
rdm2:-1/(rand(4)+2);
ppp1:v^(1/3/rdm1);
ppp2:exp(4*v*rdm2);
pp1:ppp1*rdm1;
pp2:rdm2*ppp2;
p:pp1+pp2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<p><span class="correct">Correct answer, well done.</span></p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span class="partially">Your answer is partially correct.</span></p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><span class="incorrect">Incorrect answer.</span></p>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text>q1:pp1-pp2;
q2:ppp1+ppp2;
q3:ppp1-ppp2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q1,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have lost a minus sign. Remember: \[\frac{d}{dx}(f-g)=\frac{df}{dx}-\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have forgotten the coefficients @rdm1@ and @rdm2@. Remember: \[\frac{d}{dx}(k\cdot f)=k\cdot \frac{df}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(q3,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have forgotten the coefficients @rdm1@ and @rdm2@ and combined the functions together using an incorrect sign. Remember: \[\frac{d}{dx}(a\cdot f-b\cdot g)=a\frac{df}{dx}-b\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(pp1-pp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1+ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(ppp1-ppp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 47  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-3.a</text>
    </name>
    <questiontext format="html">
      <text>Expand the powers and find the following derivative:
\[ \frac{d}{d@v@}\left[\sqrt{@pp@}\right] . \]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition <hint>calc_diff_linearity_rule</hint> and the properties of exponentials <hint>alg_indices</hint>.

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left[\sqrt{\left(@pp@\right)}\right] &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @p@\right) \\
   &\displaystyle =&\displaystyle  @rdm1^(1/2)@\cdot \frac{d}{d@v@}(@ppp^(1/2)@) \\
   &\displaystyle =&\displaystyle  @rdm1^(1/2)@\cdot @diff(ppp^(1/2),v)@ \\
   &\displaystyle =&\displaystyle  @diff(p,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:y;
rdm1:2+rand(6);
ppp:v;
pp:(rdm1)*(ppp);
p:pp^(1/2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)/(rdm1^(1/2))</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>The derivative of $af(x) = a\frac{df}{dx}$

You appear to have dropped the a factor, in this case @rdm1^(1/2)@</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)/(rdm1^(1/2))</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 48  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-3.c</text>
    </name>
    <questiontext format="html">
      <text>Expand the powers and find the following derivative:
\[ \frac{d}{d@v@}\left[\left(@pp@\right)^{@rdm2@}\right] . \]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

\[\begin{array}{rcl}\displaystyle
  \frac{d}{d@v@}\left[\left(@pp@\right)^{@rdm2@}\right] &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @p@\right) \\
   &\displaystyle =&\displaystyle  @rdm1^rdm2@\cdot \frac{d}{d@v@}(@ppp^rdm2@) \\
   &\displaystyle =&\displaystyle  @rdm1^rdm2@\cdot @diff(ppp^rdm2,v)@ \\
   &\displaystyle =&\displaystyle  @rdm1^rdm2*diff(ppp^rdm2,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:y;
rdm1:1/(2+rand(2));
rdm2:rand(4)+2;
ppp:exp(v);
pp:(rdm1)*(ppp);
p:pp^(rdm2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(p/(rdm1**rdm2),v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>The derivative of $ae^x = ae^x$

You appear to have dropped the a factor, in this case @rdm1**rdm2@</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>rdm1**rdm2*rdm2*exp(v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text>The derivative of $e^{mx} = me^{mx}$.

You appear to have lost the m factor in the exponent.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)/rdm2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text>The derivative of $e^{mx} = me^x$.

You appear to have lost the m factor.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p/(rdm1**rdm2),v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>rdm1**rdm2*rdm2*exp(v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)/rdm2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 49  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-4.c</text>
    </name>
    <questiontext format="html">
      <text>Expand the following expresion and differentiate it with respect to @v@:\[ @p@. \]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

Since

\[@p@=@q@,\]

\[\begin{array}{rcl}\displaystyle
\frac{d}{d@v@}\left[ @p@\right]  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @q@\right)  \\
   &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @qq1@\right) + \frac{d}{d@v@}\left( @qq2@\right) + \frac{d}{d@v@}\left( @qq3@\right) + \frac{d}{d@v@}\left( @qq4@\right)  \\
   &\displaystyle =&\displaystyle  @rdm1*rdm3@\frac{d}{d@v@}\left( @qqq1@\right) + @rdm1@\frac{d}{d@v@}\left( @qqq2@\right) + [@rdm2*rdm3@]\frac{d}{d@v@}\left( @qqq3@\right) + 0  \\
   &\displaystyle =&\displaystyle  @rdm1*rdm3@\left[@diff(qqq1,v)@\right] + (@rdm1@)@diff(qqq2,v)@ + [@rdm2*rdm3@]\left[@diff(qqq3,v)@\right] \\
   &\displaystyle =&\displaystyle  @rdm1*rdm3*diff(qqq1,v)@ + @rdm1*diff(qqq2,v)@ + @rdm2*rdm3*diff(qqq3,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:2+rand(2);
rdm2:-2-rand(2);
rdm3:1+rand(5);
pppp1:v;
pppp2:1/v;
pppp3:1/v^2;
pppp4:v;
ppp1:rdm1*pppp1;
ppp2:rdm2*pppp2;
ppp3:rdm3*pppp3;
ppp4:pppp4;
pp1:ppp1+ppp2;
pp2:ppp3+ppp4;
p:pp1*pp2;
qq1:ppp1*ppp3;
qq2:ppp1*ppp4;
qq3:ppp2*ppp3;
qq4:ppp2*ppp4;
qqq1:pppp1*pppp3;
qqq2:pppp1*pppp4;
qqq3:pppp2*pppp3;
qqq4:pppp2*pppp4;
q:expand(p);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(q,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords></forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(q,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(pp1+pp2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[You seem to have used the sum rule to differentiate a product. Use to product rule instead <hint>calc_product_rule</hint>.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(pp1,v)*diff(pp2,v)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[The derivative of a product is not the product of the derivative of each factor seperately.

To differentiate a product you need to use the product rule <hint>calc_product_rule</hint>.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(pp1+pp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(pp1,v)*diff(pp2,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 50  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-4.e</text>
    </name>
    <questiontext format="html">
      <text>Expand the following expresion and differentiate it with respect to @v@:\[ \frac{@pp1@}{@pp2@}. \]

[[input:ans1]] [[validation:ans1]]
</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>

Since
\[\begin{array}{rcl}\displaystyle
\frac{@pp1@}{@pp2@} &\displaystyle =&\displaystyle  \frac{1}{@pp2@}-\frac{@ppp2@}{@pp2@} \\
   &\displaystyle =&\displaystyle  @q@
 \end{array}\]

\[\begin{array}{rcl}\displaystyle
\frac{d}{d@v@}\left[\frac{@pp1@}{@pp2@}\right]  &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left( @q@\right)  \\
   &\displaystyle =&\displaystyle  \frac{d}{d@v@}\left(@1/pp2@\right) - \frac{d}{d@v@}\left( @ppp2/pp2@\right) \\
   &\displaystyle =&\displaystyle  @diff(q,v)@
 \end{array}\] ]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:-2-rand(6);
rdm2:-2-rand(5);
ppp1:1;
ppp2:exp(rdm1*v);
pp1:ppp1-ppp2;
pp2:exp(rdm2*v);
p:pp1/pp2;
q:expand(p);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(q,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<span class='correct'>Correct answer, well done.</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span class='partially'>Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span class='incorrect'>Incorrect answer.</span>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(q,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords></forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(q,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 51  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-5.a</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use the laws of logarithms to find the following derivative: \[ \frac{d}{d@v@} \left[\ln(@pp@)\right]. \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Using the laws of logarithms alg_logarithms we get \[\begin{array}{rcl}\displaystyle \ln(@pp@) &amp;\displaystyle =&amp;\displaystyle @rdm1@@pq@ \end{array}\] So, by linearity calc_diff_linearity_rule, \[\begin{array}{rcl}\displaystyle \frac{d}{d@v@}\left[\ln(@pp@)\right] &amp;\displaystyle =&amp;\displaystyle \frac{d}{d@v@}\left[@rdm1@@pq@\right] \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@\frac{d}{dx}[@pq@] \\ &amp;\displaystyle =&amp;\displaystyle @rdm1@@diff(pq,v)@ \\ &amp;\displaystyle =&amp;\displaystyle @diff(p,v)@ \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:(1+2*rand(6))/(2+2*rand(2));
pp:v^(rdm1);
p:log(pp);
pq:log(v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<p><span class="correct">Correct answer, well done.</span></p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span class="partially">Your answer is partially correct.</span></p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><span class="incorrect">Incorrect answer.</span></p>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

<!-- question: 52  -->
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-5.b</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use the laws of logarithms to find the following derivative: \[ \frac{d}{d@v@} \left[@p@\right]. \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Using the laws of logarithms alg_logarithms we get \[\begin{array}{rcl}\displaystyle @p@ &amp;\displaystyle =&amp;\displaystyle \ln(1) - \ln(@pp@) \\ &amp;\displaystyle =&amp;\displaystyle -\ln(@ppp@) - \ln(@rdm2@) \\ &amp;\displaystyle =&amp;\displaystyle @-rdm1@@pq@ - \ln(@rdm2@) \end{array}\] So, by linearity calc_diff_linearity_rule, \[\begin{array}{rcl}\displaystyle \frac{d}{d@v@}\left[@p@\right] &amp;\displaystyle =&amp;\displaystyle \frac{d}{d@v@}\left[@-rdm1@@pq@ - \ln(@rdm2@)\right] \\ &amp;\displaystyle =&amp;\displaystyle \frac{d}{d@v@}\left[@-rdm1@@pq@\right] + \frac{d}{d@v@}\left[-\ln(@rdm2@)\right] \\ &amp;\displaystyle =&amp;\displaystyle @-rdm1@\frac{d}{d@v@}\left[@pq@\right] + 0 \\ &amp;\displaystyle =&amp;\displaystyle @-rdm1@@diff(pq,v)@ \\ &amp;\displaystyle =&amp;\displaystyle @-rdm1*diff(pq,v)@ \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>v:x;
rdm1:1/(2+rand(2));
rdm2:(rand([2,3,5,6,7]))**rdm1;
ppp:v^(rdm1);
pp:rdm2*ppp;
p:log(1/pp);
pq:log(v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d@v@}(@p@) = @diff(p,v)@ \]</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <prtcorrect format="html">
      <text><![CDATA[<p><span class="correct">Correct answer, well done.</span></p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span class="partially">Your answer is partially correct.</span></p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><span class="incorrect">Incorrect answer.</span></p>]]></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <forbidwords>diff</forbidwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>