<?xml version='1.0' encoding='utf-8'?>
<quiz><question type="stack">
    <name>
      <text>1st-Order-ODE-1-sep</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Solve the following differential equation: \[ {@v@}\frac{\mathrm{d}y}{\mathrm{d}{@v@}} + {@n2@} {@v@}+{@n1@} = 0. \] Use the letter "c" (small case) to denote the arbitrary constant in the general solution. &lt;br&gt;\(y({@v@})\) =[[input:ans1]] [[validation:ans1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;p&gt;For the equation \[ {@v@}\frac{\mathrm{d}y}{\mathrm{d}{@v@}} + {@n2@}{@v@}+{@n1@} = 0 \] rearranging gives \[\mathrm{d}y = \frac{(-{@n2@}{@v@}-{@n1@})}{@v@}\mathrm{d}{@v@}\] Therefore \[\int \mathrm{d}y = \int\frac{(-{@n2@}{@v@}-{@n1@})}{@v@}d{@v@} =\int(-{@n2@}-\frac{{@n1@}}{{@v@}})\mathrm{d}{@v@}\] Hence \[y = {@correctanswer@}\]&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>n1:rand(8)+2;
n2:rand(6)+2;
v:x;
fans:-n2-n1/v;
correctanswer:int(fans,v)+c;
wronganswer:-n1*ln(v)-n2*v+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:Result]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>\({@v@}\frac{\mathrm{d}y}{\mathrm{d}{@v@}} + {@n2@}{@v@}+{@n1@} = 0\)  gives {@correctanswer@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;&lt;span class="correct"&gt;Correct answer, well done.&lt;/span&gt;&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;&lt;span class="partially"&gt;Your answer is partially correct.&lt;/span&gt;&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;&lt;span class="incorrect"&gt;Incorrect answer.&lt;/span&gt;&lt;/p&gt;</text>
    </prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>diff(ans1,x)</sans>
        <tans>fans</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.7500000</truescore>
        <truepenalty />
        <truenextnode>2</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;Your answer satisfies the differential equation, but is not in the correct form.&amp;nbsp;&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>member(c,listofvars(ans1))</sans>
        <tans>true</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2500000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Your answer should have the general constant \(c\), but does not.&amp;nbsp;&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1645139326</deployedseed>
    <deployedseed>481337292</deployedseed>
    <deployedseed>1509361853</deployedseed>
    <deployedseed>346176631</deployedseed>
    <deployedseed>915142823</deployedseed>
    <deployedseed>963104251</deployedseed>
    <deployedseed>793273982</deployedseed>
    <deployedseed>1455145271</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(correctanswer-c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.7500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-3-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>wronganswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>correctanswer-c+Q</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.7500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-3-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>


  </quiz>