<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Complex_Quotient_2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Convert the following number into the form \(p+qi\) where \(p\) and \(q\) are real. Note: give your answer in the form \(p+qi\), in the box below. \[ \frac{2}{@c1@} + \frac{5}{ {@c2@} } \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Multiplying the top and bottom by the conjugate of the denominator gives \[ \frac{2}{@c1@}*\frac{@con1@}{@con1@}+\frac{5}{@c2@}*\frac{@con2@}{@con2@}=\] \[\frac{{@2*n1@}-{@2*n2@}i}{{@n1*n1@}-{@n12@}i+{@n12@}i-{@n2*n2@}i^2}+\frac{{@5*n3@}+{@5*n4@}i}{{@n3*n3@}-{@n34@}i+{@n34@}i-{@n4*n4@}i^2} \] and because \[i^2= -1\] becomes \[\frac{{@2*n1@}-{@2*n2@}i}{@n1*n1+n2*n2@}+\frac{{@5*n3@}+{@5*n4@}i}{@n3*n3+n4*n4@}=\] \[\frac{{@2*n1*(n3*n3+n4*n4)@}-{@2*n2*(n3*n3+n4*n4)@}i}{@(n1*n1+n2*n2)*(n3*n3+n4*n4)@}+\frac{{@5*n3*(n1*n1+n2*n2)@}+{@5*n4*(n1*n1+n2*n2)@}i}{@(n1*n1+n2*n2)*(n3*n3+n4*n4)@}={@correctanswer@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(5)+1;
  n2:rand(5)+1;
  n3:rand(5)+1;
  n4:rand(5)+1;
  c1:ev(n1+n2*i,simp);
  con1:ev(n1-n2*i,simp);
  c2:ev(n3-n4*i,simp);
  con2:ev(n3+n4*i,simp);
  n12:n1*n2;
  n34:n3*n4;
  pval:(2*n1*(n3*n3+n4*n4)+5*n3*(n1*n1+n2*n2))/((n1*n1+n2*n2)*(n3*n3+n4*n4));
  qval:(-2*n2*(n3*n3+n4*n4)+5*n4*(n1*n1+n2*n2))/((n1*n1+n2*n2)*(n3*n3+n4*n4));
  correctanswer:ev(pval+qval*i,simp);
  pwrong:(2*n1*(n3*n3+n4*n4)+5*n3*(n1*n1+n2*n2));
  qwrong:(-2*n2*(n3*n3+n4*n4)+5*n4*(n1*n1+n2*n2));
  wrong1:ev(pwrong+qwrong*i,simp)/ev((n1*n1+n2*n2)*(n3*n3+n4*n4),simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\frac{2}{@c1@} + \frac{5}{ {@c2@}} = {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sansr:ev(realpart(ans1),simp);
  sansi:ev(imagpart(ans1),simp);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sansr</sans>
        <tans>pval</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The real part of your answer is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sansi</sans>
        <tans>qval</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The imaginary part of your answer is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is equivalent to the correct answer, but not in the required form \(p + qi\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1983864448</deployedseed>
    <deployedseed>863999758</deployedseed>
    <deployedseed>2063008379</deployedseed>
    <deployedseed>217094918</deployedseed>
    <deployedseed>2047168758</deployedseed>
    <deployedseed>1664741582</deployedseed>
    <deployedseed>1506005639</deployedseed>
    <deployedseed>1869040393</deployedseed>
    <deployedseed>564203982</deployedseed>
    <deployedseed>1770619118</deployedseed>
    <deployedseed>417076978</deployedseed>
    <deployedseed>1401318241</deployedseed>
    <deployedseed>616219738</deployedseed>
    <deployedseed>1346660649</deployedseed>
    <deployedseed>1537762140</deployedseed>
    <deployedseed>931546217</deployedseed>
    <deployedseed>1597149896</deployedseed>
    <deployedseed>1639117202</deployedseed>
    <deployedseed>1005637699</deployedseed>
    <deployedseed>1154136534</deployedseed>
    <deployedseed>1988046325</deployedseed>
    <deployedseed>37244055</deployedseed>
    <deployedseed>1536469769</deployedseed>
    <deployedseed>1470164981</deployedseed>
    <deployedseed>2044994735</deployedseed>
    <deployedseed>1994847114</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wrong1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer+1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer+i</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
