<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.3.2.8 task triple integral 1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Evaluate the  triple integral:
\(\displaystyle \int_0^{@cc@} \int_0^{@bb@} \int_0^{@aa@} x^3 y^2 z \;\mathrm{d}x \,\mathrm{d}y\,\mathrm{d}z\ =\ \)[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[\(\begin{align}  
\displaystyle\int_0^{@cc@} \displaystyle\int_0^{@bb@} \displaystyle\int_0^{@aa@} 
x^3y^2z\;\mathrm{d}y\,\mathrm{d}z \,\mathrm{d}x &=  \displaystyle\int_0^{@cc@} \displaystyle\int_0^{@bb@}  \displaystyle\Bigg[{@fbI11@}\Bigg]_0^{@aa@} \;\mathrm{d}z \,\mathrm{d}x
\\ &= \displaystyle\int_0^{@cc@} \displaystyle\int_0^{@bb@}  \displaystyle{@I11@}\;\mathrm{d}z \,\mathrm{d}x
 \\ &= \displaystyle\int_0^{@cc@} \Bigg[{@fbI12@}\Bigg]_0^{@bb@} \;\mathrm{d}x
\\ &= \displaystyle\int_0^{@cc@} {@I12@}\;\mathrm{d}x
 \\ &= \Bigg[{@fbI13@}\Bigg]_0^{@cc@} = {@ta1@} 
 \end{align}\)]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1+rand(3)
bb: rand_with_prohib(1,3,[aa])
cc: 1+rand(3)

fbI11: integrate(x^3*y^2*z,x);
I11: integrate(x^3*y^2*z,x,0,aa);
fbI12: integrate (I11,y);
I12: integrate (I11,y,0,bb);
fbI13: integrate (I12,z);
I13: integrate (I12,z,0,cc);
ta1: I13;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>\(\displaystyle \int_0^{@cc@} \int_0^{@bb@} \int_0^{@aa@} x^3 y^2 z \;\mathrm{d}x \,\mathrm{d}y\,\mathrm{d}z={@ta1@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1086825569</deployedseed>
    <deployedseed>1204603328</deployedseed>
    <deployedseed>750775463</deployedseed>
    <deployedseed>1945814492</deployedseed>
    <deployedseed>1245705006</deployedseed>
    <deployedseed>1622687948</deployedseed>
    <deployedseed>128580004</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
