<?xml version='1.0' encoding='utf-8'?>
<quiz><question type="stack">
    <name>
      <text>2nd-Order-ODE-find-parameter</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Consider the following ODE \[\frac{\mathrm{d}^2y}{\mathrm{d}t^2} + {@n1@} \frac{\mathrm{d}y}{\mathrm{d}t} + k^2 y(t) =0.\]&lt;/p&gt;
&lt;p&gt;For what values of \(k&amp;gt;0\) does the graph of the solution look qualitatively of the form shown below? &lt;br /&gt;{@plot(qa3,[x,0,100])@}&lt;/p&gt;
&lt;p&gt;[Notes, (1) you are not being asked to solve the ODE and (2) "qualitative" means of the same general shape.]&lt;/p&gt;
&lt;p&gt;You must enter your answer as inequalities in \(k\) to represent a range of real numbers.  For example,&lt;/p&gt;
&lt;ul&gt;
&lt;li&gt;to represent the interval \( [2,3) \) you might type  &lt;span style="font-family: 'courier new', courier, monospace;"&gt;2&amp;lt;=k and k&amp;lt;3&lt;/span&gt;&lt;/li&gt;
&lt;li&gt;to represent the interval \( [0,\infty) \) you might type  &lt;span style="font-family: 'courier new', courier, monospace;"&gt;0&amp;lt;=k&lt;/span&gt;&lt;/li&gt;
&lt;li&gt;you can connect individual inequalities using &lt;span style="font-family: 'courier new', courier, monospace;"&gt;and&lt;/span&gt; or &lt;span style="font-family: 'courier new', courier, monospace;"&gt;or&lt;/span&gt;, as the logic dictates in your answer. &lt;/li&gt;
&lt;/ul&gt;
&lt;p&gt; &lt;/p&gt;
&lt;p&gt;[[input:ans1]]&lt;/p&gt;
&lt;div&gt;[[validation:ans1]]&lt;/div&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;p&gt;Consider the ODE \[\frac{\mathrm{d}^2y}{\mathrm{d}t^2} + {@n1@} \frac{\mathrm{d}y}{\mathrm{d}t} + k^2 y(t) =0.\]&lt;/p&gt;
&lt;p&gt;Substitute in \(y(t)=e^{mt}\) so we have&amp;nbsp;\[m^2 + {@n1@} m + k^2 m =0.\]&lt;br&gt;The solutions to this equation are given by \[ m = \frac{-{@n1@} \pm \sqrt{{@n1@}^2-4\times k^2}}{2} = -{@n1/2@}\pm \sqrt{{@n1^2/4@}-k^2}\]&lt;br&gt;The only way to have oscillatory&amp;nbsp;solutions is to have a complex root, and this can only happen if \({@n1^2/4@}-k^2&amp;lt;0\), i.e. if \(k^2&amp;gt;{@n1^2/4@}\). &amp;nbsp;Since we know \(k&amp;gt;0\) we can eliminate the negative values to give \(k&amp;gt;\sqrt{{@n1^2/4@}}\), i.e. \(k&amp;gt;{@n1/2@}.\).&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>n1:1/(2*(rand(4)+2)^2);
ta:k&gt;n1/2
q:'diff(y,x,2) + n1*'diff(y,x)+(sqrt(n1/2)+2)*y = 0;
qa:ode2(q,y,x);
qa2:ic2(qa,x=0,y=1,'diff(y,x)=0)
qa3:rhs(qa2)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>{@ta@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;Your answer is partially correct.&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>k&gt;=sqrt(n1/2)</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>k&lt;sqrt(n1/2)</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>k&lt;sqrt(n1/2)</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>k^2&gt;n1/2</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;This is basically correct, but you need to write your inequality in simplified form.&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>k&gt;sqrt(n1/2) or k&lt;-sqrt(n1/2)</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;You can eliminate the negative values of \(k\) using the given condition that \(k&amp;gt;0\).&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1685192574</deployedseed>
    <deployedseed>1005426202</deployedseed>
    <deployedseed>371097510</deployedseed>
    <deployedseed>1456215945</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>k&gt;=sqrt(n1/2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>k&lt;sqrt(n1/2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>k^2&gt;n1/2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <testinput>
        <name>ans1</name>
        <value>k^2&gt;=n1/2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-6-F</expectedanswernote>
      </expected>
    </qtest>
  </question>


  </quiz>