<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.3.9.2 Factorise cubic given factor II [NOT IN QUIZ]</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Check for yourself that \(x={@aa@}\) is a root of {@eqn@} and locate the other roots algebraically. Hence record all solutions:</p><p></p>\( x = \)&nbsp;[[input:ans1]] [[validation:ans1]] (give your answer as a set of numbers).&nbsp;[[feedback:prt1]]<p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the expression \({@eqn@}\) for which we know \(x = {@aa@}\) is a root. Therefore \({@factor@}\) is a factor. Then we have \({@eqn@} = ({@factor@})(ax^2 + bx + c)\).</p>

<p>Expanding the brackets we obtain \[\begin{aligned} {@eqn@} &amp;= {@a*x^3 + b*x^2 + c*x - aa*a*x^2 - aa*b*x - aa*c@}, \\ &amp;= {@a*x^3 + (b - aa*a )*x^2 + (c - aa*b )*x - aa*c@}. \end{aligned}\]</p>

<p>Comparing coefficients we obtain \(a = 1\), \(b = {@- bb - cc @}\), and \(c = {@ bb*cc@}\).</p>
<p>Therefore we have \({@eqn@} = ({@factor@})({@quad@})\). Factorising, we obtain \[{@eqn@} = {@%eqn@}.\]</p>

<p>Therefore the solutions are \(x - \{ {@aa@}, {@bb@}, {@cc@} \} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(3)-3;
bb: (rand(3)-6);
cc: rand(3)-9;
ta: setify([aa,bb,cc]);

simp:true;
factor: x - aa;
quad: x^2 + (- bb - cc)*x + bb*cc;
%eqn: (x - cc)*(x - bb)*(x - aa);
eqn: x^3 - (aa + bb + cc)*x^2 + (aa*bb + bb*cc + cc*aa)*x - aa*bb*cc;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eqn@},  sol: [{@aa@},{@bb@},{@cc@}]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{a,b,c}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Try taking the factor of \( ({@factor@}) \) out of {@eqn@} in order to be left with a quadratic function which you can factorise yourself or solve using the quadratic formula.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1246432368</deployedseed>
    <deployedseed>465202749</deployedseed>
    <deployedseed>476513597</deployedseed>
    <deployedseed>1581570574</deployedseed>
    <deployedseed>1699434329</deployedseed>
    <deployedseed>2106935022</deployedseed>
    <deployedseed>1250647786</deployedseed>
    <deployedseed>790208410</deployedseed>
    <deployedseed>1762570545</deployedseed>
    <deployedseed>990580994</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{aa,bb,cc}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{ev(aa+1,simp),bb,cc}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{aa,bb}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
