<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.3.9.1 Factorise cubic given factor I [NOT IN QUIZ]</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Factorise {@eqn@} given that {@factor@} is a factor:</p><p>\( {@eqn@} = ( x + \) [[input:ans1]] [[validation:ans1]] \() ( x + \)&nbsp;[[input:ans2]] [[validation:ans2]] \() ( x + \)&nbsp;[[input:ans3]] [[validation:ans3]] \( ) \).&nbsp;[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the expression \({@eqn@}\) for which we know \({@factor@}\) is a factor. Then we have \({@eqn@} = ({@factor@})(ax^2 + bx + c)\).</p>

<p>Expanding the brackets we obtain \[\begin{aligned} {@eqn@} &amp;= ax^3 + bx^2 + cx + {@aa@}ax^2 + {@aa@}bx + {@aa@}c, \\ &amp;= ax^3 + (b + {@aa@}a )x^2 + (c + {@aa@}b )x + {@aa@}c. \end{aligned}\]</p>

<p>Comparing coefficients we obtain \(a = 1\), \(b = {@ bb + cc @}\), and \(c = {@ bb*cc@}\).</p>
<p>Therefore we have \({@eqn@} = ({@factor@})({@quad@})\). Factorising, we obtain \[{@eqn@} = {@%eqn@}.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(5)+5;
bb: (rand(7)-2);
cc: rand(7)-9;

simp:true;
factor: x + aa;
quad: x^2 + (bb+cc)*x + bb*cc;
%eqn: (x + cc)*(x + bb)*(x + aa);
eqn: x^3 + (aa + bb + cc)*x^2 + (aa*bb + bb*cc + cc*aa)*x + aa*bb*cc;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eqn@},  sol: [{@aa@},{@bb@},{@cc@}]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>cc</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>bb</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>aa</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>SetSA: {ans1,ans2,ans3};
SetTA: {aa,bb,cc};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>SetSA</sans>
        <tans>SetTA</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2136608191</deployedseed>
    <deployedseed>833035052</deployedseed>
    <deployedseed>1653911233</deployedseed>
    <deployedseed>1846187096</deployedseed>
    <deployedseed>1501712263</deployedseed>
    <deployedseed>207276014</deployedseed>
    <deployedseed>574946348</deployedseed>
    <deployedseed>622997222</deployedseed>
    <deployedseed>1876819272</deployedseed>
    <deployedseed>1884932523</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>cc</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>bb</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>aa</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>aa</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>bb</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>cc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>bb</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>aa</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>cc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(cc+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>bb</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>aa</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
