<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.3.7.3 Solve cubic given a root III (NOT IN QUIZ)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Verify that the given values \(x={@x1@}\) and \(x={@x4@}\)&nbsp;are solutions of the equation {@eqn@} and hence find all solutions:</p>
<p>\( x = \)&nbsp;[[input:ans1]] [[validation:ans1]]&nbsp;</p>
<p><b>Give your answer as a set of precise numbers.</b><br></p>

<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Try taking the factors of \( ({@x-x1@}) \) and&nbsp;\( ({@x-x4@}) \) out of {@eqn@} in order to be left with a quadratic function \( {@x^2+2*bb*x+bb^2-cc@} \) which you can solve using the quadratic formula to obtain all solutions {@ta@}.<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>aa: -rand(2)-1;
bb: rand(3)+1;
cc: rand(4)+5;
dd: rand(3)+4;

x1: -aa;
x2: -bb+sqrt(cc);
x3: -bb-sqrt(cc);
x4: dd;
ta: setify([x1,x2,x3,x4]);

simplify:true;
eqn: x^4 - aa*x^3 + 2*bb*x^3 - dd*x^3 - 2*aa*bb*x^2 + aa*dd*x^2 + bb^2*x^2 - 2*bb*dd*x^2 - cc*x^2 - aa*bb^2*x + 2*aa*bb*dd*x + aa*cc*x - bb^2*dd*x + cc*dd*x + aa*bb^2*dd - aa*cc*dd;
eqnattempt: (x - dd)*(x^3 + (2*bb - aa)*x^2 - (2*aa*bb - bb^2 + cc)*x - aa*bb^2 + aa*cc);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eqn@},  sol: {{@x1@},{@x2@},{@x3@},{@x4@}}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{a,b,c,d}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Try taking the factors of \( (x - {@x1@}) \) and \( (x - {@x4@}) \) out of {@eqn@} in order to be left with a quadratic function which you can solve using the quadratic formula.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1742691085</deployedseed>
    <deployedseed>1362361432</deployedseed>
    <deployedseed>970579714</deployedseed>
    <deployedseed>1275295553</deployedseed>
    <deployedseed>1214494273</deployedseed>
    <deployedseed>1702801914</deployedseed>
    <deployedseed>794857698</deployedseed>
    <deployedseed>1640313158</deployedseed>
    <deployedseed>1142869215</deployedseed>
    <deployedseed>908452208</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{x1,x2,x3,x4}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{x1,x2,x3}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{ev(x1+1,simp),x2,x3,x4}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
