<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>19.3.6.4 Exercise 6</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<h4 class="HELM_exercise">Exercise</h4>
<p>Find a particular integral for the equation</p>
\[{@deq@}.\]
<p>\(y_\text{p}(x) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We are asked to find a particular integral for the second order, constant coefficients, inhomogeneous differential equation</p>
\[{@deq@}.\]
<p>Since the right-hand side is a linear term, we try a particular integral of the for \(y_\text{p}(x) = {@pitest@}\). We first find the derivatives:</p>
\[\begin{aligned}
\frac{\mathrm{d}y_\text{p}}{\mathrm{d}x} &amp;= {@dpitest@}, \\
\frac{\mathrm{d}^2y_\text{p}}{\mathrm{d}x^2} &amp;= {@ddpitest@}.
\end{aligned}\]
<p>Substituting these into the original equation gives</p>
\[{@pieq@}.\]
<p>Equating the constant terms and the coefficients of \(x\) gives the simultaneous equations</p>
\[\begin{cases}
{@lhs(pieq1)@} &amp;= {@rhs(pieq1)@} \\
{@lhs(pieq2)@} &amp;= {@rhs(pieq2)@}
\end{cases}\]
<p>These give the values \(a = {@aval@}\) and \(b = {@bval@}\). Hence, a particular integral for the differential equation is</p>
\[y_\text{p}(x) = {@pi1@}.\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>/* Coefficient randomisation */
aa: rand_with_prohib(-3, 3, [0]);
bb: rand_with_prohib(-3, 3, [0]);

/* Equation */
f: aa + bb * x;
deq: 'diff(y, x, 2) + aa * 'diff(y, x) + bb * y = f;

/* Needed for solution */
pitest: a + b * x;
dpitest: diff(pitest, x);
ddpitest: diff(dpitest, x);
pieq: ddpitest + aa * dpitest + bb * pitest = f;
pieq1: aa * b + bb * a = aa;
pieq2: bb *  b = bb;
pieqsol: linsolve([pieq1, pieq2], [a, b]);
aval: rhs(pieqsol[1]);
bval: rhs(pieqsol[2]);
pi1: x;

ta1: x;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@deq@}, \quad y_\text{p}(x) = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1886719139</deployedseed>
    <deployedseed>245365869</deployedseed>
    <deployedseed>1263391949</deployedseed>
    <deployedseed>1667040104</deployedseed>
    <deployedseed>1206144133</deployedseed>
    <deployedseed>383421869</deployedseed>
    <deployedseed>1656692094</deployedseed>
    <deployedseed>810023668</deployedseed>
    <deployedseed>1489485195</deployedseed>
    <deployedseed>452386653</deployedseed>
    <deployedseed>782542690</deployedseed>
    <deployedseed>862203898</deployedseed>
    <deployedseed>665708911</deployedseed>
    <deployedseed>497249398</deployedseed>
    <deployedseed>1106790187</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
