<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>AlgMap-3.4</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Express the following, as a single fraction: \[{@a/b@}-{@c/d@}.\] <br /> [[input:ans]] <br /> <br /> Leave your answer in it's lowest terms in the form \(\frac{a}{b}\). [[validation:ans]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The first task in subtracting fractions is to rewrite each as an equivalent fraction with the same denominators. You can guarantee this by multiplying the two denominators together or find the Lowest Common Multiple (LCM).</p>
<p>In this example, \({@a/b@}-{@c/d@}\), the we need the lowest common multiple of {@b@} and {@d@} which is \({@lcm@}\).</p>
<p>Next re-write the expression in the form: \[\frac{a}{@lcm@}-\frac{b}{@lcm@}.\] To convert a fraction to an equivalent form we must multiply the denominator and numerator by the same number. In this case, multiply the numerator of the first fraction by \({@ev(lcm/b,simp)@}\) and the numerator of the second by \({@ev(lcm/d,simp)@}\). So we obtain: \[{@a/b@}-{@c/d@} = \frac{@a*ev((lcm/b),simp)@}{@b*ev((lcm/b),simp)@}-\frac{@c*ev((lcm/d),simp)@}{@d*ev((lcm/d),simp)@} = \frac{@ev(a*(lcm/b),simp)@}{@lcm@}-\frac{@ev(c*(lcm/d),simp)@}{@lcm@} = \frac{@ev(a*(lcm/b)-c*(lcm/d),simp)@}{@lcm@}.\] We may need to cancel any common factors in the numerator and denominator to obtain the answer in lowest terms \({@ta@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>simp:true;
b:rand([2,5]);
a:rand([3,7])^rand([0,1]);
d:rand([3,7,11]);
c:rand([2,5,13])^rand([0,1]);
ta:(a/b)-(c/d);
lcm:lcm(b,d);
simp:false;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ {@a/b@}-{@c/d@} = {@ta@}. \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>1-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>(a-c)/(b-d)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-1-T (naive adding of fractions)</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Remember that the first step in subtracting fractions is to make both denominators equal. It is not simply subtracting the numerators and the denominators.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>1-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Simplify the fraction fully into the form \(\frac{a}{b}\) and leave in it's lowest terms.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>1-2-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1167546101</deployedseed>
    <deployedseed>1634674245</deployedseed>
    <deployedseed>431854745</deployedseed>
    <deployedseed>569073827</deployedseed>
    <deployedseed>2145329188</deployedseed>
    <deployedseed>510790094</deployedseed>
    <deployedseed>186599987</deployedseed>
    <deployedseed>226443300</deployedseed>
    <deployedseed>1340561784</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>(a-c)/(b-d)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>1-1-T (naive adding of fractions)</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>a/b-c/d</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
