<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.3.2.4 Finding the exponential form</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
Express the complex number \(z = \displaystyle{@zed@}\)&nbsp; in exponential form. The argument of \(z\) should be given precisely in terms of \(\pi\), i.e. do not give your answer in decimal form:<p>\(z=\quad\)[[input:ans1]][[validation:ans1]]&nbsp;[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We want to express \(z\) in the form \(re^{i\theta}\).</p><p>We calculate \(r = \sqrt{({@re@})^2 + ({@im@})^2} = {@arr@}\) and \(\theta = \tan^{-1}\left({@im@}/{@re@}\right) = {@th@}\).</p><p>Therefore: \[z = {@arr@}e^{{@th@}i}.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023102700</text>
    </stackversion>
    <questionvariables>
      <text>th: rand([%pi/3,%pi/4,-%pi/3,-%pi/4])
arr: 2 + rand(10);
re: arr*cos(th);
im: arr*sin(th);
simp:false;
ta: arr*e^(th*%i);
zed: display_complex(re + im*%i);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ {@ta@} = {@zed@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: green; --darkreader-inline-color:#7fff7f;" data-darkreader-inline-color=""><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: orange; --darkreader-inline-color:#ffb733;" data-darkreader-inline-color=""><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: red; --darkreader-inline-color:#ff3333;" data-darkreader-inline-color=""><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>13</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1474763257</deployedseed>
    <deployedseed>1763096510</deployedseed>
    <deployedseed>380006334</deployedseed>
    <deployedseed>959660418</deployedseed>
    <deployedseed>611076165</deployedseed>
    <deployedseed>1658317878</deployedseed>
    <deployedseed>179907270</deployedseed>
    <deployedseed>1200483456</deployedseed>
    <deployedseed>580584991</deployedseed>
    <deployedseed>799573363</deployedseed>
    <deployedseed>262346972</deployedseed>
    <deployedseed>2099828996</deployedseed>
    <deployedseed>1735409900</deployedseed>
    <deployedseed>1481936188</deployedseed>
    <deployedseed>529943212</deployedseed>
    <deployedseed>749780944</deployedseed>
    <deployedseed>319202744</deployedseed>
    <deployedseed>33236441</deployedseed>
    <deployedseed>1938388878</deployedseed>
    <deployedseed>908547010</deployedseed>
    <deployedseed>697416402</deployedseed>
    <deployedseed>1019055399</deployedseed>
    <deployedseed>2090589336</deployedseed>
    <deployedseed>1834674709</deployedseed>
    <deployedseed>522951543</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>0 + 2*%i</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
