<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>18.4.4.2 Exercise 1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The sides of a right-angled triangle enclosing the right-angle are measured as \({@side_1@} \, \textrm{m}\) and \({@side_2@} \, \textrm{m}.\) The maximum errors in each measurement are \(\pm {@max_error@} \, \textrm{m}.\) Find the maximum error in the
    calculated area.</p>
<p>The maximum error is [[input:ans1]] [[validation:ans1]] \(\text{m}^2\).</p><p>Note: Keep 1 s.f. in your answer.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The area of a right-angled triangle is \[A=\frac{1}{2}xy\] where \(x, \, y\) are the sides enclosing the right-angle.</p>
<p>First we calculate \(\dfrac{\partial A}{\partial x}\) and \(\dfrac{\partial A}{\partial y}.\)</p>
<p>\[ \displaystyle\frac{\partial A}{\partial x} = {@f_x@} \quad \text{and} \quad \displaystyle \frac{\partial A}{\partial y} = {@f_y@} \]</p>
<p>Therefore, \[\delta A \approx \frac{\partial A}{\partial x} \delta x + \frac{\partial A}{\partial y} \delta y = {@f_x@} \delta x + {@f_y@} \delta y\]</p>
<p>Letting \(x={@side_1@}, \, y={@side_2@}\) and \( \delta x = \delta y = {@max_error@}\), we obtain \[\delta A \approx \frac{{@side_2@}}{2} \delta x + \frac{{@side_1@}}{2} \delta y = {@(side_2/2)*max_error@}+{@(side_1/2)*max_error@} = {@((side_2/2)*max_error)+((side_1/2)*max_error)@}\text{m}^2.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Randomisation needs rethinking. Ideally we want max error<1 and keep 1 s.f */

exp: x*y/2;
f_x: diff(exp,x);
f_y: diff(exp,y);
side_1: rand([2,3,4,5]);
side_2: rand_with_prohib(2,7,[side_1]);
max_error: rand([0.1,0.2]);
ta1: significantfigures(((side_2/2)*max_error)+((side_1/2)*max_error),1);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>sides:\({@side_1@} \, \textrm{m}, \, {@side_2@} \, \textrm{m}\) with max error \(\pm {@max_error@} \, \textrm{m}\) so max error in Area is \( {@ta1@} \textrm{m}^2\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumSigFigs</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>849619846</deployedseed>
    <deployedseed>1833940146</deployedseed>
    <deployedseed>1263211245</deployedseed>
    <deployedseed>997696974</deployedseed>
    <deployedseed>1434072722</deployedseed>
    <deployedseed>503678667</deployedseed>
    <deployedseed>1663885729</deployedseed>
    <deployedseed>1425372707</deployedseed>
    <deployedseed>1922890774</deployedseed>
    <deployedseed>1298262089</deployedseed>
    <deployedseed>483763035</deployedseed>
    <deployedseed>547332148</deployedseed>
    <deployedseed>622381059</deployedseed>
    <deployedseed>316402484</deployedseed>
    <deployedseed>1624076338</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+0.05,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
