<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.4.4.2  The vector product</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>Find \(\underline{p}\times\underline{q}\) where&nbsp;<span class="math inline">\(\underline{p}=-{@p1@}\underline{ i}-{@p2@}\underline{ j}\)</span> and <span class="math inline">\(\underline{ q}={@q1@}\underline{ i}+{@q2@}\underline{ j}\)</span>.</p><p>[[input:ans1]] \(\underline{k}\) [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We have&nbsp;</p><p><span class="math inline">\(\quad\underline{p}\times\underline{q}=\left|
\begin{matrix}
\underline{i} &amp; \underline{j} &amp; \underline{k} \\
-{@p1@} &amp; -{@p2@} &amp; 0 \\ {@q1@} &amp; {@q2@} &amp; 0 \\
\end{matrix}
\right|\)</span>&nbsp;</p><p>which, when evaluated, gives&nbsp;</p><p><span class="math inline">\(\qquad\underline{p}\times \underline{q}=\underline{i}((-{@p2@})(0)-(0)({@q2@})) - \underline{j}((-{@p1@})(0)-(0)({@q1@})) + \underline{k}((-{@p1@})({@q2@})-(-{@p2@})({@q1@}))\)</span></p><p><span class="math inline">\(\qquad\qquad \; ={@ta@}\underline{k}\)</span>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>p1: 2+rand(1)
p2: 2+rand(2)
q1: 2+rand(3)
q2: 2+rand(4)
ta: -p1*q2 + p2*q1</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( (-{@p1@}\underline{ i}-{@p2@}\underline{ j}) \times ({@q1@}\underline{ i}+{@q2@}\underline{ j} = {@ta@}\underline{k}) \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1997058798</deployedseed>
    <deployedseed>1904871081</deployedseed>
    <deployedseed>1892790120</deployedseed>
    <deployedseed>1214274705</deployedseed>
    <deployedseed>2141936867</deployedseed>
    <deployedseed>1626272991</deployedseed>
    <deployedseed>2017049925</deployedseed>
    <deployedseed>825732853</deployedseed>
    <deployedseed>2025673787</deployedseed>
    <deployedseed>2066887844</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
