<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.4.3.2 Evaluate the vector product</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>Use the key point of this section to find the vector product of<br>
<span class="math inline">\(\underline{ p}={@p1@}\underline{ i}+{@p2@}\underline{ j}\)</span> and <span class="math inline">\(\underline{ q}={@q1@}\underline{ i}-{@q2@}\underline{ j}\)</span>.</p>

<p>(Note that in this case there are no <span class="math inline">\(\underline{k}\)</span> components so <span class="math inline">\(a_3\)</span> and <span class="math inline">\(b_3\)</span> are both zero).</p><p></p><p><span class="math inline">\(\qquad \underline{ p}\times\underline{ q}=\)&nbsp;</span>[[input:ans1]] [[validation:ans1]] \(\underline{k}\) [[feedback:prt1]]</p><p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are ginen the vectors: \[ \underline{ p}={@p1@}\underline{ i}+{@p2@}\underline{ j} \, \text{ and } \, \underline{ q}={@q1@}\underline{ i}-{@q2@}\underline{ j}.\]</p><p>Identifying <span class="math inline">\(a_1={@p1@}\)</span>, <span class="math inline">\(a_2={@p2@}\)</span>, <span class="math inline">\(a_3=0\)</span>, <span class="math inline">\(b_1={@q1@}\)</span>, <span class="math inline">\(b_2=-{@q2@}\)</span>, <span class="math inline">\(b_3=0\)</span> we find</p>
<p><span class="math display">\[\begin{aligned}
\underline{a}\times \underline{b}&amp;= (({@p2@})(0)-(0)(-{@q2@}))\underline{i}-(({@p1@})(0)-(0)({@q1@}))\underline{j}+(({@p1@})(-{@q2@})-({@p2@})({@q1@}))\underline{k}\\
&amp;= {@ta@}\underline{k}\end{aligned}.\]</span></p><p><br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>p1: 2+rand(3)
p2: 2+rand(4)
q1: 2+rand(5)
q2: 2+rand(6)
ta: -p1*q2 -p2*q1</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(({@p1@}\underline{ i}+{@p2@}\underline{ j}) \times ({@q1@}\underline{ i}-{@q2@}\underline{ j}) = {@ta@}\underline{ k} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2035175511</deployedseed>
    <deployedseed>755940633</deployedseed>
    <deployedseed>121058625</deployedseed>
    <deployedseed>372506200</deployedseed>
    <deployedseed>2145459699</deployedseed>
    <deployedseed>1550917835</deployedseed>
    <deployedseed>809479233</deployedseed>
    <deployedseed>897124644</deployedseed>
    <deployedseed>1076608172</deployedseed>
    <deployedseed>19453885</deployedseed>
    <deployedseed>228579085</deployedseed>
    <deployedseed>2021917302</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
