<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>AlgMap-9.5</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Solve the following equation for {@v@}:
\[ {@b1*(cf1*v+k1)+b2*(cf2*v+k2)@} = 0.\]  <br>

\({@v@}\) = [[input:ans]] [[validation:ans]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[The objective is to get the equation into the form \({@v@} = p\) where \(p\) is a rational number. To obtain \({@v@}\) on one side of the equation it may be necessary to under-take several operations which must be performed on both sides of the equation. In this example: <br>
\[ {@b1*(cf1*v+k1)+b2*(cf2*v+k2)@} = 0.\]
We are first required to expand the brackets and collect like terms:
\[ {@b1*cf1*v+b1*k1+b2*cf2*v+b2*k2@} = 0. \]
Then move the constants to one side keeping the variables on the other:
\[ {@(b1*cf1+b2*cf2)*v@} = {@-b1*k1-b2*k2@}.\]
Then divide by the coefficient of \({@v@}\) (\({@b1*cf1+b2*cf2@}\)):
\[ {@v@} = {@ta@}. \]
And so we now have obtained a value for {@v@}.]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:rand([x,y,z,t]);
cf1:rand(3)+1;
cf2:cf1+rand(3)+1;
k1:-rand(10)-1;
k2:rand(10)+1;
b1:-rand(10)-1;
b2:rand(10)+1;
ta:(-b1*k1-b2*k2)/(b1*cf1+b2*cf2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ {@b1*(cf1*v+k1)+b2*(cf2*v+k2)@} = 0 \Rightarrow {@v@} = {@ta@}. \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>1-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1508306552</deployedseed>
    <deployedseed>984086795</deployedseed>
    <deployedseed>405859273</deployedseed>
    <deployedseed>1981537611</deployedseed>
    <deployedseed>814212812</deployedseed>
    <deployedseed>972311538</deployedseed>
    <deployedseed>973208210</deployedseed>
    <deployedseed>70984117</deployedseed>
    <deployedseed>1852284840</deployedseed>
    <deployedseed>1392717463</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>1-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
