<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.3.3.3 Green's Theorem in the Plane Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Verify Greens' theorem in the plane when applied to the integral </p>
<p> \(\qquad \displaystyle\oint_C \left\{ (5x + 2y - 7) \mathrm{d}x + (3x - 4y + 5) \mathrm{d}y \right\}, \) </p>
<p>where \(C\) represents the perimeter of the trapezium with vertices at \((0,0)\), \(({@aa@},0)\), \(({@bb@},{@cc@})\) and \(({@cc@},{@cc@})\).</p>

<ol class="HELM_parts">
<li><p>First, let \(\, P=5x+2y-7 \, \) and \(\, Q=3x-4y+5\, \) and find \(\, \displaystyle\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\)</p>
<p> \(\displaystyle\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y} =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li>

<li><p>Now find \(\displaystyle\int\!\!\int \left(\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\right)\, \mathrm{d}x \, \mathrm{d}y\, \) over the trapezium</p>
<p> \(\displaystyle\int\!\!\int \left(\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\right)\mathrm{d}x\mathrm{d}y =\)  [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>

<li><p>Now find \(\displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y)\) along each of the four sides of the trapezium, beginning with the line from \((0,0)\) to \(({@aa@},0)\), and then proceeding anti-clockwise.</p>
<ol class="HELM_parts">
<li><p> For the line from \((0,0)\) to \(({@aa@},0)\), \(\displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]] </p></li>

<li><p> For the line from \(({@aa@},0)\) to \(({@bb@},{@cc@})\), \(\displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) = \) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]] </p></li>

<li><p> For the line from \(({@bb@},{@cc@})\) to \(({@cc@},{@cc@})\), \(\displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) = \) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]] </p></li>

<li><p> For the line from \(({@cc@},{@cc@})\) to \((0,0)\), \(\displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) = \) [[input:ans6]] [[validation:ans6]] [[feedback:prt6]] </p></li>
</ol>
</li>
<li><p> Sum the integrals of each of the 4 sides to find the total integral  \(\displaystyle \oint_C (P\mathrm{d}x +Q\mathrm{d}y)\) </p>
<p>\(\displaystyle \oint_C (P\mathrm{d}x +Q\mathrm{d}y) = \)  [[input:ans7]] [[validation:ans7]] [[feedback:prt7]] </p></li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We are asked to verify Greens' theorem in the plane when applied to the integral </p>
<p> \(\qquad \displaystyle\oint_C \left\{ (5x + 2y - 7) \mathrm{d}x + (3x - 4y + 5) \mathrm{d}y \right\}, \) </p>
<p>where \(C\) represents the perimeter of the trapezium with vertices at \((0,0)\), \(({@aa@},0)\), \(({@bb@},{@cc@})\) and \(({@cc@},{@cc@})\).</p>

<ol class="HELM_parts">
<li><p>First, we have \(\, P=5x+2y-7 \, \) and \(\, Q=3x-4y+5\, \) and we are asked to find \(\, \displaystyle\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\)</p>
<p> \(Q=3x-4y+5 \implies \dfrac{\partial Q}{\partial x} = 3,\) and
<p> \(P = 5x + 2y-7 \implies \dfrac{\partial P}{\partial y} = 2,\) so
<p> \(\displaystyle\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y} = 3 - 2 = 1.\)</p></li>


<li><p>Now we are asked to find \(\displaystyle\int\!\!\int \left(\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\right)\, \mathrm{d}x \, \mathrm{d}y\, \) over the trapezium.</p>
<p>\(\begin{align}
\qquad \displaystyle\int\!\!\int \left(\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\right)\, \mathrm{d}x \, \mathrm{d}y &= \int\!\!\int 1 \, \mathrm{d}x \, \mathrm{d}y \\
&= \int\!\!\int 1 \, \mathrm{d}A \\
&= \text{Area of Trapezium} \end{align}\)</p>
<p> The area of a trapezium is given by \(\dfrac{1}{2} h(a+b)\) where \(h\) is the height and \(a\) and \(b\) are the lengths of the 2 parallel sides. Here, \(y\) ranges from \(0\) to \({@cc@}\), so \(h={@cc@}.\) At \(y=0, \, x\) ranges from \(0\) to \({@aa@}\) so \(a={@aa@}\), and at \(y={@cc@}, \, x\) ranges from \({@cc@}\) to \({@bb@}\) so \(b={@bb@}-{@cc@} = {@bb-cc@}.\)</p>
<p> So Area \(= \dfrac{1}{2} \times {@cc@}\left({@aa@}+{@bb-cc@}\right) = {@ta2@}\), so </p>
<p> \(\displaystyle\int\!\!\int \left(\dfrac{\partial Q}{\partial x}-\dfrac{\partial P}{\partial y}\right)\, \mathrm{d}x \, \mathrm{d}y\, = {@ta2@}.\)</p>
 </li>


<li><p>Now we find \(\displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y)\) along each of the four sides of the trapezium, beginning with the line from \((0,0)\) to \(({@aa@},0)\), and then proceeding anti-clockwise.</p>
<ol class="HELM_parts">
<li><p> For the line from \((0,0)\) to \(({@aa@},0)\): </p>
<p> Here \(y=0 \implies \mathrm{d}y=0,\) and we have \(x: 0 \to {@aa@}\), so</p>
<p> \(\begin{align}
\qquad \displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) &= \int P \, \mathrm{d}x \quad \text{since} \quad \mathrm{d}y=0, \\
&=\int_0^{@aa@} \left(5x+2y-7\right) \, \mathrm{d}x \\
&= \int_0^{@aa@} \left(5x-7\right) \, \mathrm{d}x \quad \text{since} \quad y=0, \\
&= \left[{@indef1@}\right]_0^{@aa@} \\
&= {@ta3@}.
\end{align}\)</p></li>


<li><p> For the line from \(({@aa@},0)\) to \(({@bb@},{@cc@})\):</p>
<p>This is the line with equation \(y={@dd*x@}-{@dd*aa@} \implies \frac{\mathrm{d}y}{\mathrm{d}x} = {@dd@} \implies \mathrm{d}y = {@dd@}\mathrm{d}x,\) and we have \(x: {@aa@} \to {@bb@}\), so</p>
 <p>\(\begin{align}
\qquad \displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) &= \left\{ (5x + 2y - 7) \mathrm{d}x + (3x - 4y + 5) \mathrm{d}y \right\}, \quad \text{and subbing in} \, y \, \text{and} \, \mathrm{d}y \, \text{simplifies to} \\
&= \int_{@aa@}^{@bb@} {@exp2@} \, \mathrm{d}x \\
&= \left[{@indef2@}\right]_{@aa@}^{@bb@} \\
&= {@ta4@}.
\end{align}\)  </p></li>


<li><p> For the line from \(({@bb@},{@cc@})\) to \(({@cc@},{@cc@})\):</p>
<p> Here \(y={@cc@} \implies \mathrm{d}y=0,\) and we have \(x: {@bb@} \to {@cc@}\), so </p>
 <p>\(\begin{align}
\qquad \displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) &= \int_{@bb@}^{@cc@} (5x+2y-7) \, \mathrm{d}x \\
&= \int_{@bb@}^{@cc@} ({@5*x+2*cc - 7@}) \, \mathrm{d}x \\
&= \left[{@indef3@}\right]_{@bb@}^{@cc@} \\
&= {@ta5@}.
\end{align}\)  </p></li>


<li><p> For the line from \(({@cc@},{@cc@})\) to \((0,0)\):</p>
<p> This is the line \(y=x\), \implies \(\mathrm{d}y=\mathrm{d}x,\) and we have \(x: {@cc@} \to 0\), so </p>
<p> \(\begin{align}
\qquad \displaystyle\int (P\mathrm{d}x +Q\mathrm{d}y) &= \left\{ (5x + 2y - 7) \mathrm{d}x + (3x - 4y + 5) \mathrm{d}x \right\} \\
&= \int_{@cc@}^0 (6x-2) \, \mathrm{d}x \\
&= \left[{@indef4@}\right]_{@cc@}^0 \\
&= {@ta6@}.
\end{align}\)</p></li>

</ol>
</li>
<li><p> We are now asked to sum the integrals of each of the 4 sides to find the total integral  \(\displaystyle \oint_C (P\mathrm{d}x +Q\mathrm{d}y)\) </p>
<p>\(\displaystyle \oint_C (P\mathrm{d}x +Q\mathrm{d}y) = {@ta3@} + {@ta4@} \, {@ta5@} \, {@ta6@} = {@ta7@}\) </p></li>

</ol>
<p> As expected, the sum of the integrals over each of the 4 sides in part (d) is exactly equal to our answer in part (b), as Green's Theorem tells us.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(3) + 3;
bb: rand(3) + 6;
cc: rand(2) + 1;
dd: cc/(bb-aa);
exp1: 5*x-7;
exp2: (5+5*dd - 4*dd^2)*x - 2*aa*dd + 4*aa*dd^2 - 7 + 5*dd;
exp3: 5*x + (2*cc - 7);
exp4: 6*x-2;
indef1: int(exp1,x);
indef2: int(exp2,x);
indef3: int(exp3,x);
indef4: int(exp4,x);
ta1: 1;
ta2: cc*(aa+bb-cc)/2;
ta3: int(exp1,x,0,aa);
ta4: int(exp2,x,aa,bb);
ta5: int(exp3,x,bb,cc);
ta6: int(exp4,x,cc,0);
ta7: ta3+ta4+ta5+ta6;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>Vertices: \((0,0)\), \(({@aa@},0)\), \(({@bb@},{@cc@})\) and \(({@cc@},{@cc@})\), and Green's Theorem: {@ta2@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>797897301</deployedseed>
    <deployedseed>2029815442</deployedseed>
    <deployedseed>162551209</deployedseed>
    <deployedseed>2118706594</deployedseed>
    <deployedseed>1541829231</deployedseed>
    <deployedseed>708053257</deployedseed>
    <deployedseed>150616740</deployedseed>
    <deployedseed>1885869222</deployedseed>
    <deployedseed>76600640</deployedseed>
    <deployedseed>1966123029</deployedseed>
    <deployedseed>1651905504</deployedseed>
    <deployedseed>407822518</deployedseed>
    <deployedseed>2072406023</deployedseed>
    <deployedseed>894463778</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4+1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5+1</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6+1</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt7-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
