<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.3.1.3 Stokes' Theorem Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Using cylindrical polar coordinates with \(z=0\), verify Stokes' theorem for the vector field \(\underline{F} = \rho \underline{\hat{\rho}}+ \rho \cos \left(\displaystyle\frac{\pi \rho}{2}\right) \underline{\hat{\phi}}\) and the semi-circle \(\rho \le 1\), \(-\dfrac{\pi}{2} \le \phi \le \dfrac{\pi}{2}\).</p>
<ol class="HELM_parts">
<li> <p> First find the \(z\)- component of the curl, \(({\underline \nabla} \times \underline{F}) \cdot \hat{z},\) and then select the correct answer from the options below.</p>
<p>[Hint: remember the formula for the curl of a vector field in polar coordinates, and how to find the component in one direction.]</p>
<p> 1. \(({\underline \nabla} \times \underline{F}) \cdot \hat{z} = 2 \rho \cos \left(\dfrac{\pi}{2} \rho\right) - \dfrac{\pi}{2} \rho^2 \sin \left(\dfrac{\pi}{2} \rho \right) \)</p>
<p> 2.  \(({\underline \nabla} \times \underline{F}) \cdot \hat{z} = 2 \cos \left(\dfrac{\pi}{2} \rho\right) - \dfrac{\pi}{2} \rho \sin \left(\dfrac{\pi}{2} \rho\right) \)</p>
<p> 3.  \(({\underline \nabla} \times \underline{F}) \cdot \hat{z} = 2 \cos \left(\dfrac{\pi}{2} \rho\right)\)</p>
<p> \(({\underline \nabla} \times \underline{F}) \cdot \hat{z} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

<li> <p> Now evaluate the double integral \(\displaystyle \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \underline{\mathrm{d}S} = \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \hat{z} \, \mathrm{d}A\), remembering the transformation of \(\mathrm{d}A\) into cylindrical polar coordinates.</p>
<p>\(\displaystyle \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \hat{z} \, \mathrm{d}A =\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

<li> <p> Now find the integral \(\displaystyle\oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) by splitting the boundary \(C\) into 3 parts: </p> 
<ol class="HELM_parts">
<li> <p> First find \(\underline{F} \cdot \underline{\mathrm{d}r},\) and then select the correct answer from the options below.</p>
<p>[Hint: you may find it useful that \(\underline{\mathrm{d}r} = \hat{\underline{\rho}}\, \mathrm{d}\rho + \hat{\underline{\phi}}\, (\rho\, \mathrm{d}\phi)\)]
<p> 1. \(\underline{F} \cdot \underline{\mathrm{d}r} = \displaystyle \rho \, \mathrm{d}\rho + \rho^3 \cos(\dfrac{\pi}{2} \rho)\, \mathrm{d}\phi\) </p>
<p> 2. \(\underline{F} \cdot \underline{\mathrm{d}r} = \displaystyle \rho \, \mathrm{d}\rho + \rho \cos(\dfrac{\pi}{2} \rho)\, \mathrm{d}\phi\) </p>
<p> 3. \(\underline{F} \cdot \underline{\mathrm{d}r} = \displaystyle \rho \, \mathrm{d}\rho + \rho^2 \cos(\dfrac{\pi}{2} \rho) \, \mathrm{d}\phi\) </p>
<p>  \(\underline{F} \cdot \underline{\mathrm{d}r} = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
</li>

<li> <p> The contribution to \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) along the radial segment \(\phi = -\dfrac{\pi}{2}\)  is [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>
<p> [Note: think about the semi-circle, which part of the boundary is represented by  \(\phi = -\dfrac{\pi}{2}\)?] </p> </li>

<li> <p>  The contribution to  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) along the arc \(\rho = 1\) is [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</p><p> [Note: think about the semi-circle, which part of the boundary is represented by  \(\rho = 1\)?] </p></li>

<li> <p> The contribution to  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) along the radial segment \(\phi = \dfrac{\pi}{2}\) is [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</p><p> [Note: think about the semi-circle, which part of the boundary is represented by  \(\phi = \dfrac{\pi}{2}\)?] </p></li>
</ol>
<li> <p> Hence find the total line integral  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) </p>
<p>  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r} = \) [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</p></li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> We are asked to verify Stokes' theorem for the vector field \(\underline{F} = \rho \underline{\hat{\rho}}+ \rho \cos \left(\displaystyle\frac{\pi \rho}{2}\right) \underline{\hat{\phi}}\) and the semi-circle \(\rho \le 1\), \(-\dfrac{\pi}{2} \le \phi \le \dfrac{\pi}{2},\) using cylindrical polar coordinates with \(z=0\).</p>
<ol class="HELM_parts">
<li> <p> We are asked to first find the \(z-\) component of the curl, \(({\underline \nabla} \times \underline{F}) \cdot \hat{z}.\)</p>
<p> We know that in cylindrical coordinates we have </p>
<p> curl \(\displaystyle \underline{F} = {\underline \nabla} \times \displaystyle \underline{F} = \frac{1}{\rho} \left|\begin{array}{ccc} \underline{\hat{\rho}} &amp; \rho \underline{\hat{\phi}} &amp; \underline{\hat{z}} \\ \\ \displaystyle \frac{\partial}{\partial \rho} &amp; \displaystyle \frac{\partial}{\partial \phi} &amp; \displaystyle \frac{\partial}{\partial z} \\ \\ F_{\rho} &amp; \rho F_{\phi} &amp; F_z \end{array} \right|\).</p>

<p> Expanding the determinant to find the \(z\)- component, we get </p>
<p> \(\qquad ({\underline \nabla} \times \underline{F}) \cdot \hat{z} = \dfrac{1}{\rho}\left(\dfrac{\partial}{\partial \rho}\!\big(\rho F_{\phi}\big) \;-\; \dfrac{\partial F_{\rho}}{\partial \phi}\right)\)</p>

<p> So here, with \(F_{\rho}=\rho, \, F_{\phi} = \rho \cos \left(\dfrac{\pi}{2} \rho \right), \,  F_z=0,\) the \(z\)- component of the curl becomes </p>
<p> \(\begin{align}
\qquad ({\underline \nabla} \times \underline{F}) \cdot \hat{z} &= \dfrac{1}{\rho}\dfrac{\partial}{\partial \rho}\!\left(\rho \rho \cos \left(\dfrac{\pi}{2} \rho\right)\right) \;-\; \dfrac{1}{\rho}\dfrac{\partial \rho}{\partial \phi} \\
&= \dfrac{1}{\rho} \dfrac{\partial}{\partial \rho} \left(\rho^2 \cos \left(\dfrac{\pi}{2} \rho \right)\right), \quad \text{since} \, \rho \, \text{is independent of} \, \phi \\
&= 2 \cos \left(\dfrac{\pi}{2} \rho\right) - \dfrac{\pi}{2} \rho \sin \left(\dfrac{\pi}{2} \rho\right), \quad \text{by using the product rule.} \end{align}\)</p></li>


<li> <p> We are now asked to evaluate the double integral \(\displaystyle \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \underline{\mathrm{d}S} = \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \hat{z} \, \mathrm{d}A\), remembering the transformation of \(\mathrm{d}A\) into cylindrical polar coordinates. We know that \(\mathrm{d}A\) becomes \(\rho \, \mathrm{d}\rho \, \mathrm{d}\phi,\) and so we have </p>
<p>\(\begin{align}
\qquad \displaystyle \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \underline{\mathrm{d}S} &= \int \int_S ({\underline \nabla} \times \underline{F}) \cdot \hat{z} \, \mathrm{d}A \\
&= \int_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \int_0^1 \left(2 \cos \left(\dfrac{\pi}{2} \rho\right) - \dfrac{\pi}{2} \rho \sin \left(\dfrac{\pi}{2} \rho\right)\right) \rho \, \mathrm{d}\rho \, \mathrm{d}\phi \\
&= \int_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \int_0^1 \left(2 \rho \cos \left(\dfrac{\pi}{2} \rho\right) - \dfrac{\pi}{2} \rho^2 \sin \left(\dfrac{\pi}{2} \rho\right)\right) \, \mathrm{d}\rho \, \mathrm{d}\phi \\
&= \int_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \left[\rho^2 \cos \left(\dfrac{\pi}{2} \rho\right)\right]_0^1 \, \mathrm{d}\phi \\
&=  \int_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \cos \left(\dfrac{\pi}{2} \right) \, \mathrm{d}\phi \\
&= \left[\phi \cos \left(\dfrac{\pi}{2} \right) \right]_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \\
&= \pi \cos \left(\dfrac{\pi}{2} \right) \\
&= 0. \end{align}\) </p></li>

<li> <p> We are now asked to find the integral \(\displaystyle\oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) by splitting the boundary \(C\) into 3 parts: </p> 
<ol class="HELM_parts">
<li> <p> We are first asked to find \(\underline{F} \cdot \underline{\mathrm{d}r}\)</p>
<p> We know that \(\underline{\mathrm{d}r} = \hat{\underline{\rho}}\, \mathrm{d}\rho + \hat{\underline{\phi}}\, (\rho\, \mathrm{d}\phi)\), and \(\underline{F} = \rho \underline{\hat{\rho}}+ \rho \cos \left(\displaystyle\frac{\pi \rho}{2}\right) \underline{\hat{\phi}}\), so </p>
<p>  \(\qquad \underline{F} \cdot \underline{\mathrm{d}r} = \rho \, \mathrm{d}\rho + \rho^2 \cos \left(\dfrac{\pi}{2} \rho \right) \mathrm{d}\phi.\) </p>
</li>

<li> <p> We are now asked to find the contribution to \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) along the radial segment \(\phi = -\dfrac{\pi}{2}\) </p>
<p> Along the radial segment \(\phi = -\dfrac{\pi}{2} \implies \mathrm{d}\phi = 0\), and with \(\rho: 0 \to 1\), the contribution becomes </p>
<p> \(\qquad \displaystyle \int_0^1 \rho \mathrm{d}\rho = \left[\dfrac{\rho^2}{2}\right]_0^1 = \dfrac{1}{2}. \)
</li>

<li> <p>  We are now asked to find the contribution to  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) along the arc \(\rho = 1\)</p>
<p> Along the arc \(\rho = 1 \implies \mathrm{d}\rho = 0\), and with \(\phi: -\dfrac{\pi}{2} \to \dfrac{\pi}{2}\), the contribution becomes </p>
<p> \(\qquad \displaystyle \int_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \rho^2 \cos \left(\dfrac{\pi}{2} \rho \right) \mathrm{d}\phi = \int_{-\frac{\pi}{2}}^{\frac{\pi}{2}} \cos \left(\dfrac{\pi}{2} \right)   \mathrm{d}\phi = \pi \cos \left(\dfrac{\pi}{2}\right) = 0. \)</li>

<li> <p> We are now asked to find the contribution to  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) along the radial segment \(\phi = \dfrac{\pi}{2}\) </p>
<p> Along the radial segment \(\phi = \dfrac{\pi}{2} \implies \mathrm{d}\phi = 0\), and with \(\rho: 1 \to 0\), the contribution becomes </p>
<p> \(\qquad \displaystyle \int_1^0 \rho \mathrm{d}\rho = \left[\dfrac{\rho^2}{2}\right]_1^0 = -\dfrac{1}{2}. \)
</li></li>
</ol>

<li> <p> We are asked to hence find the total line integral  \(\displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r}\) </p>
<p> The total line integral is equal to the sum of each of the 3 contributions along the boundary of the semi-circle, so </p>
<p>  \(\qquad \displaystyle \oint_C \underline{F} \cdot \underline{\mathrm{d}r} = \dfrac{1}{2} + 0 -\dfrac{1}{2} = 0\)</p>
<p> This is the same as our calculated surface integral in part (b), as Stokes' theorem tells us.</p> </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ta1: [["1.",false],["2.",true],["3.",false]];
ta2: 0;
ta3: [["1.",false],["2.",false],["3.",true]];
ta4: 1/2;
ta5: 0;
ta6: -1/2;
ta7: 0;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>Both integrals = {@ta2@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta1)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans3]</sans>
        <tans>mcq_correct(ta3)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_incorrect(ta1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_incorrect(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4+1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5+1</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6+1</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt7-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
