<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-13-parts-7</text>
    </name>
    <questiontext format="html">
      <text>Integrate {@p@} with respect to \(t\) using the method of integration by parts.

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[To solve this problem we will have to use the method of integration by parts 
<hint>calc_int_methods_parts</hint> 
two times as follows:

Let \(u = {@pp1@}\) and \(dv = {@pp2@}dt\) so that
\(du = {@diff(pp1,t)@}dt\) and \(v = {@int(pp2,t)@}\).

Then,
\[\begin{array}{rcl}\displaystyle
 \int {@p@} dt &\displaystyle =&\displaystyle  \int u dt \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@pp1@}{@int(pp2,t)@} - \int {@int(pp2,t)@} {@diff(pp1,t)@} dt + c \\
   &\displaystyle =&\displaystyle  {@pp1*int(pp2,t)@} - \int {@int(pp2,t)@} {@diff(pp1,t)@} dt + c
 \end{array}\]

Now, define \(u = {@diff(pp1,t)@}\) and \(dv = {@int(pp2,t)@}dt\) 
so that  \(du = {@diff(pp1,t,2)@}dt\) and \(v = {@int(int(pp2,t),t)@}\).

Then,
\[\begin{array}{rcl}\displaystyle
\int {@int(pp2,t)@}{@diff(pp1,t)@}dt &\displaystyle =&\displaystyle  \int u dt \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@diff(pp1,t)@}{@int(int(pp2,t),t)@} - \int {@int(int(pp2,t),t)@} [{@diff(pp1,t,2)@}dt] \\
   &\displaystyle =&\displaystyle  {@diff(pp1,t)*int(int(pp2,t),t)@} + \int {@-int(int(pp2,t),t)*diff(pp1,t,2)@}dt.
 \end{array}\]

And so
\[\begin{array}{rcl}\displaystyle
\int {@p@} dt &\displaystyle =&\displaystyle  {@pp1*int(pp2,t)@} - \left[{@diff(pp1,t)*int(int(pp2,t),t)@} + \int {@-int(int(pp2,t),t)*diff(pp1,t,2)@} d{@v@}\right] + c \\
   &\displaystyle =&\displaystyle  {@pp1*int(pp2,t)@} - {@diff(pp1,t)*int(int(pp2,t),t)@} - \int {@-int(int(pp2,t),t)*diff(pp1,t,2)@} dt + c.
 \end{array}\]
Hence,
\[\begin{array}{rcl}\displaystyle
2\int {@p@} dt &\displaystyle =&\displaystyle  {@pp1*int(pp2,t)-diff(pp1,t)*int(int(pp2,t),t)@} + c \\
\int {@p@} dt &\displaystyle =&\displaystyle  {@ta@} + k.
\end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:t;
pp2:exp(t);
pp1:sin(t);
p:pp1*pp2;
ta:int(p,t);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int {@p@} dt = {@ta@}+c\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta+c</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta+c</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-0-F </expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
