<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-13-parts-3.c</text>
    </name>
    <questiontext format="html">
      <text>Integrate {@p@} with respect to {@v@}.

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[If \(u\) and \(v\) are both functions of {@v@} then \[\int u dv = uv - \int v du.\]

Let \(u = {@pp2@}\) and \(dv = {@pp1@} d{@v@}\). 

Then \(du = {@diff(pp2,v)@} d{@v@}\) and \(v = {@int(pp1,v)@}\).

So, 
\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int u dv \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@pp2@}{@int(pp1,v)@} - \int {@int(pp1,v)@}{@diff(pp2,v)@}d{@v@} \\
   &\displaystyle =&\displaystyle  {@pp2*int(pp1,v)@} - \int {@int(pp1,v)*diff(pp2,v)@}d{@v@}
 \end{array}\] 

To calculate \[\int {@int(pp1,v)*diff(pp2,v)@}d{@v@}\] we will have to use parts again, by defining 

\(u = {@ppp@}\) and \(dv = {@2*pp1/(rdm+1)@}d{@v@}\) so that 

\(du = {@diff(ppp,v)@}d{@v@}\) and \(v = {@int(2*pp1/(rdm+1),v)@}\).

So,

\[\begin{array}{rcl}\displaystyle 
 \int {@int(pp1,v)*diff(pp2,v)@}d{@v@} &\displaystyle =&\displaystyle  \int u dv \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@ppp@}{@int(2*pp1/(rdm+1),v)@} - \int {@int(2*pp1/(rdm+1),v)@}{@diff(ppp,v)@}d{@v@} \\
   &\displaystyle =&\displaystyle  {@ppp*int(2*pp1/(rdm+1),v)@} - \int {@int(2*pp1/(rdm+1),v)*diff(ppp,v)@}d{@v@} \\
   &\displaystyle =&\displaystyle  {@ppp*int(2*pp1/(rdm+1),v)-int((int(2*pp1/(rdm+1),v)*diff(ppp,v)),v)@} + c
 \end{array}\] 

Hence,

\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =&\displaystyle  {@pp2*int(pp1,v)@} - \int {@int(pp1,v)*diff(pp2,v)@}d{@v@} \\
   &\displaystyle =&\displaystyle  {@pp2*int(pp1,v)@} - \left({@ppp*int(2*pp1/(rdm+1),v)-int((int(2*pp1/(rdm+1),v)*diff(ppp,v)),v)+c@}\right) \\
   &\displaystyle =&\displaystyle  {@pp2*int(pp1,v) - (ppp*int(2*pp1/(rdm+1),v)-int((int(2*pp1/(rdm+1),v)*diff(ppp,v)),v)+c)@}
 \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2019121800</text>
    </stackversion>
    <questionvariables>
      <text>v:t;
rdm:1+rand(4);
pp1:v^rdm;
ppp:log(v);
pp2:ppp^2;
p:pp1*pp2;
ta1:int(p,v)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int {@p@} d{@v@} = {@int(p,v)@}+c\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>360491604</deployedseed>
    <deployedseed>525837741</deployedseed>
    <deployedseed>1196370803</deployedseed>
    <deployedseed>1951725100</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-0-F </expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
