<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-13-parts-4.a</text>
    </name>
    <questiontext format="html">
      <text>Evaluate the following definite integral:
\[ \int^{@rdm2@}_{@rdm3@} {@p@} d{@v@}. \]

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[If \[\int f(x) dx = F(x)\] then 
\[\int^a_b f(x) dx = F(a) - F(b).\]

Using integration by parts, 
\[\int u dv = uv - \int v du,\] 
by defining \(u = {@pp1@}\) and \(dv = {@pp2@} d{@v@}\) so that 

\(du = {@diff(pp1,v)@} d{@v@}\) and \(v = {@int(pp2,v)@}\).

Then, 

\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int u dv \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@pp1@}{@int(pp2,v)@} - \int {@int(pp2,v)@}{@diff(pp1,v)@} d{@v@}
 \end{array}\] 

To calculate \[\int {@int(pp2,v)*diff(pp1,v)@}d{@v@}\] we will have to use parts again, by defining 

\(u = {@diff(pp1,v)@}\) and \(dv = {@int(pp2,v)@}d{@v@}\) so that 

\(du = {@diff(pp1,v,2)@}d{@v@}\) and \(v = {@int(int(pp2,v),v)@}\).

So,

\[\begin{array}{rcl}\displaystyle 
 \int {@int(pp2,v)*diff(pp1,v)@}d{@v@} &\displaystyle =&\displaystyle  \int u dv \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@diff(pp1,v)@}[{@int(int(pp2,v),v)@}] - \int [{@int(int(pp2,v),v)@}]{@diff(pp1,v,2)@}d{@v@} \\
   &\displaystyle =&\displaystyle  {@diff(pp1,v)*int(int(pp2,v),v)@} - \int {@int(int(pp2,v),v)*diff(pp1,v,2)@}d{@v@} \\
   &\displaystyle =&\displaystyle  {@diff(pp1,v)*int(int(pp2,v),v)-int(int(int(pp2,v),v),v)*diff(pp1,v,2)+c@}
 \end{array}\] 

Hence,

\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =&\displaystyle  {@pp1*int(pp2,v)@} - ({@diff(pp1,v)*int(int(pp2,v),v)-int(int(int(pp2,v),v),v)*diff(pp1,v,2)+c@}) \\
   &\displaystyle =&\displaystyle  {@pp1*int(pp2,v)-(diff(pp1,v)*int(int(pp2,v),v)-int(int(int(pp2,v),v),v)*diff(pp1,v,2)+c)@} \\
   &\displaystyle =&\displaystyle  {@q-c@}
 \end{array}\] 

Lastly, we evaluate this expression at the limits to give

\[ \left[{@q@}\right]_{@rdm3@}^{@rdm2@} = {@int(p,v,rdm2,rdm3)@}.\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm1:1;
rdm2:rand([pi/4,pi/2,3*pi/2,2*pi]);
rdm3:-rand([0,pi/4,pi/2,3*pi/2,2*pi]);
pp1:v^2;
pp2:cos(rdm1*v);
p:pp1*pp2;
q:int(p,v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int^{@rdm3@}_{@rdm2@} {@p@}d{@v@} = {@int(p,v,rdm2,rdm3)@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v,rdm3,rdm2)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>int(p,v,rdm3,rdm2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>diff(ans1,v)</sans>
        <tans>p</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>This is a definite integration problem, so your answer should be a number.  This is the correct integral, but you have not evaluated it.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1464428665</deployedseed>
    <deployedseed>1135219946</deployedseed>
    <deployedseed>107302024</deployedseed>
    <deployedseed>1876476599</deployedseed>
    <deployedseed>660358305</deployedseed>
    <deployedseed>1185021824</deployedseed>
    <deployedseed>818622199</deployedseed>
    <deployedseed>1402852878</deployedseed>
    <deployedseed>929111151</deployedseed>
    <deployedseed>422673398</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v,rdm3,rdm2),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F </expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
