<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-13-parts-2.a</text>
    </name>
    <questiontext format="html">
      <text>Evaluate the following definite integral:
\[ \int^{@rdm3@}_{@rdm4@} {@p@} d{@v@}. \]

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[If \[\int f(x) dx = F(x)\] then \[\int^a_b f(x) dx = F(a) -
 F(b).\]

So, we need to calculate \[\int {@p@} d{@v@}.\]

Using integration by parts, that is if \(u\) and \(v\) are both functions of {@v@} then \[\int u dv = uv - \int v du,\] by defining \(u = {@ppp1@}\) and \(dv = {@ppp2@} d{@v@}\) so that 

\(du = {@diff(ppp1,v)@} d{@v@}\) and \(v = {@int(ppp2,v)@}\).

Then, 

\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int u dv \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  {@ppp1@}\left[{@int(ppp2,v)@}\right] - \int \left[{@int(ppp2,v)@}\right]{@diff(ppp1,v)@} d{@v@} \\
   &\displaystyle =&\displaystyle  {@ppp1*int(ppp2,v)@} - \int {@int(ppp2,v)*diff(ppp1,v)@} d{@v@} \\
   &\displaystyle =&\displaystyle  {@ppp1*int(ppp2,v)-int((int(ppp2,v)*diff(ppp1,v)),v)@}
 \end{array}\] 

So, let \(F({@v@}) = {@q@}\)

Then, 

\[\begin{array}{rcl}\displaystyle 
 \int^{@rdm3@}_{@rdm4@} {@p@} d{@v@} &\displaystyle =&\displaystyle  F({@rdm3@}) - F({@rdm4@}) \\ 
   &\displaystyle =&\displaystyle  {@subst(rdm3,v,q)@} - {@subst(rdm4,v,q)@} \\
   &\displaystyle =&\displaystyle  {@int(p,v,rdm4,rdm3)@}
 \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm2:-2-rand(2);
rdm3:rand([pi/(4*rdm2),pi/(2*rdm2),pi/(2*rdm2),pi/(rdm2),2*pi/(4*rdm2)]);
rdm4:0;
ppp1:v;
ppp2:cos(rdm2*v);
p:ppp1*ppp2;
q:int(p,v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int^{@rdm3@}_{@rdm4@} {@p@}d{@v@} = {@int(p,v,rdm4,rdm3)@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v,rdm4,rdm3)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>int(p,v,rdm4,rdm3)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>diff(ans1,v)</sans>
        <tans>p</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>Remember, this is a definite integration problem so your answer should be a number.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2070934317</deployedseed>
    <deployedseed>65991505</deployedseed>
    <deployedseed>1298288302</deployedseed>
    <deployedseed>1853963510</deployedseed>
    <deployedseed>1266628242</deployedseed>
    <deployedseed>1127084853</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v,rdm4,rdm3),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F </expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
