<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>HELM 31.2 E4 Simpson's rule, intervals</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**31.2 E4 Composite Simpson's rule, precision**

The function \\\(f\\\) is known to have a fourth derivative with the property that
\\\[\|f^{(iv)}(x)\| < {@ddddf@}\\\]
for \\\(x\\\) between {@a@} and {@b@}. Determine how many subintervals are required so that the composite Simpson’s rule used to approximate

\\\( \displaystyle {@'integrate(f(x),x,a,b)@} \\\)

incurs an error that is less than {@(stackfltfmt:"~f",err)@}.

\\\(N=\\\) [[input:S_iv]] [[validation:S_iv]] 
[[feedback:iv]]]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024012900</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a: rand_with_step(-5,5,1);
b: a+rand_with_step(1,5,1);
ddddf: rand_with_step(6,12,1);
err: rand([0.01, 0.001, 0.0001]);

/*[a,b,ddddf,err]: [-1,5,6,0.001];*/

h: ev(h,solve(err=ddddf*(b-a)*h^4/180,h));
ivr: ceiling(abs((b-a)/h));
if oddp(ivr) then iv:ivr+1 else iv:ivr;

stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_value_EN.mac");]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@[a,b,ddddf,err]@} {@ivr@} {@iv@}</text>
    </questionnote>
    <questiondescription format="markdown">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>S_iv</name>
      <type>numerical</type>
      <tans>iv</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>integer</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>iv</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>correct value?</description>
        <answertest>AlgEquiv</answertest>
        <sans>S_iv</sans>
        <tans>iv</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>iv-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>iv-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_unit(S_iv,iv , 0.005)@}</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>is N even?</description>
        <answertest>AlgEquiv</answertest>
        <sans>evenp(S_iv)</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>0</truenextnode>
        <trueanswernote>iv-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>iv-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>\(N\) must be even for Simpson's rule.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1666272798</deployedseed>
    <deployedseed>36725630</deployedseed>
    <deployedseed>1006023209</deployedseed>
    <deployedseed>1131128933</deployedseed>
    <deployedseed>186153235</deployedseed>
    <deployedseed>859648525</deployedseed>
    <deployedseed>949516628</deployedseed>
    <deployedseed>1961907446</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_iv</name>
        <value>iv</value>
      </testinput>
      <expected>
        <name>iv</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>iv-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
