<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>HELM 31.2 E2 Composite trapezium rule, intervals</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**31.2 E2 Composite trapezium rule**

The function \\\(f\\\) is known to have a second derivative with the property that
\\\[\|f′′(x)\| < {@ddf@}\\\]
for \\\(x\\\) between {@a@} and {@b@}. Using the error bound given in HELM 31.2 determine how many
subintervals are required so that the composite trapezium rule used to approximate

\\\( \displaystyle {@'integrate(f(x),x,a,b)@} \\\)

can be guaranteed to have an error in it that is less than {@(stackfltfmt:"~f",err)@}.

\\\(N=\\\) [[input:S_iv]] [[validation:S_iv]] 
[[feedback:iv]]]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024072400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a: rand_with_step(1,5,1);
b: a+rand_with_step(1,5,1);
ddf: rand_with_step(6,12,1);
err: rand([0.01, 0.001, 0.0001]);

/* h: ev(h,solve(err=ddf*(b-a)*h^2/12,h)); iv: ceiling(abs(1/h));*/

iv:  ceiling(sqrt( ddf*(b-a)/err)*(b-a)/(2*sqrt(3)));



stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_value.mac");]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@[a,b,ddf,err]@} {@iv@}</text>
    </questionnote>
    <questiondescription format="markdown">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>S_iv</name>
      <type>numerical</type>
      <tans>iv</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>integer</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>iv</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>S_iv</sans>
        <tans>iv</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>iv-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>iv-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_unit(S_iv,iv , 0.005)@}</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1354993404</deployedseed>
    <deployedseed>274379906</deployedseed>
    <deployedseed>708883317</deployedseed>
    <deployedseed>1666272798</deployedseed>
    <deployedseed>391962638</deployedseed>
    <deployedseed>36725630</deployedseed>
    <deployedseed>1006023209</deployedseed>
    <deployedseed>1131128933</deployedseed>
    <deployedseed>186153235</deployedseed>
    <deployedseed>688591538</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_iv</name>
        <value>iv</value>
      </testinput>
      <expected>
        <name>iv</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>iv-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
