<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>19.2.9.6 Temperature ODE</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The temperature \(\theta\) (measured in degrees) of a body immersed in an atmosphere of varying
    temperature is given by</p>
\[\frac{\mathrm{d}\theta}{\mathrm{d}t} + {@l@} \theta = {@g@}.\]
<p>Find the temperature at time \(t\) if \(\theta = {@theta0@}^{\circ} \mathrm{C}\) when \(t = 0\).</p>
<p>\(\theta(t) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>This is a linear differential equation that is already written in standard form.
Note that here \(f(t) = {@l@}\) and \(g(t) = {@g@}\).</p>
<p>We use the key point to find the integrating factor:</p>
\[\text{I.F.} = \exp \left[ \int f(t) \,\mathrm{d}t \right] = {@IF@}\]
<p>Next we transform the equation into exact form by multiplying both sides with the
    integrating factor. It becomes:
</p>
\[{@IF@} \cdot \frac{\mathrm{d}\theta}{\mathrm{d}t} + {@l * IF@} \cdot \theta = {@IF@} \cdot ({@g@})\]
<p>Note that the left-hand side is the derivative of \({@IF@} \cdot \theta\) with respect to \(t\).
    Our equation then becomes:
</p>
\[\frac{\mathrm{d}}{\mathrm{d}t} \left( {@IF@} \cdot \theta \right) = {@IF * g@}\]
<p>Therefore, we can find an expression for \({@IF@} \cdot \theta\) by integrating the right-hand side.
    This can be done using integration by parts, which you can review in previous workbooks.
    The result is:
</p>
\[{@IF@} \cdot \theta = {@I@}\]
<p>Which then gives the following general solution:</p>
\[\theta = {@gsol@}\]
<p>We can now solve for \(C\) and use the given boundary conditions:</p>
\[C = {@Cexp@}\]
\[\theta(0) = {@theta0@} \implies C = {@C0@}\]
<p>This gives the particular solution we seek:</p>
\[\theta(t) = {@ta@}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomisation */
l: rand([0.1, 0.2, 0.5]);
a: rand_with_step(3, 7, 2);
b: float(a / 2);
theta0: rand_with_step(30, 80, 10);

g: a - b * t;
IF: e^(int(l, t));
I: expand(int(IF * g, t)) + C;
gsol: expand(I / IF);
Cexp: rhs(solve(theta = gsol, C)[1]);
C0: float(ev(Cexp, t=0, theta=theta0));
ta: ev(gsol, C=C0);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\theta' + {@l@} \theta = {@g@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>207087991</deployedseed>
    <deployedseed>2123924174</deployedseed>
    <deployedseed>560644367</deployedseed>
    <deployedseed>1205434721</deployedseed>
    <deployedseed>839656071</deployedseed>
    <deployedseed>2125588050</deployedseed>
    <deployedseed>1999063934</deployedseed>
    <deployedseed>78417243</deployedseed>
    <deployedseed>1936429163</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
