<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>19.2.9.3 First Order DE Ex2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find the solution of the equation \(x\dfrac{\mathrm{d}y}{\mathrm{d}x} - y = {@aa*x@}\) subject to the condition \(y(1) = {@bb@}\).</p>
<p>\(y =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the equation \(x\dfrac{\mathrm{d}y}{\mathrm{d}x} - y = {@aa*x@}\).</p>

<p>Divide through by \(x\) to obtain the standard form of the equation \[\dfrac{\mathrm{d}y}{\mathrm{d}x} - \dfrac{y}{x} = {@aa@}.\] The coefficient in front of \(y\) is \(-\dfrac{1}{x}\). So the integrating factor is \[e^{\int -\frac{1}{x}\ \mathrm{d}x}
    = e^{-\log(x)} = \dfrac{1}{x}.\]</p>

<p>Multiply the standard form of the equation by the integrating factor to get \[ \dfrac{1}{x} \dfrac{\mathrm{d}y}{\mathrm{d}x} - \dfrac{1}{x^2}y = \dfrac{{@aa@}}{x}.\] Recognise that this is now exact and construct an integral \[\frac{\mathrm{d}}{\mathrm{d}x}\left(\frac{1}{x}y\right)
    = \dfrac{{@aa@}}{x} \iff \frac{1}{x}y = \int \dfrac{{@aa@}}{x}\ \mathrm{d}x.\]</p>

<p>Integrating gives \[\frac{1}{x}y = {@aa*log(x)@} + C\] which implies \[y = {@aa@}x\log(x) + Cx.\]</p>

<p>We are given the condition \(y(1)={@bb@}\). Substituting \(x=1\) and \(y={@bb@}\) yields \(C={@bb@}\). Therefore we have \[ y = {@aa*x@}\log(x) + {@bb@}x.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>aa:1 + rand(5);
bb: 1 + rand(5);

ta1: (aa*x)*log(x) + bb*x</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(x\dfrac{\mathrm{d}y}{\mathrm{d}x} - y = {@aa*x@}, y(1) = {@bb@}, \qquad y = {@ta1@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1882527242</deployedseed>
    <deployedseed>134489771</deployedseed>
    <deployedseed>478082984</deployedseed>
    <deployedseed>1833475486</deployedseed>
    <deployedseed>1862970062</deployedseed>
    <deployedseed>1787170587</deployedseed>
    <deployedseed>1253648548</deployedseed>
    <deployedseed>1803285922</deployedseed>
    <deployedseed>556806537</deployedseed>
    <deployedseed>1897741298</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(2*ta1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
