<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>19.2.2.8 Separation of Variables Ex6</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The equation \(iR + L\dfrac{\mathrm{d}i}{\mathrm{d}t} = E\) where \(R\), \(L\) and \(E\) are constants arisis in electrical circuit theory. this equation can be solved by separation of variables. Find the particular solution which satisfies the condition \(i(0) = {@aa@}\).</p>
<p>\(i =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the equation \(iR + L\dfrac{\mathrm{d}i}{\mathrm{d}t} = E\). This is equivalent to \(L\dfrac{\mathrm{d}i}{\mathrm{d}t} = E - iR\). We can then turn both sides of the equation into integrals \[ \int \dfrac{1}{E-iR} \ \mathrm{d}i = \int \dfrac{1}{L} \ \mathrm{d}t. \]</p>
<p>Performing the integration and rearranging we obtain 
  \[ \begin{aligned} -\dfrac{1}{R}\log(E-iR) &amp;= \frac{1}{L}t + C, \qquad \text{for some constant} \; C,\\
  \implies \log(E-iR) &amp;= -\frac{R}{L}t - RC, \\
  \implies E - iR &amp;= e^{-\frac{R}{L}t-RC} = Ae^{-\frac{R}{L}t}, \qquad \text{for some constant} \; A, \\
  \implies iR &amp;= E - Ae^{-\frac{R}{L}t}, \\
  \implies i &amp;= \frac{E}{R} - \frac{A}{R}e^{-\frac{R}{L}t}. \end{aligned} \]</p>

<p>We are given the condition \(i(0) = {@aa@}\). Substitution and rearranging yields</p>
<p>\[ \begin{aligned} \frac{E}{R} &amp;= {@aa@} + \frac{A}{R}, \\
  \implies A + {@aa@}R &amp;= E, \\
  \implies A &amp;= E - {@aa@}R, \\
  \implies i &amp;= \frac{E}{R} - \frac{{@AA@}}{R}e^{-\frac{R}{L}t}. \end{aligned} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(5);

AA: E - aa*R;

ta1: E/R - (AA/R)*%e^(-R*t/L);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(iR + L\dfrac{\mathrm{d}i}{\mathrm{d}t} = E, i(0) = {@aa@}, i = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>17</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1215835544</deployedseed>
    <deployedseed>1109657786</deployedseed>
    <deployedseed>545246032</deployedseed>
    <deployedseed>1061736733</deployedseed>
    <deployedseed>1358371061</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(2*ta1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
