<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>19.2.2.4 Separation of Variables Ex2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Solve the equation \(\dfrac{\mathrm{d}y}{\mathrm{d}x} = {@ff@}\) subject to the condition \(y(0) = 1\). Note: Type "log" or "ln" to represent the natural logarithm.</p>
<p>\(y =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the equation \[ \dfrac{\mathrm{d}y}{\mathrm{d}x} = {@ff@}\]</p>
<p>Rearrange to obtain integrals on both sides \[\displaystyle\int {@%e^(bb*y)@}\ \mathrm{d}y = \displaystyle\int {@aa*x^(nn)@}\ \mathrm{d}x\]</p>
<p>Perform the integration and rearrange to find \(y\). \[\begin{aligned} {@1/(bb)*%e^(bb*y)@} &amp;= {@aa/(nn+1)*x^(nn+1)@} + C, \\ {@%e^(bb*y)@} &amp;= {@aa*bb/(nn+1)*x^(nn+1)@} + {@bb@}C, \\ {@bb*y@} &amp;= \log\left({@aa*bb/(nn+1)*x^(nn+1)@} + {@bb@}C\right),
    \\ y &amp;= {@1/bb@}\log\left({@aa*bb/(nn+1)*x^(nn+1)@} + {@bb@}C\right) \end{aligned}\]</p>

<p>We are given the condition \(y(0)=1\). Substituting \(x=0\) and setting the equation equal to \(1\) we may find \(C\). \[\begin{aligned} 1 &amp;= {@1/bb@}\log({@bb@}C), \\ \implies {@%e^bb@} &amp; = {@bb@}C, \\ \implies C &amp;= {@%e^(bb)/bb@} \end{aligned}\]</p>

<p>Then substituting \(C = {@%e^(bb)/bb@}\) into our general equation for \(y\) yields \[y = {@ta1@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021120900</text>
    </stackversion>
    <questionvariables>
      <text>aa: 3+rand(3);
nn: 2+rand(3);
bb:2+rand(3);

cc: %e^(bb)/(bb);

ff: aa*x^(nn)*%e^(-bb*y);



ta1: log(aa*bb/(nn+1)*x^(nn+1)+bb*cc);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ \dfrac{\mathrm{d}y}{\mathrm{d}x} = {@ff@}, y = {@ta1@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>137313246</deployedseed>
    <deployedseed>863590365</deployedseed>
    <deployedseed>624534801</deployedseed>
    <deployedseed>341522769</deployedseed>
    <deployedseed>2081490967</deployedseed>
    <deployedseed>1931087069</deployedseed>
    <deployedseed>1763038892</deployedseed>
    <deployedseed>198549496</deployedseed>
    <deployedseed>1644094461</deployedseed>
    <deployedseed>884795138</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
