<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Diff-09-misc-3.c</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the following derivative: \[ \frac{d}{d{@v@}}\left[\left({@pp@}\right)^{@rdm2@}\right]. \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To solve this problem you will need to use the chain rule calc_chain_rule and the quotient rule calc_quotient_rule. Let \[f({@v@})={@p@}\] and \[g({@v@})={@pp@}.\] Then \(f({@v@})=[g({@v@})]^{@rdm2@}\). So, \[\begin{array}{rcl}\displaystyle \frac{d}{d{@v@}}\left[{@p@}\right] &amp;\displaystyle =&amp;\displaystyle \frac{df}{dg}\cdot \frac{dg}{dx} \\ &amp;\displaystyle =&amp;\displaystyle \frac{d}{dg}[g^{@rdm2@}]\cdot \frac{d}{d{@v@}}[g] \end{array}\] In order to find \(\frac{dg}{d{@v@}}\) you have to define two new functions: \(h({@v@})={@ppp1@}\) and \(i({@v@})={@ppp2@}\). So, \[\begin{array}{rcl}\displaystyle \frac{d}{d{@v@}}(g) &amp;\displaystyle =&amp;\displaystyle \frac{d}{d{@v@}}\left[{@pp@}\right] \\ &amp;\displaystyle =&amp;\displaystyle \frac{d}{d{@v@}}\left(\frac{h}{i}\right) \\ &amp;\displaystyle =&amp;\displaystyle \frac{\frac{dh}{d{@v@}}\cdot i - h\cdot \frac{di}{d{@v@}}}{i^2} \\ &amp;\displaystyle =&amp;\displaystyle \frac{\left[\frac{d}{d{@v@}}({@ppp1@})\right]\cdot ({@ppp2@}) - ({@ppp1@})\cdot \left[\frac{d}{d{@v@}}({@ppp2@})\right]}{({@ppp2@})^2} \\ &amp;\displaystyle =&amp;\displaystyle \frac{[{@diff(ppp1,v)@}]\cdot ({@ppp2@}) - ({@ppp1@})\cdot [{@diff(ppp2,v)@}]}{@ppp2^2@} \\ &amp;\displaystyle =&amp;\displaystyle \frac{@factor(diff(ppp1,v)*ppp2-ppp1*diff(ppp2,v))@}{@ppp2^2@}. \end{array}\] Now, \[\begin{array}{rcl}\displaystyle \frac{d}{dg}[g^{@rdm2@}] &amp;\displaystyle =&amp;\displaystyle {@rdm2@}g^{@rdm2-1@} \\ &amp;\displaystyle =&amp;\displaystyle {@rdm2*pp^(rdm2-1)@} \end{array}\] Hence, \[\begin{array}{rcl}\displaystyle \frac{d}{d{@v@}}\left[{@p@}\right] &amp;\displaystyle =&amp;\displaystyle \frac{d}{dg}[g^{@rdm2@}]\cdot \frac{d}{d{@v@}}[g] \\ &amp;\displaystyle =&amp;\displaystyle \frac{@factor(diff(ppp1,v)*ppp2-ppp1*diff(ppp2,v))@}{@ppp2^2@}\cdot {@rdm2*pp^(rdm2-1)@} \\ &amp;\displaystyle =&amp;\displaystyle {@factor(diff(p,v))@} \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm1:2+rand(7);
rdm2:(2+2*rand(5))/(3+2*rand(5));
ppp1:1-v^rdm1;
ppp2:1+v^rdm1;
pp:ppp1/ppp2;
p:pp^(rdm2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\frac{d}{d{@v@}}({@p@}) = {@ratsimp(diff(p,v))@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>factor(diff(p,v))</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>q:(ppp2/ppp1)^(rdm2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>factor(diff(p,v))</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>-factor(diff(p,v))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have lost the minus sign as your answer is the negative of what was expected.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>factor(diff(p,v))/diff(pp,v)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You need to use the chain rule calc_chain_rule as \[{@p@}\] is a function of {@v@}.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>factor(diff(q,v))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-3-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have differentiated \[{@q@}\] instead of the other way up.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-3-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2134062042</deployedseed>
    <deployedseed>883492275</deployedseed>
    <deployedseed>437165572</deployedseed>
    <deployedseed>1586258341</deployedseed>
    <deployedseed>1353997515</deployedseed>
    <deployedseed>85608009</deployedseed>
    <deployedseed>1439380078</deployedseed>
    <deployedseed>728084717</deployedseed>
    <deployedseed>2130230078</deployedseed>
    <deployedseed>1211343946</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(factor(diff(p,v)),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(-factor(diff(p,v)),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(factor(diff(p,v))/diff(pp,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(factor(diff((ppp2/ppp1)^(rdm2),v)),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
