<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>35.4.3.3 Jobs</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<h4 class="HELM_exercise"> Exercise </h4>
<p> A company advertises a job in three newspapers, \(P_A\), \(P_B\) and \(P_C\). It
is known that these papers attract undergraduate readerships in the
proportions \(2:3:1\). The probabilities that an undergraduate replies
to the job advertisement after they see it in these papers&nbsp;are \(\displaystyle{@pra@}\), \(\displaystyle{@prb@}\) and \(\displaystyle{@prc@}\) respectively. Assume that the undergraduate sees only one job advertisement. Let\[\begin{aligned}A &amp;= \{\text{Person is a reader of paper }P_A\},\\ B &amp;= \{\text{Person is a reader of paper }P_B\},\\ C &amp;= \{\text{Person is a reader of paper }P_C\},\\ R &amp;= \{\text{Reader applies for the job}\}.\end{aligned}\]
</p>
<ol class="HELM_parts">
<li> If the company receives only one reply to it advertisements, calculate the probability that the applicant has seen the job advertised in \(P_A\), \(P_B\) or \(P_C\). Specifically, calculate that probabilities \(\mathbb{P}(A|R)\), \(\mathbb{P}(B|R)\), \(\mathbb{P}(C|R)\).</li>
<li> If the company receives two replies, what is the probability that both
applicants saw the job advertised in paper \(P_A\)? I.e. let \(D = \{\text{Both applicants read }P_A\}\) and calculate \(\mathbb{P}(D)\). Assume the readerships and replies are independent. </li>
</ol>
<p> Give all probabilities as fractions, not decimals: </p>
<ol class="HELM_parts">
<li> \(\mathbb{P}(A|R)=\quad\)[[input:ans1]][[validation:ans1]][[feedback:prt1]], \(\quad\mathbb{P}(B|R)=\quad\)[[input:ans2]][[validation:ans2]][[feedback:prt2]], \(\quad\mathbb{P}(C|R)=\quad\)[[input:ans3]][[validation:ans3]][[feedback:prt3]] </li>
<li> \(\mathbb{P}(D)=\quad\)[[input:ans4]][[validation:ans4]][[feedback:prt4]] </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h4 class="HELM_solution"> Solution </h4>
<p> We have the probabilities\[\begin{aligned}\mathbb{P}(A) &amp;= {@pa@};\quad\mathbb{P}(R|A)={@pra@}\\ \mathbb{P}(B) &amp;={@pb@};\quad\mathbb{P}(R|B)={@prb@}\\ \mathbb{P}(C) &amp;= {@pc@};\quad\mathbb{P}(R|C)={@prc@}.\end{aligned}\]Therefore:</p>
<ol class="HELM_parts">
<li> \[\mathbb{P}(A|R) = \frac{\mathbb{P}(R|A)\mathbb{P}(A)}{\mathbb{P}(R|A)\mathbb{P}(A)+\mathbb{P}(R|B)\mathbb{P}(B)+\mathbb{P}(R|C)\mathbb{P}(C)}={@ta1@}.\]Similarly \(\mathbb{P}(B|R)={@ta2@}\) and \(\mathbb{P}(C|R)=\displaystyle {@ta3@}.\)</li>
<li> Since each reader is independent we have \(\mathbb{P}(D) = \mathbb{P}(A|R)\times\mathbb{P}(A|R) =\displaystyle {@ta4@}.\)</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>pa: 1/3;
pb: 1/2;
pc: 1/6;
pra: rand([1/1000,2/1000,3/1000,4/1000,5/1000]);
prb: rand([1/1000,2/1000,3/1000,4/1000,5/1000]);
prc: rand([1/1000,2/1000,3/1000,4/1000,5/1000]);
den: pa*pra + pb*prb + pc*prc;
ta1: pa*pra / den;
ta2: pb*prb / den;
ta3: pc*prc / den;
ta4: ta1 * ta1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\mathbb{P}(A|R) = {@ta1@}\]
\[\mathbb{P}(B|R) = {@ta2@}\]
\[\mathbb{P}(C|R) = {@ta3@}\]
\[\mathbb{P}(D) = {@ta4@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: green; --darkreader-inline-color:#72ff72;" data-darkreader-inline-color=""><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: orange; --darkreader-inline-color:#ffae1a;" data-darkreader-inline-color=""><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: red; --darkreader-inline-color:#ff1a1a;" data-darkreader-inline-color=""><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ans1*ans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Answer is valid based on the incorrect entry for part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>596069078</deployedseed>
    <deployedseed>1480230072</deployedseed>
    <deployedseed>1143048416</deployedseed>
    <deployedseed>1357995077</deployedseed>
    <deployedseed>1322596302</deployedseed>
    <deployedseed>826772</deployedseed>
    <deployedseed>372112744</deployedseed>
    <deployedseed>1317869805</deployedseed>
    <deployedseed>2138561527</deployedseed>
    <deployedseed>882793466</deployedseed>
    <deployedseed>1067182412</deployedseed>
    <deployedseed>1830847021</deployedseed>
    <deployedseed>789917607</deployedseed>
    <deployedseed>991559287</deployedseed>
    <deployedseed>718337277</deployedseed>
    <deployedseed>868344177</deployedseed>
    <deployedseed>1771961366</deployedseed>
    <deployedseed>151037292</deployedseed>
    <deployedseed>1913743157</deployedseed>
    <deployedseed>831468015</deployedseed>
    <deployedseed>1426354232</deployedseed>
    <deployedseed>1168238036</deployedseed>
    <deployedseed>811400013</deployedseed>
    <deployedseed>1122294163</deployedseed>
    <deployedseed>151547296</deployedseed>
    <deployedseed>719301503</deployedseed>
    <deployedseed>745783057</deployedseed>
    <deployedseed>1617419732</deployedseed>
    <deployedseed>1586490112</deployedseed>
    <deployedseed>1162853152</deployedseed>
    <deployedseed>2129766752</deployedseed>
    <deployedseed>1746104108</deployedseed>
    <deployedseed>871779087</deployedseed>
    <deployedseed>1269887753</deployedseed>
    <deployedseed>1742356199</deployedseed>
    <deployedseed>1357348690</deployedseed>
    <deployedseed>1759332761</deployedseed>
    <deployedseed>1202829094</deployedseed>
    <deployedseed>1936012647</deployedseed>
    <deployedseed>1147032625</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
