<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>28.3.3.8 cylindrical polar coordinates div and curl</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>For \(\underline{F} = {@aa@}\rho\hat{\underline{\rho}} + ({@bb@}\rho \sin \phi + z) \hat{\underline{\phi}} + \rho z \hat{\underline{z}}\), find \(\underline{\nabla} \cdot \underline{F}\) and \(\underline{\nabla} \times \underline{F}\).</p>

<p>\(\underline{\nabla} \cdot \underline{F} =\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>\(\underline{\nabla} \times \underline{F} =\) [[input:ans2]][[validation:ans2]]\(\hat{\underline{\rho}} +\)[[input:ans3]][[validation:ans3]]\(\hat{\underline{\phi}} + \)[[input:ans4]][[validation:ans4]]\(\hat{\underline{z}}\) [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We know that \(\underline{F} = {@aa@}\rho\hat{\underline{\rho}} + ({@bb@}\rho \sin \phi + z) \hat{\underline{\phi}} + \rho z \hat{\underline{z}}\) and aim to find \(\underline{\nabla} \cdot \underline{F}\) and \(\nabla \times \underline{F}\).</p>
<p>\[
\begin{aligned}
\underline{\nabla} \cdot \underline{F} &= \frac{1}{\rho}\left[\frac{\partial}{\partial \rho}(\rho({@aa@}\rho)) + \frac{\partial}{\partial \phi}({@bb@}\rho\sin\phi+z) + \frac{\partial}{\partial z}(\rho(\rho z))\right] \\ 
&= \frac{1}{\rho}\left[{@2*aa@}\rho + {@bb@}\rho\cos\phi +\rho^2\right] \\
&= {@2*aa@}+ {@bb@}\cos\phi + \rho \\
&\\
\nabla \times \underline{F} &= \frac{1}{\rho}\begin{vmatrix}
\hat{\underline{\rho}} & \rho \hat{\underline{\phi}} & \hat{\underline{z}} \\
\frac{\partial}{\partial \rho} & \frac{\partial}{\partial \phi} & \frac{\partial}{\partial z} \\
{@aa@}\rho & ({@bb@}\rho^2\sin\phi + \rho z) & \rho z
\end{vmatrix} \\
&= \frac{1}{\rho}\left[
\hat{\underline{\rho}} \left( \frac{\partial}{\partial \phi} (\rho z) - \frac{\partial}{\partial z} ({@bb@}\rho^2\sin\phi + \rho z) \right)
-\rho \hat{\underline{\phi}} \left(\frac{\partial}{\partial \rho}(\rho z) -  \frac{\partial}{\partial z}({@aa@}\rho) \right)
+ \hat{\underline{z}} \left(\frac{\partial}{\partial \rho}({@bb@}\rho^2\sin\phi + \rho z) - \frac{\partial}{\partial \phi}({@aa@}\rho) \right)
\right]\\
&= \frac{1}{\rho}\left[
\hat{\underline{\rho}}(0 - \rho)
-\rho \hat{\underline{\phi}}(z-0)
+\hat{\underline{z}}({@2*bb@}\rho\sin\phi + z - 0)
\right]\\
&= \frac{1}{\rho}\left[
- \rho\hat{\underline{\rho}}
-\rho z \hat{\underline{\phi}}
+({@2*bb@}\rho\sin\phi + z)\hat{\underline{z}}
\right]\\
&= - \hat{\underline{\rho}} - z \hat{\underline{\phi}} +({@2*bb@}\sin\phi + \frac{z}{\rho})\hat{\underline{z}}
\end{aligned}
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(6)+2;
bb:rand(6)+2;

ta1:2*aa+bb*cos(phi)+rho;
ta2i:-1;
ta2j:-z;
ta2k:2*bb*sin(phi)+z/rho;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\[\underline{F} = {@aa@}\rho\hat{\underline{\rho}} + ({@bb@}\rho \sin \phi + z) \hat{\underline{\phi}} + \rho z \hat{\underline{z}}\]
\[\underline{\nabla} \cdot \underline{F} ={@ta1@}\]
\[\underline{\nabla} \times \underline{F} = {@ta2i@}\hat{\underline{\rho}} + {@ta2j@}\hat{\underline{\phi}} + {@ta2k@}\hat{\underline{z}}\]</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2i</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta2j</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta2k</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2i</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta2j</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta2k</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>242377006</deployedseed>
    <deployedseed>90858585</deployedseed>
    <deployedseed>1167235472</deployedseed>
    <deployedseed>1119196832</deployedseed>
    <deployedseed>1461051630</deployedseed>
    <deployedseed>986977652</deployedseed>
    <deployedseed>1504276574</deployedseed>
    <deployedseed>1944506879</deployedseed>
    <deployedseed>1534175040</deployedseed>
    <deployedseed>1535318285</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2i</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta2j</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta2k</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+5</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2i+9</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta2j</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta2k</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
