<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>13.3.2.7 Area under a curve</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p><span style="font-size: calc(0.90375rem + 0.045vw);">Find the area under \(y = \displaystyle{@exp4@}\) between \(x = {@l4@}\) and \(x = {@u4@}\). [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</span></p><p><span style="font-size: calc(0.90375rem + 0.045vw);">You should check that the required area lies entirely above the horizontal axis.<br></span></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Since the required area lies entirely above the horizontal axis for each question, the area under the curve \(y\) between \(x = a\) and \(x = b\) is \(\displaystyle\int_{a}^{b} y(x) \, \mathrm{d}x\).</p>
<p>\(\displaystyle \int_{@l4@}^{@u4@} {@exp4@} \, \mathrm{d}x = \left[ {@int(exp4,x)@} \right]_{@l4@}^{@u4@} = {@ta4@}.\)</p>

<p>A sketch of the graph is shown below.</p>
<p style="text-align: center;"></p>
<div style="margin-left:4em;margin-right:auto;">
[[jsxgraph width='350px' height='350px']] 
JXG.Options.axis.ticks.insertTicks = false; /*don't insert extra ticks between major ticks*/ 
JXG.Options.axis.ticks.majorHeight = 10; 
JXG.Options.axis.ticks.minorHeight = 0; /*if minor ticks needed adjust height*/
JXG.Options.text.fontSize = 16; 
JXG.Options.axis.ticks.ticksDistance=1000; 
var board = JXG.JSXGraph.initBoard(divid, {boundingbox: {#BB#}, axis: false, showCopyright:false, showNavigation:false}); 
var f = board.jc.snippet('{#exp4#}', true, 'x', true); 
var ff = board.create('functiongraph',[f,1,{#ev(u4+1)#}]); 
var i1 = board.create('integral', [{#[l4,u4]#}, ff], {fillColor:'#003399',fillOpacity:0.15}); 
i1.curveLeft.setAttribute({visible:false}); /* hide the gliders */ 
i1.curveRight.setAttribute({visible:false});
i1.label.setAttribute({visible:false}); /*hide the integral calculation*/ 
var xaxis = board.create('axis', [[0,0],[1,0]], {name:'x', withLabel: true, label: {position: 'rt', offset: [-10, 10] /* (in pixels)*/}}); 
board.create('ticks', [xaxis, {#[l4,u4]#}],{strokeColor:'#003399', majorHeight:15, drawLabels: true, label: { offset: [0, -5],anchorX: 'middle', anchorY: 'top' }}); 
var yaxis = board.create('axis', [[0, 0], [0, 1]], {name:'y', withLabel: true, label: {position: 'rt', offset: [-20, 0] }}); 
[[/jsxgraph]]
</div>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>dd: rand([2, 4, 6, 8]);
exp4: dd*x^(-2);
l4:1;                   /* lower limit*/
u4: rand(7)+5;   /* upper limit*/
ta4: int(exp4, x, l4, u4);

/*jsxgraph settings */ 
BB: [-1, dd*(l4)^(-2), u4+2, -(1/6)*dd*(l4)^(-2)]     /* boundingbox:[left, top, right, bottom]*/ 
a: ev(exp4, x=l4);    /* y axis ticks*/
b: ev(exp4, x=u4);


plot1: plot(exp4, [x, l4-1,u4+1], [y, dd*(u4)^(-2)-1, dd*(l4)^(-2)+1], [box, false], [axes,solid]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \displaystyle \int_{@l4@}^{@u4@}  {@exp4@} \, \mathrm{d}x = {@ta4@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>t</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1651997333</deployedseed>
    <deployedseed>26321727</deployedseed>
    <deployedseed>1024763884</deployedseed>
    <deployedseed>75244493</deployedseed>
    <deployedseed>592131890</deployedseed>
    <deployedseed>1657208556</deployedseed>
    <deployedseed>1390433546</deployedseed>
    <deployedseed>982601167</deployedseed>
    <deployedseed>1394793454</deployedseed>
    <deployedseed>891334864</deployedseed>
    <deployedseed>1026384303</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans4</name>
        <value>ev(ta4+1,simp)</value>
      </testinput>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
