<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>13.3.3.3 Finding areas</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>

<ol class="HELM_parts">

<li> Find the total area enclosed between the \(x\)-axis and the curve \(y={@exp1@}\) between \(x={@l1@}\) and \(x={@u1@}\). [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li> 

<li>Fnd the area under \(y={@exp2@}\) from \(t={@l2@}\) to \(t={@u2@}\). Give your answer to \(4\) significant figures. [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>

<li> Find the area enclosed by \(y={@exp3@}\) and the \(x\)-axis from \(x={@l3@}\) to \(x={@u3@}\). [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</li>

</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">

    <li>
        <p> Since the curve \(y={@exp1@}\) is above the \(x\) axis between 0 and {@u1@}, and below the \(x\) axis between {@l1@} and 0. The area under the curve is: \[-\displaystyle\int_{@l1@}^{0} {@exp1@} \, \mathrm{d}x + \displaystyle\int_{0}^{@u1@} {@exp1@}
            \, \mathrm{d}x = {@ta1@}.\]</p>
    </li>

    <li><p>The curve \(y={@exp2@}\) is above the \(t\) axis from \(t={@l2@}\) to \(t={@u2@}\). The area is: \[\displaystyle\int_{@l2@}^{@u2@} {@exp2@} \, \mathrm{d}t = \left[ {@int(exp2,t)@} \right]_{@l2@}^{@u2@} = {@ta2@}.\]</p>
  </li>

    <li><p> The area under the curve \(y={@exp3@}\) is above the \(x\) axis from \(x={@l3@}\) to \(x={@u3@}\), hence the area is: \[\displaystyle \int_{@l3@}^{@u3@} ({@exp3@}) \, \mathrm{d}x =\left[ {@int(exp3,x)@} \right]_{@l3@}^{@u3@} = {@ta3@}.\]</p>
  </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand([1,1/2, 1/3]);
exp1: aa*x^3;
l1: rand([-1,-2]);
u1: -l1;
ta1:  2*int(exp1, x, 0, l1); 

exp2: cos(2*t);
l2: 0;
u2: 0.5;
ta2: significantfigures(float(int(exp2, t, l2, u2)),4); 

bb: rand([4,9,16]);
exp3: bb-x^2;
l3: rand(2);
u3: abs(sqrt(bb));
ta3: int(exp3, x, l3, u3); 

/*Maxima Plots: don't used here*/
plot1: plot(exp1, [x, l1-1,u1+1], [y, aa*(l1)^(3)-1, aa*(u1)^(3)+1], [box, false], [axes,solid]);
plot2: plot(exp2, [t, l2-0.5,u2+0.5], [y, 0, 1], [box, false], [axes,solid]);
plot3: plot(exp3, [x, l3-1,u3+1], [y, bb-(u3)^(2)-1, bb-(l3)^(2)+1], [box, false], [axes,solid]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[-\displaystyle\int_{@l1@}^{0} {@exp1@} \, \mathrm{d}x + \displaystyle\int_{0}^{@u1@} {@exp1@} \, \mathrm{d}x = {@ta1@}, \quad  \displaystyle\int_{@l2@}^{@u2@} {@exp2@} \, \mathrm{d}t = {@ta2@}, \quad \displaystyle\int_{@l3@}^{@u3@} {@exp3@} \, \mathrm{d}x = {@ta3@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>dispsf(ta2,4)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumSigFigs</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1946208534</deployedseed>
    <deployedseed>1610864983</deployedseed>
    <deployedseed>1020268674</deployedseed>
    <deployedseed>2009380371</deployedseed>
    <deployedseed>1329910820</deployedseed>
    <deployedseed>75739690</deployedseed>
    <deployedseed>428960350</deployedseed>
    <deployedseed>774172075</deployedseed>
    <deployedseed>382830586</deployedseed>
    <deployedseed>1426246876</deployedseed>
    <deployedseed>832837760</deployedseed>
    <deployedseed>1567198831</deployedseed>
    <deployedseed>429815010</deployedseed>
    <deployedseed>686294221</deployedseed>
    <deployedseed>760638577</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>significantfigures(ta2,4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta3+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
    <hint format="html">
      <text><![CDATA[<p>If you are unsure of the limits of integration draw a diagram.</p>]]></text>
    </hint>
  </question>
</quiz>
