<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limits_2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A \({\bf rational}\mbox{ } {\bf function}\) is a quotient of two polynomials. Evaluate the following limit of a rational function: \[ \lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p1@}{@p2@}\right)}. \] If the limit does not exist, enter the value 1111 as your answer.</p>
<p>The limit is: [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The limit of the function \[ \lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p1@}{@p2@}\right)}. \] <br /> is the same as the limit found by keeping only the highest power terms in the numerator and the denominator. Hence resulting in a limit of {@correctanswer@}.</p>
<p>The limit of the function can be found by dividing the numerator and denominator by {@v^degr1@}.</p>
<p>The function can thus be expressed as \[ \lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p3@}{@p4@}\right)}. \]</p>
<p>As {@v@} tends to {@set1[ind1]@}, all terms apart from the constants tend to zero.</p>
<p>Hence the limit is {@correctanswer@}.</p>
<p> </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[v:x;
 degr1:rand_with_step(3, 9, 1);
 q:0;
 p1:block(for i:1 while i<=degr1 do (q:q+v^i*(-(rand(9)+1))),return(q));
 q1:0;
 p2:block(for i:1 while i<=degr1 do (q1:q1+v^i*(rand(9)+1)),return(q1));
 set1:[inf, minf];
 ind1:rand(2)+1;
 f:p1/p2;
 correctanswer:limit(p1/p2,v,set1[ind1]);
 p3 : expand(p1/(v^degr1));
 p4 : expand(p2/(v^degr1));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p1@}{@p2@}\right)}={@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords> limit</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1073083350</deployedseed>
    <deployedseed>603466032</deployedseed>
    <deployedseed>1341407777</deployedseed>
    <deployedseed>13262703</deployedseed>
    <deployedseed>1083555479</deployedseed>
    <deployedseed>619712693</deployedseed>
    <deployedseed>335572368</deployedseed>
    <deployedseed>1041993197</deployedseed>
    <deployedseed>2129968739</deployedseed>
    <deployedseed>685173265</deployedseed>
    <deployedseed>1184581491</deployedseed>
    <deployedseed>39276041</deployedseed>
    <deployedseed>1209037266</deployedseed>
    <deployedseed>1703222263</deployedseed>
    <deployedseed>1800308012</deployedseed>
    <deployedseed>141867793</deployedseed>
    <deployedseed>1326845592</deployedseed>
    <deployedseed>1861292374</deployedseed>
    <deployedseed>608595611</deployedseed>
    <deployedseed>1596001320</deployedseed>
    <deployedseed>1250801491</deployedseed>
    <deployedseed>194296305</deployedseed>
    <deployedseed>1764909867</deployedseed>
    <deployedseed>997365339</deployedseed>
    <deployedseed>842533674</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>1111</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer+1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
