<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-16-trigonometric-6</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \[ \int {@p@} d{@v@}. \] Hint: square an identity for \(\cos^2(A)\). [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>For this question we need to use the identity \[cos^2(A) = \frac{1}{2}[1+cos(2A)]\] as follows: \[\begin{array}{rcl}\displaystyle {@p@} &amp;\displaystyle =&amp;\displaystyle [{@pp@}]^2\\ &amp;\displaystyle =&amp;\displaystyle \left[{@q@}\right]^2 \\ &amp;\displaystyle =&amp;\displaystyle {@expand(q^2)@} \end{array}\] So, \[ \int {@p@} d{@v@} = \int {@expand(q^2)@} d{@v@} \] To solve this integral we need to express \[{@q2^2@}\] in a simpler form using the same identity: \[\begin{array}{rcl}\displaystyle {@q2^2@} &amp;\displaystyle =&amp;\displaystyle {@q1^2@}{@(2*q2)^2@} \\ &amp;\displaystyle =&amp;\displaystyle {@q1^2@}\left[{@q4@}\right] \\ &amp;\displaystyle =&amp;\displaystyle {@expand(q1^2*q4)@} \end{array}\] Hence, \[\begin{array}{rcl}\displaystyle \int {@p@} d{@v@} &amp;\displaystyle =&amp;\displaystyle \int {@expand(q1^2*q4)@} + {@2*q1*q2@} + {@q1^2@} d{@v@} \\ &amp;\displaystyle =&amp;\displaystyle \int {@expand(q1^2*q4)+2*q1*q2+q1^2@} d{@v@} \\ &amp;\displaystyle =&amp;\displaystyle {@int(expand(q1^2*q4)+2*q1*q2+q1^2,v)+c@} \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:t;
rdm1:1+rand(3);
k:2;
pp:(cos(rdm1*v))^2;
p:pp^k;
q1:1/2;
q2:cos(2*rdm1*v)/2;
q:q1+q2;
q3:cos(4*rdm1*v);
q4:q3/2+q1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int {@p@} d{@v@} = {@expand(int(p,v))@}+c\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>expand(int(p,v))+c</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>expand(int(p,v))+c</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2020338937</deployedseed>
    <deployedseed>1074994269</deployedseed>
    <deployedseed>1726757420</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
