<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.4.3.4 jacobian transformation task</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the Jacobian for the following transformation: <br>\(x={@xexp@}\), <br> \(y={@xexp@}\), <br>\(z={@xexp@}\)</p>
<p>\(|J|=\,\)[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[Evaluating the partial derivatives, <br>

  \(\qquad \dfrac{\partial x}{\partial u}= {@dxdu@}\), 
  \(\qquad\dfrac{\partial x}{\partial v}= {@dxdv@}\), 
  \(\qquad\dfrac{\partial x}{\partial w}= {@dxdw@}\),  <br>
  \(\qquad \dfrac{\partial y}{\partial u}= {@dydu@}\),
  \(\qquad\dfrac{\partial y}{\partial v}=  {@dydv@}\),
  \(\qquad\dfrac{\partial y}{\partial w}=  {@dydw@}\),  <br>
  \(\qquad \dfrac{\partial z}{\partial u}=  {@dzdu@}\),
  \(\qquad  \dfrac{\partial z}{\partial v}= {@dzdv@}\), 
  \(\qquad  \dfrac{\partial z}{\partial w}= {@dzdw@}\)  <br> so the Jacobian is  
 \[\left| \begin{array}{ccc} \dfrac{\partial x}{\partial u} &
\dfrac{\partial x}{\partial v} & \dfrac{\partial x}{\partial w} \\
\\
\dfrac{\partial y}{\partial u} & \dfrac{\partial y}{\partial v} &
\dfrac{\partial y}{\partial w} \\ \\ \dfrac{\partial z}{\partial u}
& \dfrac{\partial z}{\partial v} & \dfrac{\partial z}{\partial w}
\end{array} \right| = \left| \begin{array}{ccc} {@dxdu@} &
{@dxdv@}  & {@dxdw@}  \\
\\
{@dydu@} & {@dydv@} & {@dydw@} 
\\ \\ {@dzdu@} & {@dzdv@} & {@dzdw@}
\end{array}\right|  = {@ta1@} \] 
where expansion of the determinant has taken place down the first
column.]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1 + rand(6);
bb: 1 + rand(6);
cc: 1 + rand(6);
dd: 1 + rand(6);
ee: 1 + rand(6);




xexp: aa*u+bb*v-w;
yexp: v-dd*w;
zexp: u+ee*w;

dxdu:diff(xexp,u);
dxdv:diff(xexp,v);
dxdw:diff(xexp,w);


dydu:diff(yexp,u);
dydv:diff(yexp,v);
dydw:diff(yexp,w);


dzdu:diff(zexp,u);
dzdv:diff(zexp,v);
dzdw:diff(zexp,w);

M: matrix ([dxdu, dxdv, dxdw], [dydu, dydv, dydw], [dzdu, dzdv, dzdw]);


ta1: determinant(M);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>\(x={@xexp@}, \quad y={@xexp@} \quad z={@xexp@} \implies {@ta1@} \)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1891036394</deployedseed>
    <deployedseed>1310057289</deployedseed>
    <deployedseed>145662818</deployedseed>
    <deployedseed>1375870548</deployedseed>
    <deployedseed>1735946362</deployedseed>
    <deployedseed>1398255162</deployedseed>
    <deployedseed>1983765011</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
