<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 311  -->
  <question type="stack">
    <name>
      <text>Using Block: Lagrange multipliers 02 (3 variables, 1 constraint)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function \(f(x,y,z) = {@ff@}\) and the constraint \({@constraint@}\).</p>
<p>We want to optimize \(f\) under this constraint using the method of Lagrange multipliers.</p>
<ul>
    <li>How many equations will the corresponding system of equations have?</li>
</ul>
<p>Answer [[input:ans1]] [[validation:ans1]]</p>
<ul>
    <li>And how many variables (unknowns) does this system have?</li>
</ul>
<p>Answer: [[input:ans2]] [[validation:ans2]]</p>
<ul>
    <li>Write down the equations:<br></li>
</ul>
<p>
[[repeat id="1"]]
Equation <repeatindex>: [[input:ans3]] [[validation:ans3]]<br>
[[/repeat]]
[[repeatbutton title="Add another equation" repeat_ids="1" /]]
</p>
<p><em>Hint</em>: Use the Greek letters \(\lambda, \mu, \nu\) in this order if you need further variables</p>
<p><em>Hinweis</em>: You can type \(\lambda\) as <code>lambda</code>, \(\mu\) as <code>mu</code>, and \(\nu\) as <code>nu</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Lösung</strong>: Die Anzahl Nebenbedingungen entspricht genau der Anzahl Lagrange Multiplikatoren (also, die Anzahl weitere Variablen). Wir haben somit in diesem Fall \(1\) Langrange Multiplikator. Wir bezeichnen diesen mit \(\lambda\). Man muss nun \(x\), \(y\), \(z\) und \(\lambda\) bestimmen (falls diese überhaupt existieren!) so dass folgende Gleichungen erfüllt sind:</p>
<p>\[\nabla f(x,y,z) = \lambda\nabla g(x,y,z) \ \ \text{und} \ \ g(x,y,z) = 0,\]</p>
<p>wobei \(f(x,y,z) = {@ff@}\) und \(g(x,y,z) = {@gg@}\). In unserem Fall lauten die Gleichungen also</p>
<p>\[\left\{\begin{array}{l} {@eq1@}, &amp; \\ {@eq2@}, &amp; \\ {@eq3@}, &amp; \\ {@eq4@}. \end{array} \right.\]</p>
<p>Und damit ist es klar, dass man \({@ta1@}\) Gleichungen und \({@ta2@}\) Variablen hat.</p>]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025042500</text>
    </stackversion>
    <questionvariables>
      <text>/*Written by Meike Akveld, George Ionita, Andreas Steiger, ETH Zürich
Contact and more information:
stack@math.ethz.ch - www.math.ethz.ch/stack*/
/*Author of the Script: Michael Kallweit, Ruhr Universität Bochum
You find further information on our website:
https://www.ruhr-uni-bochum.de/hdm/stack.en.shtml*/
aa:rand([-4,-3,-2]);
if aa=-4 then bb:rand([3,5]) else if aa=-3 then bb:rand([2,4,5]) else bb:rand([3,5]);
cc:rand_with_prohib(2,5,[abs(aa),bb,gcd(abs(aa),bb)]);
/*kk:rand_with_prohib(-3,3,[0,1,-1]);
ll:rand_with_prohib(-3,3,[0,1,-1,kk,-kk]);
mm:-4;
nn:rand([1,2,3]);
pp:rand_with_prohib(1,3,[kk]);
qq:rand_with_prohib(1,3,[kk,ll]);*/
position:1+rand(2);
ordergreat(x,y,z);
/*fflist:[aa*x+bb*y+cc*z, kk*x^2+ll*y^2+mm*z^2, x^nn*y^pp*z^qq];*/
/*ff:fflist[position];*/
ff:aa*x+bb*y+cc*z;
constraintlist:[x^2+y^2+z^2=4, x^2+y^2+z^2=3];
constraint:constraintlist[position];
gglist:[x^2+y^2+z^2-4, x^2+y^2+z^2-3];
gg:gglist[position];
ta1:4;
/*ta2:[x, y, z, lambda];*/
ta2:4;
ffx:diff(ff,x);
ffy:diff(ff,y);
ffz:diff(ff,z);
ggx:diff(gg,x);
ggy:diff(gg,y);
ggz:diff(gg,z);

eq1:ffx=lambda*ggx; eq1alt1:ffx=-lambda*ggx;
eq2:ffy=lambda*ggy; eq2alt1:ffy=-lambda*ggy;
eq3:ffz=lambda*ggz; eq3alt1:ffz=-lambda*ggz;
eq4:gg=0;
ta3:[eq1, eq2, eq3, eq4];
ta3alt1:[eq1alt1, eq2alt1, eq3alt1, eq4];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]] [[feedback:prt2]] [[feedback:prt3]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(f(x,y,z) = {@ff@}\) subject to \({@constraint@}\), \(g(x,y,z) = {@gg@}\). Equation 1: \({@eq1@}\), Equation 2: \({@eq2@}\), Equation 3: \({@eq3@}\), Equation 4: \({@eq4@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Richtige Antwort, gut gemacht!]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Ihre Antwort ist teilweise korrekt.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Falsche Antwort.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Die Anzahl Gleichungen ist richtig.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Anzahl Gleichungen ist falsch.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Die Anzahl Variablen ist richtig.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Anzahl Variablen ist falsch.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>variablesineq(eq) := listofvars(eq);

solveforvars(eq) := block(
    [vars, sols],
    vars : variablesineq(eq),
    sols : [],
    for var in vars do (
        sols : cons(rhs(first(solve(eq, var))), sols)
    ),
    sols
);

solvelist(L) := map(lambda([eq], solveforvars(eq)), L);
SLans3: solvelist(ans3);
SLta3: solvelist(ta3);
SLta3alt1: solvelist(ta3alt1);

listtoset(lst):=setify(lst);

/* Convert the list of lists to a list of sets */
listofsets(L):=map(listtoset,L);

/* Convert the list of sets to a set of sets */
setofsets(L):=setify(listofsets(L));

com: intersection(setofsets(SLans3), setofsets(SLta3));
comalt1: intersection(setofsets(SLans3), setofsets(SLta3alt1));

if cardinality(com)=4 or cardinality(comalt1)=4 then equations4:true else equations4:false;
if cardinality(com)=3 or cardinality(comalt1)=3 then equations3:true else equations3:false;
if cardinality(com)=2 or cardinality(comalt1)=2 then equations2:true else equations2:false;
if cardinality(com)=1 or cardinality(comalt1)=1 then equations1:true else equations1:false;

if equations4=true and cardinality(setify(ans3))=4 then fullycorrect:true else fullycorrect:false;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>fullycorrect</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Das Gleichungssystem ist richtig.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Das Gleichungssystem ist leider falsch.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>equations4</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.75</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ihre Gleichungen sind korrekt, aber Sie haben zu viele.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>cardinality(setify(ans3))</sans>
        <tans>4</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt3-5-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt3-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben zu viele oder zu wenige Gleichungen, Sie sollten genau \({@ta1@}\) Gleichungen haben.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>equations3</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.75</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Drei Ihrer Gleichungen sind richtig. Bitte überprüfen Sie die Lösung, um herauszufinden, welche.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>equations2</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Zwei Ihrer Gleichungen sind richtig. Bitte überprüfen Sie die Lösung, um herauszufinden, welche.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt3-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>equations1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Nur eine Ihrer Gleichungen ist richtig. Bitte überprüfen Sie die Lösung, um herauszufinden, welche.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-6-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1815673680</deployedseed>
    <deployedseed>845066603</deployedseed>
    <deployedseed>731580177</deployedseed>
    <deployedseed>1042657083</deployedseed>
    <deployedseed>1879361070</deployedseed>
    <deployedseed>1168274040</deployedseed>
    <deployedseed>1658085489</deployedseed>
    <deployedseed>1029093734</deployedseed>
    <deployedseed>1117989814</deployedseed>
    <deployedseed>1696727951</deployedseed>
    <deployedseed>1181698641</deployedseed>
    <deployedseed>1080911372</deployedseed>
    <deployedseed>967010804</deployedseed>
    <deployedseed>2014495962</deployedseed>
    <deployedseed>498718360</deployedseed>
    <deployedseed>717113710</deployedseed>
    <deployedseed>798947642</deployedseed>
    <deployedseed>839238556</deployedseed>
    <deployedseed>363544625</deployedseed>
    <deployedseed>1654769667</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>MAES24</text>
</tag>
      <tag><text>MK</text>
</tag>
    </tags>
  </question>

</quiz>