<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 310  -->
  <question type="stack">
    <name>
      <text>Using Block: Polynomdivision 4.grades 02 (1 double, 2 simple real roots)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find all roots of the polynomial \(p(x)\). Potential multiple roots only have to be listed once.</p>
<p>\[p(x) = {@poly@}.\]</p>
<p>
[[repeat id="root"]]
\(x_<repeatindex> = \) [[input:ans1]] [[validation:ans1]]<br>
[[/repeat]]
[[repeatbutton title="Add another root" repeat_ids="root" /]]
</p>
<p><strong>Remark</strong>: You can't delete input fields. If you add more fields than you require in the end, you can simply leave them blank/empty. They will not be evaluated.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Lösung</strong>: Dies ist ein Polynom vierten Grades und wahrscheinlich kennen Sie die Formel für die Nullstellen nicht. </p>
<p>In dem Fall empfiehlt es sich zu versuchen eine Nullstelle zu raten. Man sieht hier ziemlich schnell, dass \(x = {@zero1@}\) eine Nullstelle von \(p(x)\) ist, d.h. \(p({@zero1@}) = 0\) (am besten probieren Sie als Nullstellen die ganze Zahlen von \(-2\) bis \(2\)).</p>
<p>Wenn \(x_1 = {@zero1@}\) eine Nullstelle ist, so ist \((x-x_1) = ({@x-zero1@})\) ein Faktor des Polynoms. Führen Sie eine Polynomdivision durch.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025042500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Written by Meike Akveld, George Ionita, Andreas Steiger, ETH Zürich
Contact and more information:
stack@math.ethz.ch - www.math.ethz.ch/stack*/
/*Author: Michael Kallweit, Ruhr Universität Bochum
You find further information on our website:
https://www.ruhr-uni-bochum.de/hdm/stack.en.shtml*/
texput(multsgnonlyfornumberssym, "\\cdot");
n:4;
zero1:rand([-1,2]); /*DOUBLE ROOT - 2 possibilities*/
zero2:rand_with_prohib(-3,3,[0,1,zero1]); /*SIMPLE ROOT - 4 possibilities*/
zero3:rand_with_prohib(-3,3,[0,1,zero1,zero2]); /*SIMPLE ROOT - 3 possibilities*/
/*24/2 = 12 variants in total*/
nst:([zero1, zero1, zero2, zero3]);
p:expand(apply("*",makelist((x-nst[i]),i,1,n)));
ta:nst;
poly:expand(p);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@poly@} = 0 \implies x_1 = x_2 = {@zero1@}, x_3, x_4 \in {@union(set(zero2), set(zero3))@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text>MEIKE'S BET: NO ONE WILL GIVE MORE THAN 4 SOLUTIONS!</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Richtige Antwort, gut gemacht!]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Ihre Antwort ist teilweise korrekt.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Falsche Antwort.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[if length(ans1)>4 then toomany:true else toomany:false;
if elementp(zero1,setify(ans1))=true then foundzero1:true else foundzero1:false;
if elementp(zero2,setify(ans1))=true then foundzero2:true else foundzero2:false;
if elementp(zero3,setify(ans1))=true then foundzero3:true else foundzero3:false;
if length(ans1)=4 and elementp(zero1,setify(ans1))=true and elementp(zero2,setify(ans1))=true and elementp(zero3,setify(ans1))=true then doublerootwrong:true else doublerootwrong:false;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>toomany</sans>
        <tans>false</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben zu viele Nullstellen gefunden. Ein Polynom vom Grad \({@n@}\) hat maximal \({@n@}\) reelle Nullstellen (und man zählt mit Multiplizität). In diesem Beispiel gibt's eine doppelte Nullstelle, also maximal drei verschiedene Nullstellen.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>setify(ans1)</sans>
        <tans>setify(ta)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>foundzero1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sehr gut, Sie haben die doppelte Nullstelle \({@zero1@}\) richtig identifiziert.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Schade, Sie haben die doppelte Nullstelle \({@zero1@}\) nicht gefunden.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>foundzero2</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sehr gut, Sie haben die Nullstelle \({@zero2@}\) richtig identifiziert.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Schade, Sie haben die Nullstelle \({@zero2@}\) nicht gefunden.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>foundzero3</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sehr gut, Sie haben die Nullstelle \({@zero3@}\) richtig identifiziert.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Schade, Sie haben die Nullstelle \({@zero3@}\) nicht gefunden.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>doublerootwrong</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ihre Nullstelle \({@listify(setdifference(setify(ans1),setify(ta)))[1]@}\) ist leider falsch.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>619277900</deployedseed>
    <deployedseed>1010586411</deployedseed>
    <deployedseed>1326669131</deployedseed>
    <deployedseed>1223140985</deployedseed>
    <deployedseed>1202215767</deployedseed>
    <deployedseed>30632570</deployedseed>
    <deployedseed>1777631147</deployedseed>
    <deployedseed>1860241011</deployedseed>
    <deployedseed>1679009199</deployedseed>
    <deployedseed>851725381</deployedseed>
    <deployedseed>1035359516</deployedseed>
    <deployedseed>1529426851</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>MK</text>
</tag>
      <tag><text>MAEZ24</text>
</tag>
    </tags>
  </question>

</quiz>