<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>39.1.8.2 Normal Distribution Probability Interval Exercise</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>What is the centred \({@displaydp(numa,2)@}\%\) probability interval for the lifetime of a bulb, \(T\), when the lifetimes of such bulbs are normally distributed with a mean of \({@mu@}\) hours and standard deviation of \({@sigma@}\) hours?</p><p>(You <b>must</b> keep <b>full</b>&nbsp;accuracy throughout the question).</p>

<p><br></p>

<p>First sketch the standard normal curve marking the values \(z_1, z_2\) between which \({@displaydp(numa,2)@}\%\) of the area under the curve is located.</p>

<p>Now, find the corresponding interval \(z\)-values, \((z_1, z_2)\), to 2 decimal places using the Table.</p><p>\(\text{P}(\) [[input:ans1]] \(&lt; Z &lt;\) [[input:ans2]] \() = {@displaydp(numa,2)@}\%\)</p><p>[[validation:ans1]][[validation:ans2]]</p><p>[[feedback:prt1]][[feedback:prt2]]</p>
<p><span style="font-size: 0.9375rem;"><br></span></p><p><span style="font-size: 0.9375rem;">Now, calculate the corresponding interval values, \((t_1, t_2)\), to 1 decimal place that correspond to normal distribution for the bulb lifetime.</span><br></p><p><span style="font-size: 0.9375rem;"></span></p><p>\(\text{P}(\) [[input:ans3]] \(&lt; T &lt;\) [[input:ans4]] \() = {@displaydp(numa,2)@}\%\)</p><p>[[validation:ans3]][[validation:ans4]]</p><p>[[feedback:prt3]][[feedback:prt4]]</p><br>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h4 class="HELM_solution">Solution</h4>

<p>For the interval \((z_1, z_2)\) to be centred, we have that \(z_1 = -z_2\). This means that \(\text{P}(z_1 &lt; Z &lt; z_2) = 2 \times \text{P}(0 &lt; Z &lt; z_2)\)<span style="font-size: 0.9375rem;">. For this probability to be equal to&nbsp;\({@displaydp(numa,2)@}\%\) then we need to find the value of \(z_2\) such that \(\text{P}(0 &lt; Z &lt; z_2) = {@displaydp(prob,4)@}\). Reading from the table, we see that \(z_2 = {@displaydp(ta2,2)@}\) to 2 decimal places. Therefore:</span></p><p><span style="font-size: 0.9375rem;">\[\text{P}( {@displaydp(ta1,2)@} &lt; Z &lt; {@displaydp(ta2,2)@}) = {@displaydp(numa,2)@}\%\]</span></p><p>The distribution for the bulb life time is \(T \sim N({@mu@}, {@sigma@}^2)\). The transformation from the standard normal to this distribution is \(T = {@mu@} + {@sigma@}Z\). This means that the corresponding lower and upper bounds of the centred interval, rounded to the nearest whole hour, are:</p><p>\[\begin{align*} t_1 &amp; = {@mu@} + {@sigma@} z_1 = {@displaydp(ta3,0)@},\\ t_2 &amp; = {@mu@} + {@sigma@} z_2 = {@displaydp(ta4,0)@}. \end{align*}\]</p><p>Hence, the \({@numa@}\%\) centred probability interval for the lifetimes is \(({@displaydp(ta3,0)@},{@displaydp(ta4,0)@})\) hours.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024072400</text>
    </stackversion>
    <questionvariables>
      <text>mu:2000
sigma:40
z: decimalplaces(float(rand_with_step(2.47,2.67,0.01)),2);
prob: decimalplaces(float(cdf_normal(z,0,1)-0.5),4);
numa: decimalplaces(100*2*prob,2);

ta1: -z;
ta2: z;
ta3: decimalplaces(float(mu-sigma*z),0);
ta4: decimalplaces(float(mu+sigma*z),0);

checkz1: decimalplaces(float(quantile_normal(0.5 - numa/200,0,1)),2);
checkz2: decimalplaces(float(quantile_normal(0.5 + numa/200,0,1)),2);
checkt1: decimalplaces(float(quantile_normal(0.5 - numa/200,mu,sigma)),0);
checkt2: decimalplaces(float(quantile_normal(0.5 + numa/200,mu,sigma)),0);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\((z_1,z_2) = ({@displaydp(ta1,2)@}, {@displaydp(ta2,2)@}), (t_1, t_2) = ({@displaydp(ta3,0)@}, {@displaydp(ta4,0)@})\). Check: \((z_1, z_2) = ({@displaydp(checkz1,2)@}, {@displaydp(checkz2,2)@}), (t_1, t_2) = ({@displaydp(checkt1,0)@}, {@displaydp(checkt2,0)@})\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>dispdp(ta1,2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 2 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:2, mindp:2</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>dispdp(ta2,2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 2 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:2, mindp:2</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>dispdp(ta3,0)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>whole number</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:0, mindp:0</options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>dispdp(ta4,0)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>whole number</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:0, mindp:0</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1351767835</deployedseed>
    <deployedseed>904746595</deployedseed>
    <deployedseed>523064303</deployedseed>
    <deployedseed>861905995</deployedseed>
    <deployedseed>898655772</deployedseed>
    <deployedseed>893597759</deployedseed>
    <deployedseed>1248299862</deployedseed>
    <deployedseed>1106723249</deployedseed>
    <deployedseed>1472002833</deployedseed>
    <deployedseed>1494967825</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta1,2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta2,2)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ta3,0)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispdp(ta4,0)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta2,2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta1,2)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(-ta4,0)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispdp(-ta3,0)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta1+1,2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta2+1,2)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ta3+1,0)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispdp(ta4+1,0)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
