<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>35.2.1.8 Shock and scratch resistant paint</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<h4 class="HELM_exercise"> Exercise </h4>
<p> A new type of paint to be used in the manufacture of garden equipment is tested
for impact shock resistance to damage and scratch resistance to damage. The
results of the \({@num@}\) tests are as follows.</p>
<br>
Of the tests
<ul>
<li> \({@gg@}\) showed both good shock and scratch resistance. </li>
<li> \({@gp@}\) showed good shock resistance but poor scratch resistance. </li>
<li> \({@pg@}\) showed poor shock resistance but good scratch resistance. </li>
<li> \({@pp@}\) showed both poor shock and scratch resistance. </li>
</ul>
<p> Now, if \(A\) is the event \(\{\)Good shock resistance\(\}\) and \(B\) is the event \(\{\)Good scratch resistance\(\}\), the consider the following: </p>
<ol class="HELM_parts_inline">
<li> \(A\cup B\) </li>
<li> \(A\cap B\) </li>
<li> \(A'\) </li>
<li> \(B'\) </li>
</ol>
<p> and determine the number of samples in each:</p>
<ol class="HELM_parts">
<li> \(|A\cup B|=\quad\)[[input:ans1]][[validation:ans1]][[feedback:prt1]] </li>
<li> \(|A\cap B|=\quad\)[[input:ans2]][[validation:ans2]][[feedback:prt2]] </li>
<li> \(|A'|=\quad\)[[input:ans3]][[validation:ans3]][[feedback:prt3]] </li>
<li> \(|B'|=\quad\)[[input:ans4]][[validation:ans4]][[feedback:prt4]] </li>
</ol>
<p>Note: here the notation \(|A|\) represents the number of elements in \(A\). </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h4 class="HELM_solution"> Solution </h4>
<ol class="HELM_parts">
<li> The event \(A\cup B\) consists of those samples which have either good shock and good
scratch resistance (or both) so that \(|A\cup B|={@ta1@}\). </li>
<li> The event \(A\cap B\) consists of those samples which have both good shock and good scratch
resistance so that \(|A\cap B|={@ta2@}\). </li>
<li> The event \(A'\) consists of those samples which do not have good shock resistance so that \(|A'|={@ta3@}\). </li>
<li> The event \(B'\) consists of those samples which do not have good scratch resistance so that \(|B'|={@ta4@}\). </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>gg: 50 + rand(50);
gp: 20 + rand(20);
pg: 20 + rand(20);
pp: 3 + rand(10);
num: gg + gp + pg + pp;
ta1: gg+gp+pg;
ta2: gg
ta3: pg + pp
ta4: gp + pp
/* Tried to use a table but gave up */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ |A\cup B| = {@ta1@} \]
\[ |A\cap B| = {@ta2@} \]
\[ |A'| = {@ta3@} \]
\[ |B'| = {@ta4@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: green; --darkreader-inline-color:#72ff72;" data-darkreader-inline-color=""><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: orange; --darkreader-inline-color:#ffae1a;" data-darkreader-inline-color=""><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color: red; --darkreader-inline-color:#ff1a1a;" data-darkreader-inline-color=""><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>583140752</deployedseed>
    <deployedseed>557256621</deployedseed>
    <deployedseed>989695901</deployedseed>
    <deployedseed>73054750</deployedseed>
    <deployedseed>820204247</deployedseed>
    <deployedseed>233864065</deployedseed>
    <deployedseed>536065433</deployedseed>
    <deployedseed>1068285323</deployedseed>
    <deployedseed>515241887</deployedseed>
    <deployedseed>755167289</deployedseed>
    <deployedseed>1214970463</deployedseed>
    <deployedseed>264919088</deployedseed>
    <deployedseed>1090142959</deployedseed>
    <deployedseed>298905898</deployedseed>
    <deployedseed>389686625</deployedseed>
    <deployedseed>1350202280</deployedseed>
    <deployedseed>317891025</deployedseed>
    <deployedseed>990182177</deployedseed>
    <deployedseed>1528625301</deployedseed>
    <deployedseed>1028816909</deployedseed>
    <deployedseed>1459945248</deployedseed>
    <deployedseed>1835266606</deployedseed>
    <deployedseed>353455924</deployedseed>
    <deployedseed>721505679</deployedseed>
    <deployedseed>1693956234</deployedseed>
    <deployedseed>134914059</deployedseed>
    <deployedseed>267669552</deployedseed>
    <deployedseed>956410300</deployedseed>
    <deployedseed>321016559</deployedseed>
    <deployedseed>932307846</deployedseed>
    <deployedseed>1239909788</deployedseed>
    <deployedseed>2051122932</deployedseed>
    <deployedseed>1019584553</deployedseed>
    <deployedseed>1274150362</deployedseed>
    <deployedseed>1709961411</deployedseed>
    <deployedseed>1178262889</deployedseed>
    <deployedseed>1301756567</deployedseed>
    <deployedseed>1045064076</deployedseed>
    <deployedseed>1435912732</deployedseed>
    <deployedseed>233288130</deployedseed>
    <deployedseed>2078550814</deployedseed>
    <deployedseed>1974679635</deployedseed>
    <deployedseed>327619503</deployedseed>
    <deployedseed>902976291</deployedseed>
    <deployedseed>1531838061</deployedseed>
    <deployedseed>282226115</deployedseed>
    <deployedseed>989380741</deployedseed>
    <deployedseed>1458826222</deployedseed>
    <deployedseed>1542796246</deployedseed>
    <deployedseed>757220104</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+pp</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2-1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
