<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>16.2.5.2 Ratio test convergence</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Use the ratio test alone to examine the convergence of the series:</p>
<ol class="HELM_parts">
    <li>\(\displaystyle \sum_{p=1}^\infty \frac{e^{@a1@}}{(2p+1)^{p+1}}\). The ratio test is [[input:ans1]].[[validation:ans1]] [[feedback:prt1]]</li>
    <li>\(\displaystyle \sum_{p=1}^\infty \frac{p^{@a2@}}{p!}\). The ratio test is [[input:ans2]]. [[validation:ans2]] [[feedback:prt2]]</li>
    <li>\(\displaystyle \sum_{p=1}^\infty \frac{1}{\sqrt{p}}\). The ratio test is [[input:ans3]]. [[validation:ans3]] [[feedback:prt3]]</li>
    <li>\(\displaystyle \sum_{p=1}^\infty \frac{1}{({@a3@})^p}\). The ratio test is [[input:ans4]]. [[validation:ans4]] [[feedback:prt4]]</li>
    <li>\(\displaystyle \sum_{p=1}^\infty \frac{(-1)^{p+1}}{{@a4@}^p}\). The ratio test is [[input:ans5]]. [[validation:ans5]] [[feedback:prt5]]</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">
    <li>
        <p>The general term of this series is \(\displaystyle a_p = \frac{e^{@a1@}}{(2p + 1)^{p+1}}\). Then</p>
        \[\begin{aligned} \lambda = \lim_{p \to \infty} \frac{a_{p+1}}{a_p} &amp;= \lim_{p \to \infty} \frac{e^{@a1@}}{(2(p+1) + 1)^{p+2}} \cdot \frac{(2p + 1)^{p+1}}{e^{@a1@}} \\ &amp;= \lim_{p \to \infty} \frac{(2p + 1)^{p+1}}{(2(p+1) + 1)^{p+2}} \\ &amp;=
        \lim_{p \to \infty} \left(\frac{2p + 1}{2p + 3}\right)^{p+1} \cdot \frac{1}{2p + 3} \\ &amp;\leq 1 \cdot \lim_{p \to \infty} \frac{1}{2p + 3} \\ &amp;= 1 \cdot 0 = 0. \end{aligned}\]
        <p>But the limit is clearly nonnegative, so we conclude that \(\lambda = 0\). Since \(\lambda &lt; 1\), the series <b>converves</b>.</p>
    </li>
    <li>
        <p>The general term of this series is \(\displaystyle a_p = \frac{p^{@a2@}}{p!}\). Then</p>
        \[\begin{aligned} \lambda = \lim_{p \to \infty} \frac{a_{p+1}}{a_p} &amp;= \lim_{p \to \infty} \frac{(p+1)^{@a2@}}{(p+1)!} \cdot \frac{p!}{p^{@a2@}} \\ &amp;= \lim_{p \to \infty} \left(\frac{p+1}{p}\right)^{@a2@} \cdot \frac{1}{p+1} \\ &amp;= 1 \cdot
        0 = 0. \end{aligned}\]
        <p>Since \(\lambda &lt; 1\), the series <b>converges</b>.</p>
    </li>
    <li>
        <p>The general term of this series is \(\displaystyle a_p = \frac{1}{\sqrt{p}}\). Then</p>
        \[\begin{aligned} \lambda = \lim_{p \to \infty} \frac{a_{p+1}}{a_p} &amp;= \lim_{p \to \infty} \frac{1}{\sqrt{p+1}} \cdot \frac{\sqrt{p}}{1} \\ &amp;= \lim_{p \to \infty} \frac{\sqrt{p}}{\sqrt{p+1}} \\ &amp;= \lim_{p \to \infty} \sqrt{\frac{p}{p+1}} \\
        &amp;= 1. \end{aligned}\]
        <p>Since \(\lambda = 1\), the ratio test is <b>inconclusive</b>. Nevertheless, note that \(\displaystyle \frac{1}{\sqrt{p}} &gt; \frac{1}{p}\) for all \(p\), so our series diverges by comparison with the harmonic series.</p>
    </li>
    <li>
        <p>The general term of this series is \(\displaystyle a_p = \frac{1}{({@a3@})^p}\). Then</p>
        \[\begin{aligned} \lambda = \lim_{p \to \infty} \frac{a_{p+1}}{a_p} &amp;= \lim_{p \to \infty} \frac{1}{({@a3@})^{p+1}} \cdot \frac{({@a3@})^p}{1} \\ &amp;= \lim_{p \to \infty} \frac{({@a3@})^{p}}{({@a3@})^{p+1}} \\ &amp;= \lim_{p \to \infty} \frac{1}{{@a3@}}
        \\ &amp;= \frac{1}{{@a3@}} \approx {@decimalplaces(1/a3,3)@}. \end{aligned}\]
        <p>Since \(\lambda &gt; 1\), the series <b>diverges</b>.</p>
    </li>
  <li><p>The general term of this series is \(\displaystyle a_p = \frac{(-1)^{p+1}}{{@a4@}^p}\). Note that this is not always positive, so we cannot apply the ratio test. It is thus <b>inconclusive</b>. Nevertheless, this series is absolutely convergent, and thus convergent.</p></li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a1: rand(3)+4;
a2: rand([3,5,7]);
a3: rand([0.3, 0.4, 0.5, 0.6]);
a4: rand(5)+3;
ta1: [[1, true, "convergent"], [2, false, "divergent"], [3, false, "inconclusive"]]; 
ta2: [[1, true, "convergent"], [2, false, "divergent"], [3, false, "inconclusive"]];
ta3: [[1, false, "convergent"], [2, false, "divergent"], [3, true, "inconclusive"]];
ta4: [[1, false, "convergent"], [2, true, "divergent"], [3, false, "inconclusive"]];
ta5: [[1, false, "convergent"], [2, false, "divergent"], [3, true, "inconclusive"]];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ \sum_{p=1}^\infty \frac{e^{@a1@}}{(2p+1)^{p+1}}, \quad \sum_{p=1}^\infty \frac{p^{@a2@}}{p!}, \quad \sum_{p=1}^\infty \frac{1}{\sqrt{p}}\]
\[ \sum_{p=1}^\infty \frac{1}{({@a3@})^p}, \quad \sum_{p=1}^\infty \frac{(-1)^{p+1}}{{@a4@}^p} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1244974533</deployedseed>
    <deployedseed>1339730380</deployedseed>
    <deployedseed>1660903861</deployedseed>
    <deployedseed>1314674997</deployedseed>
    <deployedseed>1975382733</deployedseed>
    <deployedseed>1892784971</deployedseed>
    <deployedseed>1991916417</deployedseed>
    <deployedseed>1480292418</deployedseed>
    <deployedseed>291921780</deployedseed>
    <deployedseed>1413917380</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>3</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
