<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>2.6.3.3 Find centre and radious</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Obtain the radius and the coordinates of the centre for each of the following circles.</p>

<ol class="HELM_parts">
    <li>
        <p>{@c1@}</p>
        <p>\( C \big ( \) [[input:ans1]] [[validation:ans1]],[[input:ans2]] [[validation:ans2]] \( \big) \) [[feedback:prtC1]]</p>
        <p>\( r= \)[[input:ans3]] [[validation:ans3]] [[feedback:prtR1]]</p>
    </li>

    <li>
        <p>{@c2@}</p>
        <p>\( C \big ( \) [[input:ans4]] [[validation:ans4]],[[input:ans5]] [[validation:ans5]] \( \big) \) [[feedback:prtC2]]</p>
        <p>\( r= \)[[input:ans6]] [[validation:ans6]] [[feedback:prtR2]]</p>
    </li>

    <li>
        <p>{@c3@}</p>
        <p>\( C \big ( \) [[input:ans7]] [[validation:ans7]],[[input:ans8]] [[validation:ans8]] \( \big) \) [[feedback:prtC3]]</p>
        <p>\( r= \)[[input:ans9]] [[validation:ans9]] [[feedback:prtR3]]</p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">
    <li>
        <p>We are given a cirlce with equation \({@c1@}.\)</p>
        <p>Begin by completing the square separately on the \(x-\)terms and the \(y-\)terms:</p>
        <p>\( {@expand(exp1)@} = {@exp1@} \)</p>
        <p>\( {@expand(exp2)@} = {@exp2@} \)</p>
        <p>So the original equation becomes:</p>
        <p>\(\begin{align} {@exp1@}+{@exp2@} &amp;= 0 \\ {@(x-x1)^2@}+{@(y-y1)^2@} &amp;= {@r1^2@} \end{align}\)</p>
        <p>which represents a circle with centre \( ({@x1@}, {@y1@}) \) and radius {@r1@}.</p>
    </li>
  
  <li>
        <p>We are given a cirlce with equation \({@c2@}.\)</p>
        <p>Begin by completing the square separately on the \(x-\)terms and the \(y-\)terms:</p>
        <p>\( {@expand(exp3)@} = {@exp3@} \)</p>
        <p>\( {@expand(exp4)@} = {@exp4@} \)</p>
        <p>So the original equation becomes:</p>
        <p>\(\begin{align} {@exp3@}+{@exp4@} &amp;= {@aa@} \\ {@(x-x2)^2@}+{@(y-y2)^2@} &amp;= {@r2^2@} \end{align}\)</p>
        <p>which represents a circle with centre \( ({@x2@}, {@y2@}) \) and radius {@r2@}.</p>
    </li>
  
  <li>
        <p>We are given a cirlce with equation \({@c3@}.\)</p>
        <p>Begin by completing the square separately on the \(x-\)terms and the \(y-\)terms:</p>
        <p>\( {@expand(exp5)@} = {@exp5@} \)</p>
        <p>\( {@expand(exp6)@} = {@exp6@} \)</p>
        <p>So the original equation becomes:</p>
        <p>\(\begin{align} {@exp5@}+{@exp6@} - {@bb@} &amp;= 0 \\ {@(x-x3)^2@}+{@(y-y3)^2@} &amp;= {@r3^2@} \end{align}\)</p>
        <p>which represents a circle with centre \( ({@x3@}, {@y3@}) \) and radius {@r3@}.</p>
    </li>

</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021120900</text>
    </stackversion>
    <questionvariables>
      <text>ordergreat (x,y);
/* Define centres*/
x1: rand(3)+4;
y1: -(x1+1);
x2: 1;
y2: rand([-2,-3,-4]);
x3: rand([2,3,4])
y3: 0;

sipm:false;
/* Completed square*/ 
exp1: (x-x1)^2-x1^2;
exp2: (y-y1)^2-y1^2;
c1: expand(exp1+exp2)=0;
r1: ev(sqrt(x1^2+y1^2),simp);

exp3: (x-x2)^2-x2^2;
exp4: (y-y2)^2-y2^2;
aa: ev(x2+y2+rand(5)+5, simp);
c2: expand(exp3+exp4)=aa;
r2: ev(sqrt(aa+x2^2+y2^2),simp);

exp5: (x-x3)^2-x3^2;
exp6: (y-y3)^2-y3^2;
bb: ev(rand(5)+10, simp);
c3: expand(exp5+exp6)-bb=0;
r3: ev(sqrt(bb+x3^2+y3^2),simp);
simp:true;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[{@c1@}, \rightarrow C({@x1@}, {@y1@}), \, r={@r1@} \]
\[{@c2@}, \rightarrow C({@x2@}, {@y2@}), \, r={@r2@} \]
\[{@c3@}, \rightarrow C({@x3@}, {@y3@}), \, r={@r3@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>x1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>y1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>r1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>x2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>y2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>r2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>x3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>y3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>r3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prtC1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>x1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtC1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtC1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( x\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>y1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( y\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>y1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( y\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtC2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>x2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtC2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtC2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( x\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>y2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( y\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>y2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC2-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( y\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtC3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>x3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtC3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtC3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( x\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>y3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( y\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>y3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC3-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">The \( y\)&nbsp; coordinate is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtR1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>r1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtR1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtR1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtR2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>r2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtR2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtR2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtR3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>r3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtR3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtR3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1114960855</deployedseed>
    <deployedseed>1344492518</deployedseed>
    <deployedseed>1519022430</deployedseed>
    <deployedseed>1512808967</deployedseed>
    <deployedseed>2116223412</deployedseed>
    <deployedseed>1066650682</deployedseed>
    <deployedseed>1677119683</deployedseed>
    <deployedseed>1535177716</deployedseed>
    <deployedseed>468044899</deployedseed>
    <deployedseed>1388497302</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>x1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>y1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>r1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>x2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>y2</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>r2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>x3</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>y3</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>r3</value>
      </testinput>
      <expected>
        <name>prtC1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtC1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtC2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtC2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtC3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtC3-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(x1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>y1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(r1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ev(x2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>y2</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ev(r2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ev(x3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>y3</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ev(r3+1,simp)</value>
      </testinput>
      <expected>
        <name>prtC1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC1-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtC2</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC2-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtC3</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtR1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtR2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtR2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtR3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtR3-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>x1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(y1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>r1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>x2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ev(y2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>r2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>x3</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ev(y3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>r3</value>
      </testinput>
      <expected>
        <name>prtC1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtC2</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtC3</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC3-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtR1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(x1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(y1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>r1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ev(x2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ev(y2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>r2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ev(x3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ev(y3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>r3</value>
      </testinput>
      <expected>
        <name>prtC1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC1-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtC2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC2-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtC3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prtC3-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtR1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtR3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtR3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
