<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.3.4.8 Angle between vectors (NOT IN QUIZ)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>Find the angle between the vectors <span class="math inline">\(\underline{p}={@p_i@}\underline{i}+\underline{j}+\underline{k}\)</span> and <span class="math inline">\(\underline{q}={@q_i@}\underline{i}+{@q_j@}\underline{j}+2\underline{k}\)</span>.</p><p></p><p>(Express your answer in degrees and round your answer up to one decimal place)<br></p><p>[[input:ans1]]\(^\circ\) (degrees) [[validation:ans1]]</p><p>[[feedback:prt1]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>First, we compute the modulus of&nbsp;\(\underline{ p}\) and \(\underline{ q}\):&nbsp;&nbsp;\(|\underline{p}| = \sqrt{{@p_i@}^2+1^2 +1^2}=\sqrt{{@p_i^2 +2@}}\),&nbsp;\(|\underline{q}| =&nbsp;\sqrt{{@q_i@}^2+{@q_j@}^2 + 2^2}=\sqrt{{@q_i^2 +q_j^2 +4@}}\). Then, we compute \( \underline{p}\cdot \underline{q} = ({@p_i@})({@q_i@})+(1)({@q_j@}) + (1)(2)={@p_i*q_i + q_j + 2@}\). Substituting these values for&nbsp;\({\underline p}\cdot {\underline q},|\underline{p}|\)&nbsp;and&nbsp;\(|\underline{q}|\)&nbsp;into the formula for the scalar product \(&nbsp;\underline{p}\cdot\underline{q}=|\underline{p}|\,|\underline{q}| \cos\theta \) we have:</p><p>\(&nbsp;{@p_i*q_i + q_j + 2@} =&nbsp;\sqrt{{@p_i^2 +2@}} \sqrt{{@q_i^2 +q_j^2 +4@}} \cos \theta&nbsp;\)</p><p>which is equivalent to</p><p>\( \cos \theta = \dfrac{{@p_i*q_i + q_j + 2@}}{\sqrt{{@(p_i^2 +2)*(q_i^2 +q_j^2 +4)@}}} \)<br></p><p>finally having that</p><p>\( \theta = \cos^{-1} \left(&nbsp;\dfrac{{@p_i*q_i + q_j + 2@}}{\sqrt{{@(p_i^2 +2)*(q_i^2 +q_j^2 +4)@}}} \right) = {@ta@}^\circ\)</p><p></p><p><br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>p_i: 2 + rand(3);
q_i: 2 + rand(3);
q_j: 2 + rand(3);
argCos: (p_i*q_i + q_j + 2)/sqrt((p_i^2 + 2)*(q_i^2 + q_j^2 +4));
invCos: float(acos(argCos));
ta: 0.1*round(10*float((invCos*180)/%pi));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \cos^{-1}({@argCos@}) = {@ta@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:1</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>277299630</deployedseed>
    <deployedseed>512681250</deployedseed>
    <deployedseed>1319515272</deployedseed>
    <deployedseed>1345830566</deployedseed>
    <deployedseed>387901401</deployedseed>
    <deployedseed>445328152</deployedseed>
    <deployedseed>1059255342</deployedseed>
    <deployedseed>1712766271</deployedseed>
    <deployedseed>1217464304</deployedseed>
    <deployedseed>179716181</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>32</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
