<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.3.4.4 Perpendicular vectors</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>If <span class="math inline">\(\underline{a}\)</span> and <span class="math inline">\(\underline{b}\)</span> are perpendicular, simplify <span class="math inline">\((\underline{a}-{@b1@}\underline{b})\cdot(3\underline{a}+{@b2@}\underline{b})\)</span>.</p><p>Please write your answer as an algebraic expression in terms of \(a^2\) and \(b^2\) recalling that \(\underline{a} \cdot \underline{a} = a^2\) and \(\underline{b} \cdot \underline{b} = b^2\). Write the modulus of \(\underline{a}\) and \(\underline{b}\) as <code>a</code> and <code>b</code>, respectively.</p><p>[[input:ans1]] [[validation:ans1]]</p><p>[[feedback:prt1]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We proceed to compute&nbsp;\((\underline{a}-{@b1@}\underline{b})\cdot(3\underline{a}+{@b2@}\underline{b})\) as follows:</p><p>\(&nbsp; (\underline{a}-{@b1@}\underline{b})\cdot(3\underline{a}+{@b2@}\underline{b}) = 3 \underline{a} \cdot \underline{a} +&nbsp;{@b2@}&nbsp;\underline{a} \cdot \underline{b} +(3)(-{@b1@}) \underline{a} \cdot \underline{b} + (-{@b1@})({@b2@}) \underline{b} \cdot \underline{b}&nbsp; \)</p><p>As \(\underline{a} \) and \(\underline{b}\) are perpendicular, their dot product is zero, thus:</p><p>\(&nbsp; (\underline{a}-{@b1@}\underline{b})\cdot(3\underline{a}+{@b2@}\underline{b}) = 3 \underline{a} \cdot \underline{a}&nbsp; + (-{@b1@})({@b2@}) \underline{b} \cdot \underline{b} = 3a^2 - {@b1*b2@} b^2 \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>b1: 2 + rand(4);
b2: 3 + rand(4);
ta: 3*a^2 - (b1*b2)*b^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ (\underline{a}-{@b1@}\underline{b})\cdot(3\underline{a}+{@b2@}\underline{b}) = {@ta@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1921675788</deployedseed>
    <deployedseed>893515743</deployedseed>
    <deployedseed>1915881358</deployedseed>
    <deployedseed>1648499619</deployedseed>
    <deployedseed>503684668</deployedseed>
    <deployedseed>110173693</deployedseed>
    <deployedseed>674098965</deployedseed>
    <deployedseed>2138496274</deployedseed>
    <deployedseed>1619847251</deployedseed>
    <deployedseed>11684619</deployedseed>
    <deployedseed>738567055</deployedseed>
    <deployedseed>1983904085</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>2*a^2 -4*b^2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
