<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.5.2.2 Direction ratios and cosines in three dimensions, Exercise</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>In the following question, note that you can use the command <code>sqrt(5)</code> for the square root of \(5\), for example.</p> 
<p>Points \(A\) and \(B\) have position vectors \({@stackvector(a)@}={@A1@}\), and \({@stackvector(b)@}= {@B1@}\) respectively. Find</p>
<ol class="HELM_parts">
 <li><p> \(\overrightarrow{AB}=\)[[input:AB11]]  [[validation:AB11]] {@stackvector(i)@} \(+\) [[input:AB12]] [[validation:AB12]] {@stackvector(j)@} \(+\) [[input:AB13]] [[validation:AB13]]{@stackvector(k)@}</p>
<p>[[feedback:diff]]</p></li>

<li> <p>\(|\overrightarrow{AB}|=\)[[input:magAB1]] [[validation:magAB1]]</p>
<p>[[feedback:mag]]</p></li>

<li><p> The direction ratios of \(\overrightarrow{AB}\): [[input:dr1]] [[validation:dr1]] : [[input:dr2]] [[validation:dr2]] : [[input:dr3]] [[validation:dr3]]</p>
<p>[[feedback:DR]]</p> </li>

<li><p>The direction cosines of \(\overrightarrow{AB}\), \((\ell,m,n)=\) \( \Big( \) [[input:dc1]] [[validation:dc1]], [[input:dc2]] [[validation:dc2]], [[input:dc3]] [[validation:dc3]] \( \Big) \)<br></p>
<p>[[feedback:DC]]</p></li>

<li><p>\(\ell^2+m^2+n^2=\) [[input:magCos]]. [[validation:magCos]]</p>
<p>[[feedback:sumDC]]</p></li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Points \(A\) and \(B\) have position vectors \({@stackvector(a)@}={@A1@}\), and \({@stackvector(b)@}= {@B1@}\) respectively.<br></p><p>We have that</p>
<ol class="HELM_parts">
 <li><p> \(\overrightarrow{AB}= \underline{b} - \underline{a} = {@R1@}\)</p></li>
<li><p> \(\lvert\overrightarrow{AB}\rvert={@magr1@}\)</p></li>
<li><p>The direction ratios of \(\overrightarrow{AB}\) are {@r1[1,1]@}&nbsp;:&nbsp;{@r1[2,1]@}&nbsp;:&nbsp;{@r1[3,1]@}</p></li>
<li><p> The direction cosines of \(\overrightarrow{AB}\) are 
\[(\ell,m,n) = \left({@l1@}, {@m1@}, {@n1@}\right).\]</p></li>
<li><p>We have \(\ell^2+m^2+n^2=1\). Can you see why this is true in general?</p></li></ol>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
a1: matrix([5],[5],[3]) - rand(matrix([10],[10],[3]));
b1: matrix([5],[5],[8]) - rand(matrix([10],[10],[2]));
r1:  b1-a1;
magr1: sqrt(dotproduct(r1,r1));
u1: matrix([i],[j],[k]);
A1: texboldatoms(dotproduct(a1,u1));
B1: texboldatoms(dotproduct(b1,u1));
R1:  texboldatoms(dotproduct(r1,u1));
l1: (r1[1,1])/(magr1);
m1: (r1[2,1])/(magr1);
n1: (r1[3,1])/(magr1);
dc: matrix([l1],[m1],[n1]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>The direction ratios for the vector \(\mathbf{r} = {@R1@}\) are \({@r1[1,1]@}\,:\,{@r1[2,1]@}\,:\,{@r1[3,1]@}\), and its direction cosines are \(l = {@l1@}\), \(m = {@m1@}\) and \(n = {@n1@}\).</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>symi</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>AB11</name>
      <type>algebraic</type>
      <tans>r1[1,1]</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>AB12</name>
      <type>algebraic</type>
      <tans>r1[2,1]</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>AB13</name>
      <type>algebraic</type>
      <tans>r1[3,1]</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>dc1</name>
      <type>algebraic</type>
      <tans>l1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>dc2</name>
      <type>algebraic</type>
      <tans>m1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>dc3</name>
      <type>algebraic</type>
      <tans>n1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>dr1</name>
      <type>algebraic</type>
      <tans>r1[1,1]</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>dr2</name>
      <type>algebraic</type>
      <tans>r1[2,1]</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>dr3</name>
      <type>algebraic</type>
      <tans>r1[3,1]</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>magAB1</name>
      <type>algebraic</type>
      <tans>magr1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>magCos</name>
      <type>algebraic</type>
      <tans>1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>DC</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>student_r: matrix([dc1],[dc2],[dc3]);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>dc1</sans>
        <tans>l1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>VEC-DC-X_COR</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>VEC-DC-X_INCOR</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the direction cosine in the \(x\)-direction appears to be incorrect.</p>
<p>Remember the formula for calculating the direction cosine (draw a diagram if this helps!), and try again.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>dc2</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>VEC-DC-Y_COR</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>VEC-DC-Y_INCOR</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the direction cosine in the \(y\)-direction appears to be incorrect.</p>
<p>Remember the formula for calculating the direction cosine (draw a diagram if this helps!), and try again.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>dc3</sans>
        <tans>n1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>VEC-DC-CORRECT</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>VEC-DC-Z_INCOR</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the direction cosine in the \(z\)-direction appears to be incorrect.</p>
<p>Remember the formula for calculating the direction cosine (draw a diagram if this helps!), and try again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>diff</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>AB11,AB12,AB13</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>matrix([AB11],[AB12],[AB13])</sans>
        <tans>r1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>diff-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>diff-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[Your expression for the difference seems to contain an error.<br>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>DR</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>student_r: [dr1*x=dr2*y, dr2*y=dr3*z];
taEqRat: [r1[1,1]*x=r1[2,1]*y, r1[2,1]*y=r1[3,1]*z];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>SysEquiv</answertest>
        <sans>student_r</sans>
        <tans>taEqRat</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>VEC-DR-TRUE</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>VEC-DR-FALSE</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Remember how we obtain the direction ratios for a three-dimensional vector from its expression in terms of \(\mathbf{i}\),&nbsp;\(\mathbf{j}\) and&nbsp;\(\mathbf{k}\) components, and try again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>mag</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>magAB1</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>magAB1</sans>
        <tans>magr1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>mag-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>mag-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your expression for the magnitude appears to be incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>sumDC</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>magCos</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>magCos</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>sumDC-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done. Do you expect this formula to hold in general?<br></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>sumDC-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>This answer is incorrect. Please check your calculations.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>798628930</deployedseed>
    <deployedseed>804788146</deployedseed>
    <deployedseed>782382927</deployedseed>
    <deployedseed>500428150</deployedseed>
    <deployedseed>509290389</deployedseed>
    <deployedseed>1938362763</deployedseed>
    <deployedseed>1975170051</deployedseed>
    <deployedseed>1003751379</deployedseed>
    <deployedseed>1695724250</deployedseed>
    <deployedseed>275922861</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value>r1[1,1]</value>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value>r1[2,1]</value>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value>r1[3,1]</value>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value>l1</value>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value>m1</value>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value>n1</value>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value>r1[1,1]</value>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value>r1[2,1]</value>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value>r1[3,1]</value>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value>magr1</value>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value>1</value>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-DC-CORRECT</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>diff-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>mag-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>sumDC-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value>ev(r1[1,1]+1,simp)</value>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value>ev(r1[2,1]+1,simp)</value>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value>ev(r1[3,1]+1,simp)</value>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value>l1 + 1</value>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value>m1</value>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value>n1</value>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value>r1[1,1]</value>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value>r1[2,1]</value>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value>r1[3,1]</value>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value>ev(magr1+1, simp)</value>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value>5</value>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>VEC-DC-X_INCOR</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>diff-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>mag-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>sumDC-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value>l1</value>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value>m1+1</value>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value>n1</value>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value>r1[1,1]</value>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value>r1[2,1]</value>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value>r1[3,1]</value>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value/>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value/>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>VEC-DC-Y_INCOR</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value>l1</value>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value>m1</value>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value>n1+1</value>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value>r1[1,1]</value>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value>r1[2,1]</value>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value>r1[3,1]</value>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value/>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value/>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>VEC-DC-Z_INCOR</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value>l1</value>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value>m1</value>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value>n1</value>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value>20</value>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value>20</value>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value>20</value>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value/>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value/>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-DC-CORRECT</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>VEC-DR-FALSE</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value>r1[1,1]</value>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value>r1[2,1]</value>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value>r1[3,1]</value>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value/>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value/>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value/>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value/>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value>magr1</value>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value/>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>diff-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>mag-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description/>
      <testinput>
        <name>AB11</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB12</name>
        <value/>
      </testinput>
      <testinput>
        <name>AB13</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc1</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc2</name>
        <value/>
      </testinput>
      <testinput>
        <name>dc3</name>
        <value/>
      </testinput>
      <testinput>
        <name>dr1</name>
        <value/>
      </testinput>
      <testinput>
        <name>dr2</name>
        <value/>
      </testinput>
      <testinput>
        <name>dr3</name>
        <value/>
      </testinput>
      <testinput>
        <name>magAB1</name>
        <value/>
      </testinput>
      <testinput>
        <name>magCos</name>
        <value>magCos</value>
      </testinput>
      <expected>
        <name>DC</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>diff</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>DR</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>mag</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>sumDC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>sumDC-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
