<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.1.2.9 Solve a linear equation_7</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Solve the following equations.</p><ol class="HELM_parts">

<li><p>\( \displaystyle {@a7/m@} = {@b7@} \), therefore \(m=\) [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</p></li>

<li><p>\( \displaystyle{{@a8/m@} = {@b8/(m + 1)@}} \), therefore \(m=\) [[input:ans8]] [[validation:ans8]] [[feedback:prt8]]</p></li>

</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">

    <li>
        <p>\( \displaystyle {@a7/m@} = {@b7@} \Rightarrow {@b7*m@} = {@a7@}\), therefore \(m= \displaystyle {@ta7@}.\)</p>
    </li>

    <li>
        <p>\( \displaystyle{{@a8/m@} = {@b8/(m + 1)@}} \Rightarrow {@b8*m@} = {@a8@}(m+1) \), therefore \(m={@ta8@}.\)</p>
    </li>

</ol>]]></text>
    </generalfeedback>
    <defaultgrade>12</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>a7: rand(6)+2; 
b7: rand(3)+2;
eq7: a7/m = b7;
tag: first(solve(eq7, m));
ta7: rhs(tag);

a8: rand(6)+1; 
b8: rand_with_prohib(2,5,[a8]);   /* if a8=b8 then no solution 0*m = n*/
eq8: a8/m = b8/(m+1);
tah: first(solve(eq8, m));
ta8: rhs(tah);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@eq7@} \Rightarrow m= {@ta7@}, \quad {@eq8@} \Rightarrow m= {@ta8@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1378417595</deployedseed>
    <deployedseed>2117228019</deployedseed>
    <deployedseed>163871450</deployedseed>
    <deployedseed>273587298</deployedseed>
    <deployedseed>1855518300</deployedseed>
    <deployedseed>619679884</deployedseed>
    <deployedseed>764816665</deployedseed>
    <deployedseed>333953847</deployedseed>
    <deployedseed>2117844009</deployedseed>
    <deployedseed>102512895</deployedseed>
    <deployedseed>611856798</deployedseed>
    <deployedseed>872458472</deployedseed>
    <deployedseed>1904847692</deployedseed>
    <deployedseed>1616187571</deployedseed>
    <deployedseed>1383831772</deployedseed>
    <deployedseed>684146219</deployedseed>
    <deployedseed>1419092282</deployedseed>
    <deployedseed>1910472797</deployedseed>
    <deployedseed>753457569</deployedseed>
    <deployedseed>81838332</deployedseed>
    <deployedseed>1034374355</deployedseed>
    <deployedseed>365172999</deployedseed>
    <deployedseed>1050934026</deployedseed>
    <deployedseed>477685549</deployedseed>
    <deployedseed>872793546</deployedseed>
    <deployedseed>2073009736</deployedseed>
    <deployedseed>1580779358</deployedseed>
    <deployedseed>506990993</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans7</name>
        <value>ta7-1</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8-1</value>
      </testinput>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt7-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt8-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
