<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>43.2.1.3 Spearman R</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>As a result of two tests given to 10 students studying laboratory safety, the students were placed in the following class order.</p>
<p>&nbsp;</p>
<table style="height: 437px;" border="1" width="330">
<tbody>
<tr>
<td style="width: 91.3333px; text-align: center;"><b>Student</b></td>
<td style="width: 135.333px; text-align: center;"><b>Test 1</b></td>
<td style="width: 80.6667px; text-align: center;"><b>Test 2</b></td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">A</td>
<td style="width: 135.333px; text-align: center;">\({@A1@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B1@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">B</td>
<td style="width: 135.333px; text-align: center;">\({@A2@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B2@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">C</td>
<td style="width: 135.333px; text-align: center;">\({@A3@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B3@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">D</td>
<td style="width: 135.333px; text-align: center;">\({@A4@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B4@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">E</td>
<td style="width: 135.333px; text-align: center;">\({@A5@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B5@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">F</td>
<td style="width: 135.333px; text-align: center;">\({@A6@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B6@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">G</td>
<td style="width: 135.333px; text-align: center;">\({@A7@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B7@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">H</td>
<td style="width: 135.333px; text-align: center;">\({@A8@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B8@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">I</td>
<td style="width: 135.333px; text-align: center;">\({@A9@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B9@}\)</td>
</tr>
<tr>
<td style="width: 91.3333px; text-align: center;">J</td>
<td style="width: 135.333px; text-align: center;">\({@A10@}\)</td>
<td style="width: 80.6667px; text-align: center;">\({@B10@}\)</td>
</tr>
</tbody>
</table>
<p>&nbsp;</p>
<p>Use Spearman's \(R\) to discuss the consistency of their performances, giving your answers to 3 decimal places. What is the conclusion?&nbsp;</p><p><br></p><p>Can you think of a meaningful comment regarding the two tests as a means of assessing laboratory safety?</p><p><br></p><p>Take the hypotheses to be</p>
\[ H_0 : R = 0 \hspace{2cm} H_1 : R \neq 0. \]
<p></p>
<ol class="HELM_parts">
<li><p> \(R\) = [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p></li>

<li><p> \(R_{test}\) = [[input:ans2]][[validation:ans2]][[feedback:prt2]]</p></li>
<li><p> The critical value for a two-sided \(t\)-test at the \(5\%\) level read off from tables is [[input:ans3]], therefore we [[input:ans4]] the null hypothesis at the \(5\%\) level.</p></li>
<p></p><p>[[validation:ans3]][[feedback:prt3]]</p>
<p>[[validation:ans4]][[feedback:prt4]]</p>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>&nbsp;</p>
<p>Setting up a spreadsheet to do the calculations gives:</p>
<p>&nbsp;</p>
<table style="height: 286px;" border="1" width="550">
<tbody>
<tr>
<td style="width: 122px;"><b>Test 1</b></td>
<td style="width: 122px;"><b>Test 2</b></td>
<td style="width: 134px;">\(D\)</td>
<td style="width: 143.333px;">\(D^2\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A1@}\)</td>
<td style="width: 122px;">\({@B1@}\)</td>
<td style="width: 134px;">\({@D1@}\)</td>
<td style="width: 143.333px;">\({@D1^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A2@}\)</td>
<td style="width: 122px;">\({@B2@}\)</td>
<td style="width: 134px;">\({@D2@}\)</td>
<td style="width: 143.333px;">\({@D2^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A3@}\)</td>
<td style="width: 122px;">\({@B3@}\)</td>
<td style="width: 134px;">\({@D3@}\)</td>
<td style="width: 143.333px;">\({@D3^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A4@}\)</td>
<td style="width: 122px;">\({@B4@}\)</td>
<td style="width: 134px;">\({@D4@}\)</td>
<td style="width: 143.333px;">\({@D4^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A5@}\)</td>
<td style="width: 122px;">\({@B5@}\)</td>
<td style="width: 134px;">\({@D5@}\)</td>
<td style="width: 143.333px;">\({@D5^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A6@}\)</td>
<td style="width: 122px;">\({@B6@}\)</td>
<td style="width: 134px;">\({@D6@}\)</td>
<td style="width: 143.333px;">\({@D6^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A7@}\)</td>
<td style="width: 122px;">\({@B7@}\)</td>
<td style="width: 134px;">\({@D7@}\)</td>
<td style="width: 143.333px;">\({@D7^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A8@}\)</td>
<td style="width: 122px;">\({@B8@}\)</td>
<td style="width: 134px;">\({@D8@}\)</td>
<td style="width: 143.333px;">\({@D8^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A9@}\)</td>
<td style="width: 122px;">\({@B9@}\)</td>
<td style="width: 134px;">\({@D9@}\)</td>
<td style="width: 143.333px;">\({@D9^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">\({@A10@}\)</td>
<td style="width: 122px;">\({@B10@}\)</td>
<td style="width: 134px;">\({@D10@}\)</td>
<td style="width: 143.333px;">\({@D10^2@}\)</td>
</tr>
<tr>
<td style="width: 122px;">&nbsp;</td>
<td style="width: 122px;">&nbsp;</td>
<td style="width: 134px;">&nbsp;</td>
<td style="width: 143.333px;">\( \sum D^2 = {@sumd2@} \)</td>
</tr>
<tr>
<td style="width: 122px;">&nbsp;</td>
<td style="width: 122px;">\( R = {@R@} \)</td>
<td style="width: 134px;">\( R_{test} = {@Rtest@} \)</td>
<td style="width: 143.333px;">&nbsp;</td>
</tr>
</tbody>
</table>
<p>Substituting in the formula we developed for \(R\) gives the result&nbsp;</p>
<p>&nbsp;\[ r =1 - \dfrac{6 \sum D^2}{n(n^2 - 1)} = 1 - \dfrac{6 \times {@sumd2@}}{ {@n@}({@n^2@} - 1) } = {@ta1@}\]</p>
<p>and then</p>
<p>\[R_{test} = \dfrac{R\sqrt{n-2}}{\sqrt{1 - R^2}} = {@ta2@}\]</p>
<p>Hence, since the critical value for a two-sided \(t\)-test at the \(5\%\) level read off from tables is \(2.306\) we see that since \( {@ta2@} &lt; 2.306 \) we <strong>cannot</strong> reject the null hypothesis that there is no correlation between the results. This implies that the performances of the students on the tests may not be related and we should question at least one of the tests as a means of assessing laboratory safety. One could, of course, question the usefulness of both tests!</p>
<p>&nbsp;</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024072400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[n:10;
A1:2;
A2:4
A3:3;
A4:5;
A5:1;
A6:6;
A7:8;
A8:7;
A9:9;
A10:10;

B1:rand([3,4]);
B2:rand([5,6]);
B3:rand([7,8]);
B4:rand([9,10]);
B5:rand([10,11]);
B6:rand([2,4]);
B7:rand([6,7]);
B8:rand([8,10]);
B9:rand([4,5]);
B10:rand([1]);

D1:A1 - B1;
D2:A2 - B2;
D3:A3 - B3;
D4:A4 - B4;
D5:A5 - B5;
D6:A6 - B6;
D7:A7 - B7;
D8:A8 - B8;
D9:A9 - B9;
D10:A10 - B10;

sumd2:D1^2 + D2^2 + D3^2 + D4^2 + D5^2 + D6^2 + D7^2 + D8^2 + D9^2 + D10^2;

R:decimalplaces( 1 - (6*sumd2)/(n*(n^2 - 1)) ,6);
Rtest:decimalplaces( R*sqrt(n-2)/sqrt(1-R^2),6);
ta1:decimalplaces(R,3);
ta2:decimalplaces(Rtest,3);
Ta2:decimalplaces( ta1*sqrt(n-2)/sqrt(1-ta1^2),3);
ta3:[[1.860,false],[2.228, false],[2.306,true]];
ta4:[["can",false],["cannot", true]];

]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@ta1@}, {@ta2@}, {@Ta2@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:3, maxdp:3</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:3, maxdp:3</options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>abs(ans2)</sans>
        <tans>abs(ta2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>abs(ans2)</sans>
        <tans>abs(Ta2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans3]</sans>
        <tans>mcq_correct(ta3)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>String</answertest>
        <sans>ans4</sans>
        <tans>first(mcq_correct(ta4))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1704048355</deployedseed>
    <deployedseed>1157023282</deployedseed>
    <deployedseed>443985908</deployedseed>
    <deployedseed>1027036256</deployedseed>
    <deployedseed>1580066887</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>2.306</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value><![CDATA["cannot"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>2.228</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value><![CDATA["can"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>1.86</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value><![CDATA["can"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>2.306</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value><![CDATA["cannot"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
