<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.1.1.4 line integral with trig scalar field</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find
<ol class="HELM_parts_inline">
<li>\(\displaystyle \int_C F(x,y)\ \mathrm{d}x,\)</li>
<li>\(\displaystyle \int_C F(x,y)\ \mathrm{d}y,\)</li>
<li>\(\displaystyle \int_C F(x,y)\ \mathrm{d}s\)</li>
</ol>
where \(F (x, y) = {@aa@}\sin2x\) and \(C\) is the curve \(y = \sin x\) from \((0, 0 )\) to \(\dfrac{\pi}{2}, 1)\).</p>

<p>\(\displaystyle \int_C F(x,y)\ \mathrm{d}x =\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>\(\displaystyle \int_C F(x,y)\ \mathrm{d}y =\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>\(\displaystyle \int_C F(x,y)\ \mathrm{d}s =\)[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We know that \(F (x, y) = {@aa@}\sin2x\) and \(C\) is the curve \(y = \sin x\) from \((0, 0 )\) to \(\dfrac{\pi}{2}, 1)\), and aim to find:</p>
<p>(a)\(\displaystyle \int_C F(x,y)\ \mathrm{d}x,\) (b)\(\displaystyle \int_C F(x,y)\ \mathrm{d}y,\) and (c)\(\displaystyle \int_C F(x,y)\ \mathrm{d}s\).</p>

<p><strong>(a)</strong>
\[\begin{aligned}
\int_C F(x,y)\ \mathrm{d}x &= \int^{\frac{\pi}{2}}_0 {@aa@}\sin2x  \mathrm{d}x\\
&= {@aa@}\left[\frac{-\cos2x}{2}\right]^{\frac{\pi}{2}}_0\\
&= {@aa@} \left(\frac{1}{2} - \left[-\frac{1}{2}\right]\right)\\
&={@aa@}
\end{aligned}\]</p>

<p><strong>(b)</strong> We will aim to solve by rewriting the integral to be in terms of \(x\):
\[\begin{aligned}
I_2 = \frac{\mathrm{d}y}{\mathrm{d}x} &= \cos x \\
\mathrm{d}y &= \cos x \mathrm{d}x\\
&\implies \\
\int_C F(x,y)\ \mathrm{d}y &= \int^\frac{\pi}{2}_0 {@aa@} \sin2x\cos x \mathrm{d}x \\
&= \int^\frac{\pi}{2}_0 {@aa*2@} \sin x\cos^2 x \mathrm{d}x
\end{aligned}\]</p>

<p>let \(u =\cos x \implies \frac{\mathrm{d}u}{\mathrm{d}x}=-\sin x\)
\[\begin{aligned}
I_2 & ={@-aa*2@}\int^{x=\frac{\pi}{2}}_{x=0} u^2 \mathrm{d}u\\
&= {@-aa*2@}\left[\frac{u^3}{3}\right]^{x=\frac{\pi}{2}}_{x=0}\\
&={@-aa*2@}\left[\frac{\cos^3 x}{3}\right]^{\frac{\pi}{2}}_{0}\\
&= {@ta2@}
\end{aligned}\]</p>


<p><strong>(c)</strong> we again aim to express everything in terms of \(x\)
\[\begin{aligned}
\mathrm{d}s &= \sqrt{1+(\frac{\mathrm{d}y}{\mathrm{d}x})^2}\mathrm{d}x\\
&= \sqrt{1+\cos^2 x}\mathrm{d}x
\end{aligned}\]

<p>Hence \( \int_C F(x,y)\ \mathrm{d}s =\int^{\frac{\pi}{2}}_0 \sin 2 x\sqrt{1+\cos^2 x}\mathrm{d}x \)</p>
<p>let \(u = 1+cos^2 \implies \frac{\mathrm{d}u}{\mathrm{d}x} = -\sin 2x \)</p>
<p>for  \(x=0,\ u=2 \)</p>
<p>for \(x=\frac{\pi}{2},\ u=1\)</p>

<p>So \[\begin{aligned}
\int_C F(x,y)\ \mathrm{d}s &= -\int^{1}_2\sqrt{u}\mathrm{d}x\\
&=\left[\frac{2}{3}u^\frac{3}{2}\right]^2_1\\
&={@ta3@}
\end{aligned}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(8)+2;

ta1:aa;
ta2:aa*2/3;
ta3:aa*(2*sqrt(2)-1)*2/3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\(F (x, y) = {@aa@}\sin2x, \quad \displaystyle \int_C F(x,y)\ \mathrm{d}x = {@ta1@}, \quad \displaystyle \int_C F(x,y)\ \mathrm{d}y = {@ta2@}, \quad \displaystyle \int_C F(x,y)\ \mathrm{d}s = {@ta3@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>onum</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2128011308</deployedseed>
    <deployedseed>1014071468</deployedseed>
    <deployedseed>1985565443</deployedseed>
    <deployedseed>991670240</deployedseed>
    <deployedseed>663151203</deployedseed>
    <deployedseed>399732366</deployedseed>
    <deployedseed>370624964</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
