<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.1.2.6 scalar integral of vector field task</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the vector field \(\displaystyle \underline{G}=x\underline{i}+({@aa@}x-y)\underline{j} 
\)
Let \(C_1\) and \(C_2\) be the curves from \(O=(0,0,0)\) to \(A=(1,1,1)\), 
given by 
\[\begin{eqnarray*} 
C_1&:&x=t,\qquad y=t,\qquad z=t\qquad (0\leq t\leq 1)\\ 
C_2&:&x=t^2,\qquad y=t,\qquad z=t^2\qquad (0\leq t\leq 1) 
\end{eqnarray*} \]

<ol class="HELM_parts">
<li><p>Evaluate the scalar integral \(\displaystyle\int_C{\underline G}\cdot \underline{dr}\) of each vector field along each path.</p>
<p>\(\displaystyle\int_{C_1}{\underline G}\cdot \underline{dr} =\)[[input:ans1]] [[validation:ans1]]</p>
<p>\(\displaystyle\int_{C_2}{\underline G}\cdot \underline{dr} =\)[[input:ans2]] [[validation:ans2]][[feedback:prt1]]</p>
</li>

<li><p>Find the value of  \(\displaystyle\oint_C \underline{G} \cdot  \underline{dr}\) where \(C\) is the closed path along \(C_1\) from \(O\) to \(A\) and back along \(C_2\) from \(A\) to \(O\).</p>
<p>\(\displaystyle\oint_C \underline{G} \cdot  \underline{dr} =\)[[input:ans3]] [[validation:ans3]][[feedback:prt2]]</p>
</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<ol class="HELM_parts">
<li>
<p>The path \(C_1\) is given in terms of the parameter \(t\) by \(x=t\), \(y=t\) and \(z=t\). Hence 
\(\qquad\displaystyle \frac{d x}{d t}=\frac{d y}{d t}=\frac{d z}{d t}=1\mbox{ and }
\frac{d\underline{r}}{d t}=\frac{d x}{d t}\underline{i}+\frac{d y}{d t}\underline{j}+\frac{d z}{d t}\underline{k}=\underline{i}+\underline{j}+\underline{k}\) 
Substituting for \(x=y=z=t\) in \(\underline{G}\) we have\(\qquad\displaystyle \underline{G}=t\underline{i}+{@aa-1@}t\underline{j}\mbox{ and } 
\underline{G}\cdot\frac{d\underline{r}}{d t}=t+{@aa-1@}t={@aa@}t\).</p>

<p>The limits of integration are  \(t=0\) and \(t=1\),  then\(\qquad\displaystyle \displaystyle\int_{C_1}\underline{G}\cdot\underline{d r}=\displaystyle\int_0^1\underline{G}\cdot\frac{d\underline{r}}{d t}\mathrm{d} t=\displaystyle\int_0^1 {@aa@}t\mathrm{d} 
t=\left[\phantom{\dfrac{1}{1}}\!\!\!{@aa/2@}t^2\right]_0^1={@aa/2@}\).</p>
<p>For the path \(C_2\) the parameterisation is \(x=t^2\), \(y=t\) and \(z=t^2\) so \(\displaystyle\frac{d\underline{r}}{d t}=2t\underline{i}+\underline{j}+2t\underline{k}\).
Substituting \(x=t^2\), \(y=t\) and \(z=t^2\) in \(\underline{G}\) we have\(\displaystyle \underline{G}=t^2\underline{i}+\left({@aa@}t^2-t\right)\underline{j}\mbox{ and } \underline{G}\cdot\frac{d\underline{r}}{d t}=2t^3+{@aa@}t^2-t\)  
\(\qquad\displaystyle \displaystyle\int_{C_2}\underline{G}\cdot\underline{d r}=\displaystyle\int_0^1 
\left(2t^3+{@aa@}t^2-t\right)d t=\left[\frac{1}{2}t^4+{@aa/3@}t^3-\frac{1}{2}t^2\right]_0^1={@aa/3@}\)</p></li>

<li><p>For the closed path \(C\) \(\displaystyle 
\displaystyle\oint_C\underline{G}\cdot\underline{d r}=\displaystyle\int_{C_1}\underline{G}\cdot\underline{d r}\displaystyle\int_{C_2}\underline{G}\cdot\underline{d r}={@aa/2@}-{@aa/3@}={@aa/6@}
\)
(Note: The value of the integral around the closed path is non-zero, unlike the previous Example.)</p></li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(8)+2;

ta1:aa/2;
ta2:aa/3;
ta3:aa/6;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\(\displaystyle \underline{G}=x\underline{i}+({@aa@}x-y)\underline{j}, \quad \displaystyle\int_{C_1}{\underline G}\cdot \underline{dr} ={@ta1@}, \quad \displaystyle\int_{C_2}{\underline G}\cdot \underline{dr} = {@ta2@} \quad \displaystyle\oint_C \underline{G} \cdot  \underline{dr} = {@ta3@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>11771757</deployedseed>
    <deployedseed>1338082070</deployedseed>
    <deployedseed>2146606579</deployedseed>
    <deployedseed>1505632629</deployedseed>
    <deployedseed>419373270</deployedseed>
    <deployedseed>940338638</deployedseed>
    <deployedseed>1106825472</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+4</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
