<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.1.1.2 simple line integral task</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>For \(F (x, y) = {@aa@}x + y^{@bb@}\), find
<ol class="HELM_parts_inline">
<li>\(\displaystyle \int_C F(x,y)\ \mathrm{d}x,\)</li>
<li>\(\displaystyle \int_C F(x,y)\ \mathrm{d}y,\)</li>
<li>\(\displaystyle \int_C F(x,y)\ \mathrm{d}s\)</li>
</ol>
where \(C\) is the line \(y = 2x\) from \((0, 0)\) to \((1, 2)\).</p>

<ol class="HELM_parts">
<li>\(\displaystyle \int_C F(x,y)\ dx =\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</li>
<li>\(\displaystyle \int_C F(x,y)\ dy =\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</li>
<li>\(\displaystyle \int_C F(x,y)\ ds =\)[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We know that \(F (x, y) = {@aa@}x + y^{@bb@}\), and that \(C\) is the line \(y = 2x\) from \((0, 0)\) to \((1, 2)\). We aim to evaluate the line integrals: (a) \(\displaystyle \int_C F(x,y)\ \mathrm{d}x,\) (b) \(\displaystyle \int_C F(x,y)\ \mathrm{d}y,\) and (c)\(\displaystyle \int_C F(x,y)\ \mathrm{d}s\).</p>

<p>
(a) We solve by rewriting the integral in terms of \(x\). As we have \(x\) between \(0\) to \(1\) for the defined line segment \(C\), these are the bounds of out integral. We can then substitute \(y=2x\) to get:
\[\begin{aligned}
\int_C F(x,y)\ \mathrm{d}x &= \int^1_0 {@aa*x@} + (2x)^{@bb@}\ \mathrm{d}x \\
&= \left[{@aa/2*x^2@} + {@(2^bb / (bb+1))*x^(bb+1)@}\right]^1_0 \\
&= \left({@subst(1, x, exp1), simp=false@} + {@ev((2^bb / (bb+1)), simp)@}(1)^{@ev(bb+1, simp)@}\right) - \left({@subst(0, x, exp1), simp=false@} + [[define simp="true"/]] {@ev(2^bb / (bb+1), simp)@}(0)^{@ev(bb+1,simp)@}\right)\\
&= {@aa/2+2^bb / (bb+1)@}
\end{aligned}\]
</p>

<p>
(b) We solve by rewriting the integral in terms of \(y\). As we have \(y\) between \(0\) to \(2\) for the defined line segment \(C\), these are the bounds of out integral. We can then substitute \(x=\dfrac{y}{2}\) to get:
\[\begin{aligned}
\int_C F(x,y)\ dy &= \int_0^2{@aa/2@}y + y^{@bb@}\ \mathrm{d}y \\
&=\left[{@exp2@} + \frac{y^{@bb+1@}}{{@bb+1@}}\right]^2_0 \\
&= \left({@subst(2, y, exp3), simp=false@}[[define simp="true"/]] + \frac{(2)^{@bb+1@}}{{@bb+1@}}\right) - \left({@subst(0, y, exp3), simp=false@}[[define simp="true"/]] + \frac{(0)^{@bb+1@}}{{@bb+1@}}\right) \\
&= {@aa + 2^(bb+1)/(bb+1)@}
\end{aligned}\]
</p>

<p>
(c) We solve by rewriting the integral in terms of \(x\). As we have \(x\) between \(0\) to \(1\) for the defined line segment \(C\), these are the bounds of out integral. We can rewrites \(\mathrm{d}s\) as follows:
\[\begin{aligned}
\mathrm{d}s &= \sqrt{1+\left(\frac{\mathrm{d} y}{\mathrm{d} x}\right)^2}\mathrm{d}x \\
&= \sqrt{1+(2)^2}\mathrm{d}x\\
&= \sqrt{5}\mathrm{d}x
\end{aligned}\]
We can then substitute \(y=2x\) into the line integral to get to get:
\[\begin{aligned}
\int_C F(x,y)\ \mathrm{d}s &= \int^1_0 \left({@aa*x@} + (2x)^{@bb@}\right)\sqrt{5}\ \mathrm{d}x\\
&\mbox{by (a) }\\
&= {@(aa/2+2^bb / (bb+1))*sqrt(5)@}
\end{aligned}\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa: rand(5)+2;
bb: rand(2)+2;

if (equal(aa, 2)) then (exp1:sconcat(stack_disp(disp_parens(x),""), "^2")) else (exp1:sconcat(stack_disp(aa/2*(disp_parens(x)),""), "^2"));
if (equal(aa, 4)) then (exp2:y^2) else (exp2:aa/4*y^2);
if (equal(aa, 4)) then (exp3:sconcat(stack_disp(disp_parens(y),""), "^2")) else (exp3:sconcat(stack_disp(aa/4*disp_parens(y),""), "^2"));

ta1:aa/2+2^bb / (bb+1);
ta2:aa + 2^(bb+1)/(bb+1);
ta3:(aa/2+2^bb / (bb+1))*sqrt(5);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\(F (x, y) = {@aa@}x + y^{@bb@}, \quad \displaystyle \int_C F(x,y)\ dx ={@ta1@}, \quad \displaystyle \int_C F(x,y)\ dy ={@ta2@}, \quad \displaystyle \int_C F(x,y)\ ds = {@ta3@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>onum</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1450502711</deployedseed>
    <deployedseed>1027634255</deployedseed>
    <deployedseed>1259717656</deployedseed>
    <deployedseed>2030740012</deployedseed>
    <deployedseed>815174386</deployedseed>
    <deployedseed>576910521</deployedseed>
    <deployedseed>1236401548</deployedseed>
    <deployedseed>2012848916</deployedseed>
    <deployedseed>1777662984</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+6</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
