<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>28.2.2.11 Third vector field divergence question</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the divergence of the following vector field, in general terms and at the point \(({@alpha@}, {@beta@}, {@gamma@})\). <br> \(\underline{F}_3 = x^{@aa@}z\underline{i} - 2y^{@bb@}z^{@cc@} \underline{j} + xyz^{@dd@} \underline{k}\)</p> <br>
<p> In terms of \(x, y,\) and \(z\): [[input:ans1]] [[validation:ans1]]</p>
<p> At the point  \(({@alpha@}, {@beta@}, {@gamma@})\): [[input:ans2]] [[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>For \(\underline{F}_3 = x^{@aa@}z\underline{i} - 2y^{@bb@}z^{@cc@} \underline{j} + xyz^{@dd@} \underline{k}\), we have that \[\underline{\nabla}\ \underline{F}_3 = \dfrac{\partial  (x^{@aa@}z)}{\partial x} + \dfrac{\partial (- 2y^{@bb@}z^{@cc@})}{\partial y} + \dfrac{\partial xyz^{@dd@}}{\partial z} = {@(aa*z*x^(ev(aa-1,simp))-2*bb*y^(ev(bb-1,simp))*z^cc+dd*x*y*z^(ev(dd-1,simp))), simp=false@}\]</p>

<p>To find the value at the point, we can substitute \(({@alpha@}, {@beta@}, {@gamma@})\) into the above equation: \({@(aa*z*x^(ev(aa-1,simp))-ev(2*bb,simp)*y^(ev(bb-1,simp))*z^cc+dd*x*y*z^(ev(dd-1,simp))), simp=false@} = {@(aa*gamma*alpha^(ev(aa-1,simp))-ev(2*bb,simp)*beta^(ev(bb-1,simp))*gamma^cc+dd*alpha*beta*gamma^(ev(dd-1,simp))), simp=false@}= {@ta2, simp=true@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>ordergreat(x,y,z);
aa:rand(4)+2;
bb:rand(4)+2;
cc:rand(4)+2;
dd:rand(2)+2;
orderless(aa*z*x^(aa-1),2*bb*y^(bb-1)*z^cc,dd*x*y*z^(dd-1))

alpha:rand(3)+2;
beta:rand(3)+2;
gamma:rand(3)+2;
ta1:aa*z*x^(aa-1)-2*bb*y^(bb-1)*z^cc+dd*x*y*z^(dd-1);
simp:false;
interim_sub(x,y,z):=ta1;
interim_expr = interim_sub(alpha,beta,gamma)
ta2:subst(alpha, x, subst(beta, y, subst(gamma, z, ta1)));
simp:true;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>\( \underline{F}_3 = x^{@aa@}z\underline{i} - 2y^{@bb@}z^{@cc@} \underline{j} + xyz^{@dd@} \underline{k} \quad \underline{\nabla}\ \underline{F}_3 = {@ta1@} \)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1362522837</deployedseed>
    <deployedseed>725828797</deployedseed>
    <deployedseed>2077961726</deployedseed>
    <deployedseed>1677393991</deployedseed>
    <deployedseed>1484780127</deployedseed>
    <deployedseed>1411766586</deployedseed>
    <deployedseed>83513143</deployedseed>
    <deployedseed>560102009</deployedseed>
    <deployedseed>881422668</deployedseed>
    <deployedseed>1210200382</deployedseed>
    <deployedseed>1935101367</deployedseed>
    <deployedseed>35008434</deployedseed>
    <deployedseed>694142879</deployedseed>
    <deployedseed>709850222</deployedseed>
    <deployedseed>337309963</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
