<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.0.3 (Write quadratic in standard form 1)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Write the quadratic {@qq@} in standard form. </p>


<p>[[input:ans1]] [[validation:ans1]] \(= 0\)</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We must rearrange the given equation to have zero on the right hand side. Combining the squared terms we get \({@aa1-aa2@} x^2 - {@bb@}x + {@cc@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>aa1: 2+ rand(20);
aa2: 2+ rand(10);
aa3: aa1-aa2
bb: 2+ rand(20);
cc: 2+ rand(20);


qq: aa1*x^2 = aa2*x^2 + bb*x - cc;

ta1: aa3*x^2 - bb*x +cc;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>[{@aa3@},{@bb@}, {@cc@}]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1297866370</deployedseed>
    <deployedseed>1244657784</deployedseed>
    <deployedseed>1864368055</deployedseed>
    <deployedseed>784100204</deployedseed>
    <deployedseed>1624480664</deployedseed>
    <deployedseed>1631743277</deployedseed>
    <deployedseed>285706569</deployedseed>
    <deployedseed>211252864</deployedseed>
    <deployedseed>13650716</deployedseed>
    <deployedseed>1969208359</deployedseed>
    <deployedseed>180587246</deployedseed>
    <deployedseed>559890616</deployedseed>
    <deployedseed>1070649336</deployedseed>
    <deployedseed>688921643</deployedseed>
    <deployedseed>471111754</deployedseed>
    <deployedseed>1189068073</deployedseed>
    <deployedseed>291601034</deployedseed>
    <deployedseed>1927760095</deployedseed>
    <deployedseed>750442445</deployedseed>
    <deployedseed>2034682352</deployedseed>
    <deployedseed>879669797</deployedseed>
    <deployedseed>1147121130</deployedseed>
    <deployedseed>1731331795</deployedseed>
    <deployedseed>543408719</deployedseed>
    <deployedseed>677948333</deployedseed>
    <deployedseed>1877873121</deployedseed>
    <deployedseed>835199824</deployedseed>
    <deployedseed>1948123352</deployedseed>
    <deployedseed>916200066</deployedseed>
    <deployedseed>1411382366</deployedseed>
    <deployedseed>440770986</deployedseed>
    <deployedseed>798207882</deployedseed>
    <deployedseed>1312989086</deployedseed>
    <deployedseed>1709986611</deployedseed>
    <deployedseed>1254795928</deployedseed>
    <deployedseed>961708417</deployedseed>
    <deployedseed>2109861405</deployedseed>
    <deployedseed>1563734397</deployedseed>
    <deployedseed>169933485</deployedseed>
    <deployedseed>526964533</deployedseed>
    <deployedseed>186793077</deployedseed>
    <deployedseed>1144127915</deployedseed>
    <deployedseed>1043203974</deployedseed>
    <deployedseed>1965263293</deployedseed>
    <deployedseed>1110024327</deployedseed>
    <deployedseed>942614996</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>aa3*x^2 + bb*x +cc;</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
