<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.1.2 (Factorising two terms positive)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Factorise the equation \({@expr@} = 0\).</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<br/>
<p>Thus state the solutions of the quadratic. Enter your answers as a list.</p>
<p>\(x = \)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In this case we wish to take out the common factor of \(x\) from \({@expr@} = 0\). This gives us 
\[
{@expr@} = {@ta1@}
\]</p>

<p>Therefore the solutions of the quadratic are the values for \(x\) which make this factorised form equal to zero, that is, \(x = 0 \mbox{ or } x = {@rr@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1+ rand(10);
bb: 1+ rand(20);

expr: aa*x^2 + bb*x

ta1: x*(aa*x + bb)=0;
rr: -bb/aa; 
ta2: [0, rr];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>[{@ta1@}, {@ta2@}]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>setify(ans2)</sans>
        <tans>setify(ta2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2098626233</deployedseed>
    <deployedseed>1198134958</deployedseed>
    <deployedseed>695890994</deployedseed>
    <deployedseed>1170810125</deployedseed>
    <deployedseed>1117800217</deployedseed>
    <deployedseed>240920668</deployedseed>
    <deployedseed>859747818</deployedseed>
    <deployedseed>692199796</deployedseed>
    <deployedseed>439621503</deployedseed>
    <deployedseed>963937616</deployedseed>
    <deployedseed>126865122</deployedseed>
    <deployedseed>1937604339</deployedseed>
    <deployedseed>1713912368</deployedseed>
    <deployedseed>676861613</deployedseed>
    <deployedseed>52545058</deployedseed>
    <deployedseed>1988022530</deployedseed>
    <deployedseed>1784122695</deployedseed>
    <deployedseed>322895595</deployedseed>
    <deployedseed>1012456085</deployedseed>
    <deployedseed>1816978458</deployedseed>
    <deployedseed>857168398</deployedseed>
    <deployedseed>1696794913</deployedseed>
    <deployedseed>3783893</deployedseed>
    <deployedseed>2101800453</deployedseed>
    <deployedseed>1309784364</deployedseed>
    <deployedseed>914660348</deployedseed>
    <deployedseed>1680199730</deployedseed>
    <deployedseed>2012129238</deployedseed>
    <deployedseed>411870799</deployedseed>
    <deployedseed>652981489</deployedseed>
    <deployedseed>673680038</deployedseed>
    <deployedseed>1081867803</deployedseed>
    <deployedseed>2048731708</deployedseed>
    <deployedseed>1808413397</deployedseed>
    <deployedseed>1292044525</deployedseed>
    <deployedseed>1032115569</deployedseed>
    <deployedseed>2054535321</deployedseed>
    <deployedseed>47088469</deployedseed>
    <deployedseed>2027757007</deployedseed>
    <deployedseed>1168815277</deployedseed>
    <deployedseed>2107183692</deployedseed>
    <deployedseed>847204523</deployedseed>
    <deployedseed>1975866773</deployedseed>
    <deployedseed>177066925</deployedseed>
    <deployedseed>1193547057</deployedseed>
    <deployedseed>1910287354</deployedseed>
    <deployedseed>1567727847</deployedseed>
    <deployedseed>379152479</deployedseed>
    <deployedseed>229783386</deployedseed>
    <deployedseed>469424634</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>x*(bb*x + aa)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>[rr,0]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
