<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.2.4 (Solve by completing the square x^2+2*bx+c)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write {@quad@} in completed square form:</p><p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<p>Hence find the solutions of \({@quad@}=0\). Enter your answers as a list and if there is a repeated root enter this twice.</p>
<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Since the \(x\) coefficient is {@2*b@} we use \(\dfrac{{@2*b@}}{2}={@b@}\) in the bracket:</p><p>\[ {@(x+b)^2=expand((x+b)^2)@} \]</p><p>This gives us the correct \(x^2\) and \(x\) terms. Since we want the constant term to be {@c@} we need to adjust by subtracting \({@b@}^2\) and adding {@c@}:</p><p>\[ \begin{align*} {@quad@} &amp;= {@(x+b)^2@}-{@b@}^2+{@c@} \\ &amp;= {@ta1@}. \end{align*} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>bb:rand_range(5,15);
cc:rand_range(1,11);
quad:x^2+2*bb*x+cc;
ta1:(x+bb)^2+cc-bb^2;
ta2: solve(quad,x)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@quad=ta1@}, {@ta2@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[x=?,x=?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>CompSquare</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>setify(ans2)</sans>
        <tans>setify(ta2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>351031172</deployedseed>
    <deployedseed>1677212515</deployedseed>
    <deployedseed>635563798</deployedseed>
    <deployedseed>1343448606</deployedseed>
    <deployedseed>26698679</deployedseed>
    <deployedseed>979496765</deployedseed>
    <deployedseed>488662793</deployedseed>
    <deployedseed>1003809762</deployedseed>
    <deployedseed>464093875</deployedseed>
    <deployedseed>769367680</deployedseed>
    <deployedseed>1624595078</deployedseed>
    <deployedseed>1546696974</deployedseed>
    <deployedseed>1800277861</deployedseed>
    <deployedseed>390691719</deployedseed>
    <deployedseed>1068469682</deployedseed>
    <deployedseed>1245576439</deployedseed>
    <deployedseed>641508977</deployedseed>
    <deployedseed>565199450</deployedseed>
    <deployedseed>463116645</deployedseed>
    <deployedseed>1850549422</deployedseed>
    <deployedseed>804679829</deployedseed>
    <deployedseed>402902847</deployedseed>
    <deployedseed>827058735</deployedseed>
    <deployedseed>1049546937</deployedseed>
    <deployedseed>1486115595</deployedseed>
    <deployedseed>1930371189</deployedseed>
    <deployedseed>243126099</deployedseed>
    <deployedseed>720610030</deployedseed>
    <deployedseed>666838032</deployedseed>
    <deployedseed>875926373</deployedseed>
    <deployedseed>817710222</deployedseed>
    <deployedseed>1973256059</deployedseed>
    <deployedseed>193573517</deployedseed>
    <deployedseed>1411430691</deployedseed>
    <deployedseed>465081489</deployedseed>
    <deployedseed>1032964907</deployedseed>
    <deployedseed>1844084932</deployedseed>
    <deployedseed>2002566062</deployedseed>
    <deployedseed>107439874</deployedseed>
    <deployedseed>930025343</deployedseed>
    <deployedseed>565233452</deployedseed>
    <deployedseed>1637022314</deployedseed>
    <deployedseed>1008744634</deployedseed>
    <deployedseed>36857256</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>quad</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
