<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.2.8 (Complete the square x^2-bx+c)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Write {@quad@} in completed square form:</p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Since the \(x\) coefficient is {@b@} we use \(\frac{b}{2}={@b/2@}\) in the bracket:</p><p>\[ {@(x+b/2)^2=expand((x+b/2)^2)@} \]</p><p>This gives us the correct \(x^2\) and \(x\) terms. Since we want the constant term to be {@c@} we need to adjust by subtracting {@b^2/4@} and adding {@c@}:</p><p>\[ \begin{align*} {@quad@} &amp;= {@(x+b/2)^2@}-{@(abs(b)/2)^2@}+{@c@} \\ &amp;= {@ta1@}. \end{align*} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[b:rand_with_step(-19,-1,2); /*make it odd so that b/2 is not "nice" */
c:rand_with_prohib(-9,9,[0]);
quad:x^2+b*x+c;
ta1:(x+b/2)^2+c-(b/2)^2;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@quad=ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>CompSquare</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1724527440</deployedseed>
    <deployedseed>181023042</deployedseed>
    <deployedseed>582208602</deployedseed>
    <deployedseed>1059010959</deployedseed>
    <deployedseed>931222882</deployedseed>
    <deployedseed>1649988532</deployedseed>
    <deployedseed>434401647</deployedseed>
    <deployedseed>440758566</deployedseed>
    <deployedseed>1903364007</deployedseed>
    <deployedseed>1550944024</deployedseed>
    <deployedseed>721639332</deployedseed>
    <deployedseed>1222475216</deployedseed>
    <deployedseed>100726151</deployedseed>
    <deployedseed>1929681014</deployedseed>
    <deployedseed>628938236</deployedseed>
    <deployedseed>1266803900</deployedseed>
    <deployedseed>1710487013</deployedseed>
    <deployedseed>590905290</deployedseed>
    <deployedseed>73470879</deployedseed>
    <deployedseed>678824039</deployedseed>
    <deployedseed>297317156</deployedseed>
    <deployedseed>2073199389</deployedseed>
    <deployedseed>491725236</deployedseed>
    <deployedseed>990183533</deployedseed>
    <deployedseed>2021114361</deployedseed>
    <deployedseed>424246165</deployedseed>
    <deployedseed>225219504</deployedseed>
    <deployedseed>1639513371</deployedseed>
    <deployedseed>1361068876</deployedseed>
    <deployedseed>628083098</deployedseed>
    <deployedseed>563732784</deployedseed>
    <deployedseed>1622237548</deployedseed>
    <deployedseed>246974287</deployedseed>
    <deployedseed>1418433886</deployedseed>
    <deployedseed>1066294484</deployedseed>
    <deployedseed>586027773</deployedseed>
    <deployedseed>132685461</deployedseed>
    <deployedseed>962814206</deployedseed>
    <deployedseed>1613555458</deployedseed>
    <deployedseed>929256160</deployedseed>
    <deployedseed>408511166</deployedseed>
    <deployedseed>315549419</deployedseed>
    <deployedseed>680727462</deployedseed>
    <deployedseed>186368229</deployedseed>
    <deployedseed>2005218170</deployedseed>
    <deployedseed>1169684651</deployedseed>
    <deployedseed>796391507</deployedseed>
    <deployedseed>1558377981</deployedseed>
    <deployedseed>369374030</deployedseed>
    <deployedseed>324486545</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>quad</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
