<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.3.20 (Line tangent to parabola)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The line with equation \(y= {@l@}\) is tangent to the parabola \( y = {@p@} \). Find the value of \(k.\)</p><p>\(k = \) [[input:ans1]][[feedback:prt1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The line&nbsp;\(y= {@l@}\) is tangent to the parabola \( y = {@p@} \) when these two intersect at one point.&nbsp;</p><p>So we solve \[ \begin{align*} {@l@} &amp;= {@p@} \\ {@p-l@} &amp;= 0 \end{align*}\]</p><p>To intersect at one point means that the&nbsp;discriminant of this quadratic equation should be \(0\).</p><p>We have \(a={@a@}\),&nbsp;\(b={@b-m@}\) and \(c={@c-k@}\), so</p><p>\[ \begin{align*} b^2-4ac = {@(b-m)^2@} - 4({@a@})({@c-k@}) &amp;= 0&nbsp; \\ {@expand((b-m)^2- 4*a*(c-k))@} &amp;= 0&nbsp; \\ k &amp;= {@ta1@}\end{align*} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>m: rand_with_prohib(-2,2,[0]);
a: 1;
b: 2*m;
c: rand_with_prohib(-3,3,[0]);
l: m*x+k;
p: a*x^2+b*x+c;
k_val: c-(b-m)^2/(4*a)
ta1:k_val</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(y= {@l@}\), \( y = {@p@} \) so \(k = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1440980554</deployedseed>
    <deployedseed>1659154070</deployedseed>
    <deployedseed>1121406829</deployedseed>
    <deployedseed>1665174575</deployedseed>
    <deployedseed>1941520094</deployedseed>
    <deployedseed>1873989928</deployedseed>
    <deployedseed>1492294796</deployedseed>
    <deployedseed>497400970</deployedseed>
    <deployedseed>624014301</deployedseed>
    <deployedseed>1969922913</deployedseed>
    <deployedseed>625554438</deployedseed>
    <deployedseed>894808341</deployedseed>
    <deployedseed>1876487125</deployedseed>
    <deployedseed>588498930</deployedseed>
    <deployedseed>1150005630</deployedseed>
    <deployedseed>1454165471</deployedseed>
    <deployedseed>1176340465</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
