<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Drag/drop: forms a subspace</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Prove the following theorem by dragging lines into the correct order.  (Note, you may not need to use all the lines!) </p>
<p>{@thm@}</p>

[[parsons input="ans1"]]
{# stackjson_stringify(proof_steps) #}
[[/parsons ]]
<p style="display:none">[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>{@proof_display(ta1, proof_steps)@}</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024072400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[stack_include("contribl://prooflib.mac");

thm:"Let \\(A\\) be a real \\(n\\times n\\) matrix.  Consider the set \\[ S = \\{ A\\mathbf{x} | \\mathbf{x}\\in\\mathbb{R}^n\\}.\\] Show that \\(S\\) is a subspace.";

proof_steps: [
    ["ass1",    "Assume \\(\\mathbf{b}_1, \\mathbf{b}_2\\in S\\)."],
    ["ca1",    "There exist \\(\\mathbf{x}_1\\), and \\(\\mathbf{x}_2\\) such that \\[ \\mathbf{b}_1=A\\mathbf{x}_1,\\mbox{ and } \\mathbf{b}_2=A\\mathbf{x}_2.\\]"],
    ["ca2",    "\\[ \\mathbf{b}_1+\\mathbf{b}_2=A\\mathbf{x}_1+A\\mathbf{x}_2=A(\\mathbf{x}_1+\\mathbf{x}_2)\\in S.\\]"],
    ["ca3",    "\\(\\mathbf{b}_1+\\mathbf{b}_2\\in S\\)"],
    ["ca4",    "\\(S\\) is closed under addition."],
    ["cs1",    "Let \\(c_1\\) be any scalar."],
    ["cs2",    "\\[ c\\mathbf{b}_1 = c A\\mathbf{x}_1 = A(c\\mathbf{x}_1) \\in S.\\]"],
    ["cs3",    "\\(S\\) is closed under scalar multiplication"],
    ["conc",  "Since \\(S\\) is closed under addition and scalar multiplication it is a subspace."],
    ["w1",    "There exist \\(\\mathbf{x}_1\\), and \\(\\mathbf{x}_2\\) such that \\[ \\mathbf{x}_1=A\\mathbf{b}_1,\\mbox{ and } \\mathbf{x}_2=A\\mathbf{b}_2.\\]"],
    ["w2",    "\\(S\\) is closed under multiplication."]
];
proof_steps:random_permutation(proof_steps);

ta:proof("ass1",proof_c(proof("ca1","ca2","ca3","ca4"), proof("cs1","cs2","cs3")),"conc");
ta1:proof("ass1",proof("ca1","ca2","ca3","ca4"), proof("cs1","cs2","cs3"),"conc");
ta2:proof(proof_c("ass1","cs1"),proof_c(proof("ca1","ca2","ca3","ca4"), proof("cs2","cs3")),"conc");

/* Generate the alternative proofs. */
tal:append(proof_alternatives(ta),proof_alternatives(ta2));
/* Create a set of flattened proofs. */
tas:setify(map(proof_flatten, tal));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@map(first, proof_steps)@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>This question is distributed as part of the STACK source code as an example. This is licenced as Creative Commons Attribution-ShareAlike 4.0 International License.</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>proof_parsons_key_json(ta, proof_steps)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa:proof_parsons_interpret(ans1);
[pd, saa]:proof_assessment(sa, tal);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>pd</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@proof_assessment_display(saa, proof_steps)@}</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1839968323</deployedseed>
    <deployedseed>1172319367</deployedseed>
    <deployedseed>430160505</deployedseed>
    <deployedseed>1087214573</deployedseed>
    <deployedseed>27425216</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>proof_parsons_key_json(ta, proof_steps)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
