<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.1.6 (Diff of one power)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the derivative of the function \(f(x) = {@v@}^{@rdm@}\).</p><p>\(f'(x) = \) [[input:R]]. [[validation:R]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[Using the rule for differentiation \[\frac{d}{dx}(x^n)=nx^{n-1},\] we have that \begin{align*}
f'(x)=   \frac{d}{d{@v@}}({@p@}) &amp;={@rdm@}{@v@}^{{@rdm@}-1} \\
    &amp;= {@rdm@}{@v@}^{@rdm-1@}.
  \end{align*}]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>v:x;
 rdm:(50+rand(12)+10*rand(6));
 p:v^(rdm);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\frac{d}{d{@v@}}({@p@}) = {@diff(p,v)@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>R</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>R</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>R</sans>
        <tans>diff(p,v)/rdm</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text>You have lost a coefficient. The derivative of \(x^n\) is \(nx^{n-1}\), here \(n={@rdm@}\) but you do not appear to have a {@rdm@} coefficient.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1941928446</deployedseed>
    <deployedseed>2016192492</deployedseed>
    <deployedseed>1948513017</deployedseed>
    <deployedseed>144852305</deployedseed>
    <deployedseed>995029406</deployedseed>
    <deployedseed>1232012922</deployedseed>
    <deployedseed>1178490731</deployedseed>
    <deployedseed>623443468</deployedseed>
    <deployedseed>374997970</deployedseed>
    <deployedseed>1733785217</deployedseed>
    <deployedseed>841029679</deployedseed>
    <deployedseed>1824287317</deployedseed>
    <deployedseed>1109443492</deployedseed>
    <deployedseed>1266511963</deployedseed>
    <deployedseed>1222748860</deployedseed>
    <deployedseed>1243318054</deployedseed>
    <deployedseed>1442028892</deployedseed>
    <deployedseed>95640284</deployedseed>
    <deployedseed>1977671824</deployedseed>
    <deployedseed>964144585</deployedseed>
    <deployedseed>864487406</deployedseed>
    <deployedseed>1673743114</deployedseed>
    <deployedseed>620226745</deployedseed>
    <deployedseed>1345877886</deployedseed>
    <deployedseed>450501626</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>R</name>
        <value>diff(p,v)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>R</name>
        <value>v^(rdm-1)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>R</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
