<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.1.11 (Diff poly and trig)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the derivative of the function \(f(x) = {@f@} .\)<br></p><p>\( f'(x) = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We have the function \( f(x) = {@f@}.\)&nbsp; Since \(\frac{d}{dx} \left(ax^n \right) = an x^{n-1}\), \(\frac{d}{dx} \sin(x) = \cos (x)\), and \(\frac{d}{dx} \cos(x) = - \sin(x)\), the derivative is \[f'(x) = {@d1@} {@d2@}+{@d3@} {@d4@}.\]<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a: rand(10)+2;
b: rand_with_step(2,10,1);
c: rand_with_prohib(5,15,[a]);
d: -(rand_with_prohib(2,10,[b]));
e: rand_with_prohib(-20,20,[0,a,b,c,d]);
simp: false;
f: a*x^2+b*cos(x)+c*x+d*sin(x)+e
ta1: ev(diff(f,x),simp);

d1: ev(diff(a*x^2, x),simp);       /* for the worked solution, to print things in the order I want */
d2: ev(diff(b*cos(x), x),simp);
d3: ev(diff(c*x, x), simp);
d4: ev(diff(d*sin(x), x), simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( f= {@f@}\) so \( f' = {@ta1@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>7588385</deployedseed>
    <deployedseed>1585199710</deployedseed>
    <deployedseed>992248392</deployedseed>
    <deployedseed>903171833</deployedseed>
    <deployedseed>1156913744</deployedseed>
    <deployedseed>23594785</deployedseed>
    <deployedseed>177011038</deployedseed>
    <deployedseed>197526321</deployedseed>
    <deployedseed>793954327</deployedseed>
    <deployedseed>17752259</deployedseed>
    <deployedseed>532924157</deployedseed>
    <deployedseed>511969636</deployedseed>
    <deployedseed>361097858</deployedseed>
    <deployedseed>775134688</deployedseed>
    <deployedseed>920948092</deployedseed>
    <deployedseed>156750285</deployedseed>
    <deployedseed>1060649474</deployedseed>
    <deployedseed>419373379</deployedseed>
    <deployedseed>530582749</deployedseed>
    <deployedseed>1492236198</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
