<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.1.17 (Linearity of differentiation)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Suppose \(f\) and \(g\) are differentiable functions with \(f'({@x@}) = {@A@}\) and \(g'({@x@}) = {@B@}\).&nbsp; Then \(({@l@}f + g)'({@x@}) = \)[[input:ans1]]. [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>By linearity of differentiation we have that \[({@l@}f + g)'({@x@}) = {@l@}f'({@x@}) + g'({@x@}) = ({@l@}\times{@A@}) + {@B@} = {@ta1@}.\]<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>A:rand_with_prohib(-10,10,[0]);
B:rand_with_prohib(-10,10,[0]);
l:rand_with_prohib(-10,10,[0,1]);
x:rand(10)- rand(10);
ta1:l*A+ B;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(({@l@}f + g)'({@x@}) = {@l@}\times{@A@} + {@B@}= {@ta1@}\).</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1024530722</deployedseed>
    <deployedseed>1602065753</deployedseed>
    <deployedseed>567481615</deployedseed>
    <deployedseed>1209926932</deployedseed>
    <deployedseed>1202972022</deployedseed>
    <deployedseed>839958075</deployedseed>
    <deployedseed>2068783849</deployedseed>
    <deployedseed>1198900163</deployedseed>
    <deployedseed>1997570900</deployedseed>
    <deployedseed>86263997</deployedseed>
    <deployedseed>1239088302</deployedseed>
    <deployedseed>1048973470</deployedseed>
    <deployedseed>173475413</deployedseed>
    <deployedseed>2401805</deployedseed>
    <deployedseed>810741448</deployedseed>
    <deployedseed>661425405</deployedseed>
    <deployedseed>615302256</deployedseed>
    <deployedseed>233313228</deployedseed>
    <deployedseed>1241677592</deployedseed>
    <deployedseed>2090179260</deployedseed>
    <deployedseed>2094826730</deployedseed>
    <deployedseed>1856549801</deployedseed>
    <deployedseed>1255059143</deployedseed>
    <deployedseed>881780909</deployedseed>
    <deployedseed>265760059</deployedseed>
    <deployedseed>315393664</deployedseed>
    <deployedseed>1099879025</deployedseed>
    <deployedseed>1863864500</deployedseed>
    <deployedseed>191016382</deployedseed>
    <deployedseed>1926523022</deployedseed>
    <deployedseed>1554132230</deployedseed>
    <deployedseed>1795746321</deployedseed>
    <deployedseed>254519663</deployedseed>
    <deployedseed>1535771717</deployedseed>
    <deployedseed>843375523</deployedseed>
    <deployedseed>1735047950</deployedseed>
    <deployedseed>1457975315</deployedseed>
    <deployedseed>1292096099</deployedseed>
    <deployedseed>278926004</deployedseed>
    <deployedseed>280798287</deployedseed>
    <deployedseed>263808574</deployedseed>
    <deployedseed>167758252</deployedseed>
    <deployedseed>1322445041</deployedseed>
    <deployedseed>1321400770</deployedseed>
    <deployedseed>1188593984</deployedseed>
    <deployedseed>1033747200</deployedseed>
    <deployedseed>709859685</deployedseed>
    <deployedseed>74389344</deployedseed>
    <deployedseed>915150876</deployedseed>
    <deployedseed>321603872</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
