<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.1.7 (Differentiate constant)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>What is \( \displaystyle \frac{d}{dx}{@a@}^{@b@}\)?<br></p><p>\(\frac{d}{dx}{@a@}^{@b@} = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\({@a@}^{@b@}\) is a constant which does not depend on \(x\), so the derivative is \(0\).<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a: rand(8)+2
b: rand(30)+15
f: a^b
ta1: diff(f,x)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \frac{d}{dx} {@a@}^{@b@} = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>621698009</deployedseed>
    <deployedseed>1844664284</deployedseed>
    <deployedseed>1860202021</deployedseed>
    <deployedseed>314142295</deployedseed>
    <deployedseed>38136591</deployedseed>
    <deployedseed>1881606854</deployedseed>
    <deployedseed>775888918</deployedseed>
    <deployedseed>2076233232</deployedseed>
    <deployedseed>2072605743</deployedseed>
    <deployedseed>995641166</deployedseed>
    <deployedseed>1422040550</deployedseed>
    <deployedseed>1274597230</deployedseed>
    <deployedseed>1272029598</deployedseed>
    <deployedseed>1258044005</deployedseed>
    <deployedseed>2054601834</deployedseed>
    <deployedseed>1714659849</deployedseed>
    <deployedseed>2101205764</deployedseed>
    <deployedseed>1326931430</deployedseed>
    <deployedseed>1492603070</deployedseed>
    <deployedseed>1294887903</deployedseed>
    <deployedseed>1990837602</deployedseed>
    <deployedseed>1500412815</deployedseed>
    <deployedseed>644300967</deployedseed>
    <deployedseed>1755181345</deployedseed>
    <deployedseed>362673273</deployedseed>
    <deployedseed>1742408268</deployedseed>
    <deployedseed>1083914562</deployedseed>
    <deployedseed>795875001</deployedseed>
    <deployedseed>229952063</deployedseed>
    <deployedseed>1480642317</deployedseed>
    <deployedseed>1518263050</deployedseed>
    <deployedseed>310046739</deployedseed>
    <deployedseed>34400955</deployedseed>
    <deployedseed>1639186340</deployedseed>
    <deployedseed>1972000436</deployedseed>
    <deployedseed>1883654454</deployedseed>
    <deployedseed>1524735580</deployedseed>
    <deployedseed>995692671</deployedseed>
    <deployedseed>140490314</deployedseed>
    <deployedseed>33015652</deployedseed>
    <deployedseed>2025152053</deployedseed>
    <deployedseed>67831246</deployedseed>
    <deployedseed>342093264</deployedseed>
    <deployedseed>519142460</deployedseed>
    <deployedseed>1364626764</deployedseed>
    <deployedseed>2042014764</deployedseed>
    <deployedseed>109931726</deployedseed>
    <deployedseed>142882425</deployedseed>
    <deployedseed>105011232</deployedseed>
    <deployedseed>1765548175</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
