<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.2.2 (General power CAP W10 Q16 edited)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>What is \[ \frac{d^{@kval@}}{dx^{{@kval@}}} x^{{@nval@}}?\] <br></p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We want to calculate the \({@kval@}^{th} \) derivative \( \displaystyle \frac{d^{{@kval@}}}{dx^{{@kval@}}} x^{{@nval@}}\). We have that \begin{align*}\frac{d}{dx} x^{{@nval@}} &amp;= {{@nval@}}x^{{@nval@}-1}&nbsp; \\<br>\frac{d^2}{dx^2} x^{{@nval@}} &amp;=&nbsp;&nbsp;\frac{d}{dx}&nbsp;({@nval@}x^{{@nval@}-1}) = {@nval@}({@nval@}-1)x^{{@nval@}-2} \\<br>&amp;\vdots\\<br>\frac{d^{{@kval@}}}{dx^{{@kval@}}} x^{{@nval@}} &amp;=&nbsp;{@nval@}({@nval@}-1)({@nval@}-2)\ldots({@nval@}-{@kval@}+1)x^{{@nval@}-{@kval@}} = {@ta@}.\end{align*} <br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>kval:rand_range(2,4);
nval:kval+1+rand_range(4,12);
ta:nval!/(nval-kval)!*x^(nval-kval);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\frac{d^{@kval@}}{dx^{{@kval@}}} x^{{@nval@}} = {@ta@}.\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>cross</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1288901042</deployedseed>
    <deployedseed>275236753</deployedseed>
    <deployedseed>486516919</deployedseed>
    <deployedseed>1789069733</deployedseed>
    <deployedseed>1437001197</deployedseed>
    <deployedseed>540350234</deployedseed>
    <deployedseed>421498135</deployedseed>
    <deployedseed>940041859</deployedseed>
    <deployedseed>1150254375</deployedseed>
    <deployedseed>681369450</deployedseed>
    <deployedseed>1820088276</deployedseed>
    <deployedseed>1045342799</deployedseed>
    <deployedseed>128114193</deployedseed>
    <deployedseed>2047214266</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
