<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.2.23 (Trig stationary points 2)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>What is the value of \(\tan(x)\) at the stationary points of the function \(f(x) = {@ff@}\)?</p><p>\(\tan(x) = \) [[input:ans1]]. [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Differentiating the function \(f(x) = {@ff@}\) we get \(f'(x) ={@df@}\). At a stationary point \(x\) we therefore have that \({@df@} = 0.\) If \(\cos(x) = 0\) , then \(\sin(x) = \pm1\), and the equation would not be satisfied.&nbsp; Therefore \(\cos(x) \neq0\), so we can divide by \(\cos(x)\) to get the equivalent equation \[{@aa@} + {@bb@}\tan(x) = 0,\]which implies that \(\tan(x) = {@ta1@}\).<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1+rand(10);
bb: rand_with_prohib(2,10,[a]);
ff:aa*sin(x) -bb*cos(x);
df:diff(ff,x);
ta1:-aa/bb;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\((x) = {@ff@}\) has deriv \({@df@}\) and \(\tan(x) = {@ta1@}\) at stat points.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1532560671</deployedseed>
    <deployedseed>235550983</deployedseed>
    <deployedseed>202496764</deployedseed>
    <deployedseed>1238278994</deployedseed>
    <deployedseed>1558668754</deployedseed>
    <deployedseed>2056624620</deployedseed>
    <deployedseed>182432521</deployedseed>
    <deployedseed>306459451</deployedseed>
    <deployedseed>612087044</deployedseed>
    <deployedseed>1829187151</deployedseed>
    <deployedseed>991273820</deployedseed>
    <deployedseed>1851596879</deployedseed>
    <deployedseed>1939820606</deployedseed>
    <deployedseed>1430351396</deployedseed>
    <deployedseed>824396910</deployedseed>
    <deployedseed>1493638414</deployedseed>
    <deployedseed>192834691</deployedseed>
    <deployedseed>1756694525</deployedseed>
    <deployedseed>266974239</deployedseed>
    <deployedseed>1636699313</deployedseed>
    <deployedseed>807223213</deployedseed>
    <deployedseed>553922121</deployedseed>
    <deployedseed>1290704511</deployedseed>
    <deployedseed>1906072625</deployedseed>
    <deployedseed>2099170051</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
