<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.2.21 (Match function and deriv CAP W3 Q24 used)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>For each of the functions graphed below, identify the graph of its derivative.</p>
<table>
  <tbody>
  <tr>
    <td style="text-align: center;">{@thumbnailplot(funcstoplot[1],"Function 1")@}</td>
    <td style="text-align: center;">{@thumbnailplot(funcstoplot[2],"Function 2")@}</td><td style="text-align: center;">{@thumbnailplot(funcstoplot[3],"Function 3")@}</td><td style="text-align: center;">{@thumbnailplot(funcstoplot[4],"Function 4")@}</td>
  </tr><tr>
    <td style="text-align: center;">&nbsp;[[input:ans1]] [[validation:ans1]]</td>
    <td style="text-align: center;">&nbsp;&nbsp;[[input:ans2]] [[validation:ans2]]</td><td style="text-align: center;">&nbsp;[[input:ans3]] [[validation:ans3]]</td><td style="text-align: center;">&nbsp;&nbsp;[[input:ans4]] [[validation:ans4]]</td>
  </tr>
  
  
</tbody></table><br>
<p>Choose from the graphs below:</p><ul> </ul>
<table>
  <tbody><tr>
    <td style="text-align: center;"><b>Graph A</b></td>
    <td style="text-align: center;"><b>Graph B</b></td><td style="text-align: center;">&nbsp;<b>Graph C</b></td><td style="text-align: center;">&nbsp;<b>Graph D</b></td>
  </tr>
  <tr>
    <td style="text-align: center;">{@thumbnailplot(derivs[1],"Graph A")@}</td>
    <td style="text-align: center;">{@thumbnailplot(derivs[2],"Graph B")@}</td><td style="text-align: center;">{@thumbnailplot(derivs[3],"Graph C")@}</td><td style="text-align: center;">{@thumbnailplot(derivs[4],"Graph D")@}</td>
  </tr>
  
  
</tbody></table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text>.</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[m:rand_with_prohib(-3,3,[-1,0,1]);
c:rand_with_prohib(-3,3,[0]);
k:rand_range(-2,2);
pm:rand([-1,1]);
funcs:[exp(x)-3*x+1, ln(abs(x)), 1/x, exp(2/3*x+1), 5*x*exp(x), (pm)*x^2+(k), 1/8*x^3-3/2*x, (1/16)*x^4-(1/2)*x^2+1, 1/20*x^4-2/15*x^3-2/5*x^2+8/5*x+43/60, (m)*x+(c)];
/* Choose functions to plot */
funcstoplot:rand_selection(funcs,4);
/* permute the graphs but keep note of which ones are right! */
reordering:random_permutation([1,2,3,4]);
fntoletter:zip_with("[",reordering,[A,B,C,D]);
/* apply the reordering to produce the shuffled list of options */
derivs:makelist(diff(funcstoplot[reordering[i]],x),i,1,4);
corr:maplist(lambda([ex],assoc(ex,fntoletter)), [1,2,3,4]);
ta1:maplist(lambda([ex],[ex, is(ex=assoc(1,fntoletter))]), [A,B,C,D]);
ta2:maplist(lambda([ex],[ex, is(ex=assoc(2,fntoletter))]), [A,B,C,D]);
ta3:maplist(lambda([ex],[ex, is(ex=assoc(3,fntoletter))]), [A,B,C,D]);
ta4:maplist(lambda([ex],[ex, is(ex=assoc(4,fntoletter))]), [A,B,C,D]);
thumbnailplot:lambda([ex1,ex2],plot(ex1,[x,-4,4],[y,-4,4],[box, false],[yx_ratio, 1],[xtics, -3, 6, 3],[ytics, -3, 6, 3],[axes, solid],[size,170,170],[plottags,false],[alt,ex2]));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]][[feedback:prt2]][[feedback:prt3]][[feedback:prt4]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@funcstoplot@} {#fntoletter#}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>corr[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>corr[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>corr[3]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>corr[4]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>652522794</deployedseed>
    <deployedseed>624297973</deployedseed>
    <deployedseed>699294286</deployedseed>
    <deployedseed>1057434893</deployedseed>
    <deployedseed>172931438</deployedseed>
    <deployedseed>588830452</deployedseed>
    <deployedseed>525230219</deployedseed>
    <deployedseed>795328852</deployedseed>
    <deployedseed>1212894299</deployedseed>
    <deployedseed>698707277</deployedseed>
    <deployedseed>1824586690</deployedseed>
    <deployedseed>1021950607</deployedseed>
    <deployedseed>1588883438</deployedseed>
    <deployedseed>1353933939</deployedseed>
    <deployedseed>1836229454</deployedseed>
    <deployedseed>1071447909</deployedseed>
    <deployedseed>1632345143</deployedseed>
    <deployedseed>1841103205</deployedseed>
    <deployedseed>1448447998</deployedseed>
    <deployedseed>747134516</deployedseed>
    <deployedseed>279153719</deployedseed>
    <deployedseed>868545898</deployedseed>
    <deployedseed>1853324039</deployedseed>
    <deployedseed>921637043</deployedseed>
    <deployedseed>1688151468</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>corr[1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>corr[2]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>corr[3]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>corr[4]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
