<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.2.25 (Integrate linear function)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \(\alpha\), \(\beta\), and \(\gamma\) such that the function \(f(x) = \alpha x^2 + \beta x + \gamma\) has a derivative given by {@y@} for all \(x\), and such that the graph of \(f\) passes through the point \( ({@A@}, {@B@} )\).</p><ul><li>\( \alpha = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li><li>\( \beta = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li><li>\( \gamma = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]<br></li></ul><p><br></p><p><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The function \(f(x) = \alpha x^2 + \beta x + \gamma\) has derivative \(f'(x) = 2 \alpha x&nbsp; + \beta\).&nbsp; Comparing coefficients with the given formula for the derivative \(f'(x)={@y@}\), we see that \(\alpha = {@alpha@}\) and \(\beta = {@beta@}\).<br></p><p>To find \(\gamma\) we use the further information that \[{@B@}= f({@A@}) = {@alpha@}({@A@})^2 + {@beta@}({@A@}) + \gamma.\] Hence \( \gamma = {@B@} - {@alpha*A^2@} - {@c*A@} = {@gamma@}\).<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>A: rand(5)+2
B: rand_with_prohib(1,10,[a])
m: rand_with_step(4,20,2)
c: rand(3)+2
y: m*x+c

alpha: m/2
beta: c
gamma: B-(m/2)*A^2-c*A

ta1: alpha;
ta2: beta;
ta3: gamma;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( P = ({@A@}, {@B@}) \), \( y = {@y@}\) so \( \alpha = {@ta1@}\), \( \beta = {@ta2@}\), \( \gamma = {@ta3@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>38928878</deployedseed>
    <deployedseed>1366596688</deployedseed>
    <deployedseed>403590395</deployedseed>
    <deployedseed>1256474152</deployedseed>
    <deployedseed>1216727152</deployedseed>
    <deployedseed>1787039038</deployedseed>
    <deployedseed>1699812807</deployedseed>
    <deployedseed>2092736781</deployedseed>
    <deployedseed>1395429211</deployedseed>
    <deployedseed>1772449910</deployedseed>
    <deployedseed>78203847</deployedseed>
    <deployedseed>490281308</deployedseed>
    <deployedseed>1171054136</deployedseed>
    <deployedseed>53777352</deployedseed>
    <deployedseed>986212542</deployedseed>
    <deployedseed>268936134</deployedseed>
    <deployedseed>1114393262</deployedseed>
    <deployedseed>1581463203</deployedseed>
    <deployedseed>18819926</deployedseed>
    <deployedseed>1188278271</deployedseed>
    <deployedseed>901938978</deployedseed>
    <deployedseed>1994303348</deployedseed>
    <deployedseed>1804782108</deployedseed>
    <deployedseed>1563166004</deployedseed>
    <deployedseed>1241578141</deployedseed>
    <deployedseed>1060100710</deployedseed>
    <deployedseed>498333914</deployedseed>
    <deployedseed>967172203</deployedseed>
    <deployedseed>1231982672</deployedseed>
    <deployedseed>224223885</deployedseed>
    <deployedseed>250202888</deployedseed>
    <deployedseed>1571762381</deployedseed>
    <deployedseed>1506745470</deployedseed>
    <deployedseed>1552799485</deployedseed>
    <deployedseed>1072355983</deployedseed>
    <deployedseed>1238809413</deployedseed>
    <deployedseed>1624412676</deployedseed>
    <deployedseed>850874120</deployedseed>
    <deployedseed>1672428050</deployedseed>
    <deployedseed>358001899</deployedseed>
    <deployedseed>760589089</deployedseed>
    <deployedseed>1048833972</deployedseed>
    <deployedseed>1572236159</deployedseed>
    <deployedseed>633413512</deployedseed>
    <deployedseed>1467156419</deployedseed>
    <deployedseed>598881448</deployedseed>
    <deployedseed>2088297851</deployedseed>
    <deployedseed>116855768</deployedseed>
    <deployedseed>388879554</deployedseed>
    <deployedseed>624339724</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1, simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+1, simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta3+1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
