<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.3.36 (Definite integral - find limit)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the values of \(k\) for which \( \displaystyle\int_0^k {\left( {@f@} \right)\,dx} ={@c@} \).</p><p>\(k =\) [[input:ans1]] [[feedback:prt1]] [[validation:ans1]] Give your answer as a set.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We are given the integral \[ \int_0^k {\left( {@f@} \right)\,dx} = {@c@}\]</p><p>We can also calculate</p><p>\[\begin{align*}&nbsp;\int_0^k \left( {@f@} \right)\,dx &amp;= \left[{@int_f@}\right]_0^k&nbsp; \\ &amp;= &nbsp;{@f_k@} \end{align*}\]</p><p>So we need to solve</p><p>\[\begin{align*}&nbsp; {@f_k@} &amp;= {@c@} \\&nbsp; {@integrate(p1,x,0,k)@} &amp;= {@c@} + {@integrate(p2,x,0,k)@} \\&nbsp;({@integrate(p1,x,0,k)@})^2 &amp;= ({@c@} + {@integrate(p2,x,0,k)@})^2 \\&nbsp;{@expand((k-k1)*(k-k2))@} = 0&nbsp;&nbsp; \end{align*}\]</p><p>By solving this quadratic equation we find \(k={@k1@}\) or \(k={@k2@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[v: rand([1,2,3]);
k_val: [[1,4], [2,8], [4,9]][v];
con: [[3,2],[3*sqrt(2),4], [5,6]][v];
a: con[1];
b:1;
c: con[2];
k1: k_val[1]
k2: k_val[2];

p1: a/(2*sqrt(x));
p2: b;
f: a/(2*sqrt(x))-b;
int_f: integrate(f,x);
assume(k>0);
f_k: integrate(f,x,0,k);

ta1: set(k1,k2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \displaystyle\int_0^k {\left( {@f@} \right)\,dx} ={@c@} \), \( k = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{?,?}</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>347005190</deployedseed>
    <deployedseed>1176675577</deployedseed>
    <deployedseed>33827136</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+{6}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
