<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.3.7 (Use area interpretation to find definite integral)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The graph of \(y=f(x)\) is shown below.</p>

<div style="margin-left:4em;margin-right:auto;">[[jsxgraph width='400px' height='200px']] JXG.Options.axis.ticks.majorHeight = 20; JXG.Options.axis.ticks.insertTicks = false; JXG.Options.axis.ticks.ticksDistance = 100; JXG.Options.text.fontSize = 16; var brd = JXG.JSXGraph.initBoard(divid, {axis:true,
    grid: true, boundingbox:[-0.5,4.5,6.5,-0.8], showCopyright: false, showNavigation:false}); var allsegs = brd.create('curve',[{#[0,ax,bx,cx]#},{#[0,yval,yval,0]#}],{strokeWidth:3,strokeColor:'#003399'}); var xaxis = brd.create('axis', [[0,0],[1,0]]);
    brd.create('ticks', [xaxis, [2,4,6]], {strokeColor: '#003399', majorHeight:15, drawLabels: true, label: { offset: [0, -10],anchorX: 'middle', anchorY: 'top' }}); var yaxis = brd.create('axis', [[0,0],[0,1]]); brd.create('ticks', [yaxis, [2,4,6]],
    {strokeColor: '#003399', majorHeight:15, drawLabels: true, label: { offset: [-10, 0],anchorX: 'right', anchorY: 'middle' }}); [[/jsxgraph]]
</div>

<p>Use the area interpretation of the definite integral to evaluate the following:</p>
<p></p>
\(\displaystyle {@q1@}=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
<p></p>
\(\displaystyle {@q2@}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>(a) </b>This is simply the area of a triangle, with base {#ax#} and height {#yval#}. So the area is \(\frac12\times{@ax@}\times{@yval@}={@ta1@}\).</p><p><b>(b) </b>This area is made up of two pieces:</p><p></p><ul><li>The first&nbsp;(\({@ax@}\leq x\leq{@bx@}\)) is a rectangle with base {#bx-ax#} and height {#yval#}.</li><li>The second (\({@bx@}\leq x\leq{@cx@}\)) is a triangle with&nbsp;base {#cx-bx#} and height {#yval#}.</li></ul>So overall the area is \({@(bx-ax)*yval@}+{@(cx-bx)*yval/2@}={@ta2@}\).<br><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025042500</text>
    </stackversion>
    <questionvariables>
      <text>[ax,bx,cx]:sort( rand_selection(makelist(i,i,1,6), 3) ); /* three distinct numbers from 1 to 6 */
yval:rand_range(2,4);
q1:'int(f(x),x,0,ax);
q2:'int(f(x),x,ax,cx);
ta1:ax*yval/2;
ta2:(bx-ax)*yval+(cx-bx)*yval/2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{#[ax,bx,cx]#} and {#yval#} gives {#ta1#} and {#ta2#}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>791078357</deployedseed>
    <deployedseed>546191293</deployedseed>
    <deployedseed>1985372359</deployedseed>
    <deployedseed>1797510566</deployedseed>
    <deployedseed>2093437746</deployedseed>
    <deployedseed>191100783</deployedseed>
    <deployedseed>1700312491</deployedseed>
    <deployedseed>2120049808</deployedseed>
    <deployedseed>375914480</deployedseed>
    <deployedseed>1617593956</deployedseed>
    <deployedseed>1827223439</deployedseed>
    <deployedseed>1906316178</deployedseed>
    <deployedseed>1892674796</deployedseed>
    <deployedseed>178150757</deployedseed>
    <deployedseed>170037490</deployedseed>
    <deployedseed>1317666766</deployedseed>
    <deployedseed>874332635</deployedseed>
    <deployedseed>1635189044</deployedseed>
    <deployedseed>64518583</deployedseed>
    <deployedseed>1681117828</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>jsxgraph</text>
</tag>
    </tags>
  </question>
</quiz>
