<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.3.16 (Definite integral constant)</text>
    </name>
    <questiontext format="html">
      <text>Evaluate the following definite integral:
\[ \int^{@uu@}_{@ll@}{@pp@}\:d{@v@}. \]

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>An antiderivative of {@pp@} is \(\displaystyle{@qq@}\).</p>
<p>So, using the Fundamental Theorem,
\[
\begin{align*}
 \int^{@uu@}_{@ll@}{@pp@}d{@v@} &= \left[{@qq@}\right]^{{@uu@}}_{@ll@}
\\
&= {@a1@} - {@a2@}
\\
&= {@ta1@}
\end{align*}
\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>v:x;
aa: 1+rand(10);
bb: 1+ rand(10);

ll: rand(2);
uu: ll+(1+rand(8));
pp: 1+ rand(100);

qq:int(pp,v);

a1: ev(qq, x=uu);
a2: ev(qq,x=ll);


ta1:int(pp,v,ll,uu);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ \int^{@uu@}_{@ll@}{@pp@}\:d{@v@} = {@ta1@}. \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1696237443</deployedseed>
    <deployedseed>295942388</deployedseed>
    <deployedseed>2012360947</deployedseed>
    <deployedseed>1845801086</deployedseed>
    <deployedseed>419859138</deployedseed>
    <deployedseed>403810619</deployedseed>
    <deployedseed>2051772800</deployedseed>
    <deployedseed>1698604821</deployedseed>
    <deployedseed>755601607</deployedseed>
    <deployedseed>1313608831</deployedseed>
    <deployedseed>2122707149</deployedseed>
    <deployedseed>324466885</deployedseed>
    <deployedseed>1990887729</deployedseed>
    <deployedseed>514843223</deployedseed>
    <deployedseed>952229813</deployedseed>
    <deployedseed>2140994196</deployedseed>
    <deployedseed>508791118</deployedseed>
    <deployedseed>333211626</deployedseed>
    <deployedseed>1372855101</deployedseed>
    <deployedseed>143764337</deployedseed>
    <deployedseed>440588618</deployedseed>
    <deployedseed>155583459</deployedseed>
    <deployedseed>1202679469</deployedseed>
    <deployedseed>1332254399</deployedseed>
    <deployedseed>117982657</deployedseed>
    <deployedseed>1475048220</deployedseed>
    <deployedseed>713779586</deployedseed>
    <deployedseed>331974085</deployedseed>
    <deployedseed>1909158511</deployedseed>
    <deployedseed>2114065978</deployedseed>
    <deployedseed>2005772441</deployedseed>
    <deployedseed>1104967767</deployedseed>
    <deployedseed>1885130688</deployedseed>
    <deployedseed>42324514</deployedseed>
    <deployedseed>1411484730</deployedseed>
    <deployedseed>625261175</deployedseed>
    <deployedseed>2086985623</deployedseed>
    <deployedseed>1608313273</deployedseed>
    <deployedseed>1222082773</deployedseed>
    <deployedseed>422472246</deployedseed>
    <deployedseed>1704231458</deployedseed>
    <deployedseed>1320515579</deployedseed>
    <deployedseed>299804306</deployedseed>
    <deployedseed>1123669079</deployedseed>
    <deployedseed>888719521</deployedseed>
    <deployedseed>125367594</deployedseed>
    <deployedseed>1854889726</deployedseed>
    <deployedseed>1678101807</deployedseed>
    <deployedseed>1685423584</deployedseed>
    <deployedseed>1893006525</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
