<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>10.3.2 (Understanding Riemann sums (x^2 on [0,6]))</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Set the slider to \(n={@n@}\) - you should see that the left sum is {@ls@} and the right sum is {@rs@}.</p><p>Complete the following calculation which shows how the value of the {@LorR@} sum is obtained:</p><p>\(\displaystyle \sum_{i={@imin@}}^{@imax@} f\left({@i*Rend/n@}\right)\times\frac{{@Rend@}}{@n@}\)</p><p>\(\displaystyle = {@Rend/n@}\bigg({@(firstbit),simp=false@}+\)&nbsp;[[input:ans1]] \(+{@lastbit@}\bigg)\)</p><p>[[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The missing term is when \(i={@missingterm@}\) so the missing value is \(f({@missingterm@})={@ta1@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022082900</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[f(x):=x^2;
Rend:6;
n:rand([3,6]);
[ls,rs]:[[40,112],[55,91]][n/3];
ver:rand([1,2]);
LorR:["left","right"][ver];
imin:[0,1][ver];
imax:[n-1,n][ver];
missingterm:rand(makelist(i,i,imin+1,imax-1));
evalpointsL:makelist(i*Rend/n, i, imin, missingterm-1);
evalpointsR:makelist(i*Rend/n, i, missingterm+1,imax);
fvalsL:map(f,evalpointsL);
fvalsR:map(f,evalpointsR);
simp:false;
firstbit:apply("+",map(lambda([ex],ev(ex,simp)),fvalsL));
lastbit:apply("+",map(lambda([ex],ev(ex,simp)),fvalsR));
simp:true;
ta1:ev(f(i*Rend/n), i=missingterm);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>n={#n#} {#LorR#} sum term {#missingterm#} is {@firstbit@} + ({@ta1@}) + {@lastbit@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>symi</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1529012604</deployedseed>
    <deployedseed>1819999398</deployedseed>
    <deployedseed>1319337050</deployedseed>
    <deployedseed>851265833</deployedseed>
    <deployedseed>87168363</deployedseed>
    <deployedseed>15958233</deployedseed>
    <deployedseed>1408744798</deployedseed>
    <deployedseed>1806911073</deployedseed>
    <deployedseed>1766791657</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
