<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.10.2 (Finding limits quad over quad 2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the limit of \(f(x)\) as \(x\) goes to infinity where \[f(x)&nbsp; = {@func@}.\]</p><p>\(\lim_{x \to \infty} f(x) = \) [[input:ans1]]. [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In this case, notice that we can divide every term in the function by \(x^2\) to get 
\[
{@func@} = \frac{{@aa@} + \frac{{@bb@}}{x} + \frac{{@cc@}}{x^2}}{{@dd@}+ \frac{{@ee@}}{x^2}}
\]
Thus 
\[
\lim\limits_{x \to \infty} {@func@} = \lim\limits_{x \to \infty} \frac{{@aa@} + \frac{{@bb@}}{x} + \frac{{@cc@}}{x^2}}{{@dd@} + \frac{{@ee@}}{x^2}} = \frac{{@aa@}+0+0}{{@dd@}+0} = {@ta1@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa:1+rand(10);
bb:1+rand(10);
cc: 1+rand(10);
dd: 1+rand(10);
ee: 1+rand(10);

top: aa*x^2 + bb*x + cc
bot: dd*x^2 + ee

func: top/bot;

ta1: aa/dd;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@func@} has limit {@ta1@}.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2024991490</deployedseed>
    <deployedseed>1561001250</deployedseed>
    <deployedseed>515196083</deployedseed>
    <deployedseed>1040593706</deployedseed>
    <deployedseed>2095618207</deployedseed>
    <deployedseed>1386757907</deployedseed>
    <deployedseed>2859875</deployedseed>
    <deployedseed>635928004</deployedseed>
    <deployedseed>1363918542</deployedseed>
    <deployedseed>776443223</deployedseed>
    <deployedseed>1088818896</deployedseed>
    <deployedseed>661745768</deployedseed>
    <deployedseed>551509760</deployedseed>
    <deployedseed>1835567357</deployedseed>
    <deployedseed>478608434</deployedseed>
    <deployedseed>263315026</deployedseed>
    <deployedseed>1815186581</deployedseed>
    <deployedseed>1290340593</deployedseed>
    <deployedseed>1354848136</deployedseed>
    <deployedseed>1523659116</deployedseed>
    <deployedseed>1347707954</deployedseed>
    <deployedseed>1357638053</deployedseed>
    <deployedseed>1619257950</deployedseed>
    <deployedseed>1274764816</deployedseed>
    <deployedseed>475543272</deployedseed>
    <deployedseed>2037609971</deployedseed>
    <deployedseed>619633167</deployedseed>
    <deployedseed>449097186</deployedseed>
    <deployedseed>207168661</deployedseed>
    <deployedseed>1635711672</deployedseed>
    <deployedseed>1736654959</deployedseed>
    <deployedseed>1686214179</deployedseed>
    <deployedseed>1169134350</deployedseed>
    <deployedseed>165714388</deployedseed>
    <deployedseed>21702105</deployedseed>
    <deployedseed>1963630503</deployedseed>
    <deployedseed>746936254</deployedseed>
    <deployedseed>1367515950</deployedseed>
    <deployedseed>498552835</deployedseed>
    <deployedseed>303300536</deployedseed>
    <deployedseed>418941046</deployedseed>
    <deployedseed>601082654</deployedseed>
    <deployedseed>1676151718</deployedseed>
    <deployedseed>541546585</deployedseed>
    <deployedseed>956925262</deployedseed>
    <deployedseed>1727758905</deployedseed>
    <deployedseed>538159052</deployedseed>
    <deployedseed>1494203268</deployedseed>
    <deployedseed>1554576184</deployedseed>
    <deployedseed>1513672327</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
