<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.10.11 (Solve for cubic coeff given stat point)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the value of \(a\) if the function \(f(x) ={@f@}\) has a stationary point at the point \(x={@p@}\).</p><p>\(a= \) [[input:ans1]]. [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Differentiating \({@f@}\) we get\[f'(x) = {@df@}.\]Since \(x={@p@}\) is a stationary point, we know that \(f'({@p@}) = 0\).&nbsp; Substituting this into the expression for \(f'(x)\) gives the equation\[{@lhs@}=0.\]The solution to this equation is \[a={@ta1@}.\]<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>A:2+rand(4);
B:-(1+rand(4));
C:2+rand(4);
D:rand(5);
p:rand_with_prohib(-4,4,[0]);
f:A*x^3+B*a*x^2+C*a*x+D;
df:diff(f,x);
lhs:3*A*p^2+2*B*a*p+C*a;
ta1:-3*A*p^2/(2*B*p+C);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@f@} has a stationary point at \(x={@p@}\) implies \(a={@ta1@}\).</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>239921803</deployedseed>
    <deployedseed>1934443898</deployedseed>
    <deployedseed>1521457500</deployedseed>
    <deployedseed>129365311</deployedseed>
    <deployedseed>1715566187</deployedseed>
    <deployedseed>2141373344</deployedseed>
    <deployedseed>1961604543</deployedseed>
    <deployedseed>376459736</deployedseed>
    <deployedseed>2116465326</deployedseed>
    <deployedseed>1392720824</deployedseed>
    <deployedseed>40377461</deployedseed>
    <deployedseed>1626635323</deployedseed>
    <deployedseed>6411845</deployedseed>
    <deployedseed>1392667185</deployedseed>
    <deployedseed>1295373959</deployedseed>
    <deployedseed>1369664654</deployedseed>
    <deployedseed>1435161332</deployedseed>
    <deployedseed>1698503888</deployedseed>
    <deployedseed>717766469</deployedseed>
    <deployedseed>547212183</deployedseed>
    <deployedseed>701905589</deployedseed>
    <deployedseed>690873645</deployedseed>
    <deployedseed>518843869</deployedseed>
    <deployedseed>1232336913</deployedseed>
    <deployedseed>263572201</deployedseed>
    <deployedseed>59207240</deployedseed>
    <deployedseed>1143577170</deployedseed>
    <deployedseed>1112568652</deployedseed>
    <deployedseed>1241311923</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
