<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-17-misc-10</text>
    </name>
    <questiontext format="html">
      <text>Find \[ \int^{@rdm1@}_{@rdm2@} {@p@} d{@v@}. \]

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[To solve this problem we will have to express {@p@} in a simpler way using some trigonometric identities <hint>trig_standard_identities</hint>

Now, recall that if \(u\) and \(v\) are both functions of \(x\) then 
\[\int u dv = uv - \int v du.\]

So, let \(u = {@pp3@}\) and \(dv = {@pp1@} d{@v@}\). 

Then \(du = {@diff(pp3,v)@} d{@v@}\) and \(v = {@int(pp1,v)@}\).

So, 

\[\begin{array}{rcl}\displaystyle  
\int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int u dv \\
   &\displaystyle =&\displaystyle  uv - \int v du \\
   &\displaystyle =&\displaystyle  [{@pp3@}][{@int(pp1,v)@}] - \int [{@int(pp1,v)@}][{@diff(pp3,v)@}d{@v@}] +c \\
   &\displaystyle =&\displaystyle  {@expand(pp3*int(pp1,v))@} - \int {@expand(int(pp1,v)*diff(pp3,v))@}d{@v@} +c \\
   &\displaystyle =&\displaystyle  {@expand(pp3*int(pp1,v))@} + 2\int {@pp1@}[{@pp4@}] d{@v@} +c \\
   &\displaystyle =&\displaystyle  {@expand(pp3*int(pp1,v))@} + 2\int {@expand(pp1*pp4)@} d{@v@} +c
 \end{array}\] 

Hence, 
\[3\int {@p@} d{@v@} = {@expand(pp3*int(pp1,v))@} + 2\int {@pp1@} d{@v@} +c \]

And so,

\[\begin{array}{rcl}\displaystyle 
\int {@p@} d{@v@} &\displaystyle =&\displaystyle  {@expand(pp3*int(pp1,v)/3)@}+ {@2/3*int(pp1,v)@} + \frac{c}{3} \\
   &\displaystyle =&\displaystyle  {@expand(pp3*int(pp1,v)/3)+2/3*int(pp1,v)@} + \frac{c}{3}.
 \end{array}\] 

Then
\[\begin{array}{rcl}\displaystyle   
\int^{@rdm1@}_{@rdm2@} {@p@} d{@v@} &\displaystyle =&\displaystyle  F({@rdm1@}) - F({@rdm2@}) \\
   &\displaystyle =&\displaystyle  {@subst(rdm1,v,q)@} - \left[{@subst(rdm2,v,q)@}\right] \\
   &\displaystyle =&\displaystyle  {@subst(rdm1,v,q)-subst(rdm2,v,q)@}
 \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm1:rand([2*pi/3,3*pi/4,5*pi/6,2*pi]);
rdm2:rand([0,pi/6,pi/4,pi/3,pi/2]);
pp1:sin(v);
pp2:pp1^2;
pp3:1-(cos(v))^2;
pp4:1-(sin(v))^2;
p:pp1*pp2;
q:int(p,v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ \int^{@rdm1@}_{@rdm2@} {@p@} d{@v@} = {@int(p,v,rdm2,rdm1)@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v,rdm2,rdm1)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>int(p,v,rdm2,rdm1)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1680008391</deployedseed>
    <deployedseed>1675046735</deployedseed>
    <deployedseed>731362511</deployedseed>
    <deployedseed>717859251</deployedseed>
    <deployedseed>1785730177</deployedseed>
    <deployedseed>295670300</deployedseed>
    <deployedseed>1970688357</deployedseed>
    <deployedseed>396337498</deployedseed>
    <deployedseed>256313615</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v,rdm2,rdm1),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
