<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-17-misc-16</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \[ \int^{@rdm1@}_{@rdm2@} {@p@} d{@v@} \] by letting \({@v@}=2\sin(z)\). [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We want to solve \[\int {@p@} d{@v@}\] using the substitution \({@v@} = 2\sin(z)\). So, if \({@v@} = 2\sin(z)\) then \(d{@v@} = 2\cos(z)dz\) and \[\begin{array}{rcl}\displaystyle \int{@p@}d{@v@} &amp;\displaystyle =&amp;\displaystyle \int \frac{1}{[4-(2\sin(z))^2]^{3/2}}2\cos(z)dz \\ &amp;\displaystyle =&amp;\displaystyle \int \frac{2\cos(z)dz}{[4-4\sin^2(z)]^{3/2}} \\&amp;\displaystyle =&amp;\displaystyle \int \frac{2\cos(z)dz}{[4(1-\sin^2(z))]^{3/2}} \\ &amp;\displaystyle =&amp;\displaystyle \int \frac{2\cos(z)dz}{4^{3/2}[1-\sin^2(z)]^{3/2}} \\ &amp;\displaystyle =&amp;\displaystyle \int \frac{\cos(z)dz}{4[\cos^2(z)]^{3/2}} \\ &amp;\displaystyle =&amp;\displaystyle \int \frac{\cos(z)dz}{4\cos^3(z)} \\ &amp;\displaystyle =&amp;\displaystyle \int \frac{dz}{4\cos^2(z)} \\ &amp;\displaystyle =&amp;\displaystyle \frac{1}{4}\int \sec^2(z)dz \\ &amp;\displaystyle =&amp;\displaystyle \frac{\tan(z)}{4} +c \\ &amp;\displaystyle =&amp;\displaystyle \frac{1}{4}\tan\left[\sin^{-1}\left(\frac{t}{2}\right)\right]+c \end{array}\] We need \[\int^{@rdm1@}_{@rdm2@} {@p@} d{@v@}\] so, recall that if \[\int f(x) dx = F(x)+c\] then \[\int^a_b f(x) dx = F(a) - F(b).\] Thus, define \[F({@v@}) = {@q@}\] Then, \[\begin{array}{rcl}\displaystyle \int^{@rdm1@}_{@rdm2@} {@p@} d{@v@} &amp;\displaystyle =&amp;\displaystyle F({@rdm1@}) - F({@rdm2@}) \\ &amp;\displaystyle =&amp;\displaystyle {@subst(rdm1,v,q1)@} - {@subst(rdm2,v,q1)@} \\ &amp;\displaystyle =&amp;\displaystyle {@subst(rdm1,v,q1)-subst(rdm2,v,q1)@} \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>dumvar0:triginverses:false;
v:t;
rdm1:rand([2/2^(1/2),1,3^(1/2)]);
rdm2:rand([0,-2/2^(1/2),-1,-3^(1/2)]);
p:1/(4-v^2)^(3/2);
q:tan(asin(v/2))/4;
q1:v/(4*(4-v^2)^(1/2));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ \int^{@rdm1@}_{@rdm2@} {@p@} d{@v@} = {@int(p,v,rdm2,rdm1)@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v,rdm2,rdm1)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>int(p,v,rdm2,rdm1)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1728498360</deployedseed>
    <deployedseed>1091379041</deployedseed>
    <deployedseed>71585193</deployedseed>
    <deployedseed>435243449</deployedseed>
    <deployedseed>2029635578</deployedseed>
    <deployedseed>808096431</deployedseed>
    <deployedseed>1496729721</deployedseed>
    <deployedseed>878092155</deployedseed>
    <deployedseed>1193994982</deployedseed>
    <deployedseed>1029206032</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v,rdm2,rdm1),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
