<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-17-misc-12</text>
    </name>
    <questiontext format="html">
      <text>Find \[ \int {@p@} d{@v@}. \]

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the method of integration by parts <hint>calc_int_methods_parts</hint> as follows:

Let \(u = {@pp1@}\) and \(dv = {@pp2@} d{@v@}\). 

Then \(du = {@diff(pp1,v)@} d{@v@}\) and \(v = {@int(pp2,v)@}\).

So, 

\[\begin{array}{rcl}\displaystyle  
\int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int u dv \\&\displaystyle =&\displaystyle  uv - \int v du \\&\displaystyle =&\displaystyle  {@pp1@}{@int(pp2,v)@} - \int {@int(pp2,v)@}{@diff(pp1,v)@}d{@v@}
 \end{array}\] 

In order to solve \[\int {@int(pp2,v)@}{@diff(pp1,v)@}d{@v@}\] we need to apply integration by parts again difining 

\(u = {@qq1@}\) and \(dv = {@qq2@} d{@v@}\). 

Then \(du = {@diff(qq1,v)@} d{@v@}\) and \(v = {@int(qq2,v)@}\).

So, 

\[\begin{array}{rcl}\displaystyle  
\int {@int(pp2,v)@}{@diff(pp1,v)@}d{@v@} &\displaystyle =&\displaystyle  \int u dv \\&\displaystyle =&\displaystyle  uv - \int v du \\&\displaystyle =&\displaystyle  {@qq1@}{@int(qq2,v)@} - \int {@int(qq2,v)@}[{@diff(qq1,v)@}d{@v@}] \\&\displaystyle =&\displaystyle  {@qq1@}{@int(qq2,v)@} - \int {@int(qq2,v)*diff(qq1,v)@}d{@v@} \\&\displaystyle =&\displaystyle  {@qq1@}{@int(qq2,v)@} + {@-'int(int(qq2,v)*diff(qq1,v),v)@}
 \end{array}\] 

Hence, 

\[\begin{array}{rcl}\displaystyle 
\int {@p@} d{@v@} &\displaystyle =&\displaystyle  {@pp1@}{@int(pp2,v)@} - \left({@qq1@}{@int(qq2,v)@} + {@-'int(int(qq2,v)*diff(qq1,v),v)@}\right) \\&\displaystyle =&\displaystyle  {@pp1*int(pp2,v)-(qq1*int(qq2,v)+-'int(-rdm*rdm*p,v))@}
 \end{array}\] 

Thus,

\[{@'int(p,v)-'int(-rdm*rdm*p,v)@} = {@pp1*int(pp2,v)-qq1*int(qq2,v)@}\]

And so,

\[{@'int(p,v)@} = {@(pp1*int(pp2,v)-qq1*int(qq2,v))/(rdm*rdm+1)@}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm:2+rand(4);
pp2:exp(v);
pp1:sin(rdm*v);
p:pp1*pp2;
qq1:diff(pp1,v);
qq2:diff(pp2,v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int {@p@} d{@v@} = {@int(p,v)@}+c\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v)+c</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>int(p,v)+c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1446119699</deployedseed>
    <deployedseed>370334331</deployedseed>
    <deployedseed>1533982546</deployedseed>
    <deployedseed>1316908555</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
