<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Conics-1-tangents and normals - random</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the equations of the tangent and normal to the circle \[ {@q@}\] at the point \(({@ap@},{@bp@})\).</p>
<p>Tangent: [[input:ans1]] [[validation:ans1]]</p>
<p><span>Normal: [[input:ans2]] </span><span>[[validation:ans2]]<span><br /></span></span></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Tangents:</strong> Circles are a special case of the ellipse centred at \((x_c,y_c)\)<br />\[ \frac{(x-x_c)^2}{a^2}+\frac{(y-y_c)^2}{b^2}=1.\]<br />The tangent to this ellipse, through the point \((x',y')\), is given by<br />\[ \frac{(x-x_c)(x'-x_c)}{a^2}+\frac{(y-y_c)(y'-y_c)}{b^2}=1.\]<br />In our case we have \(a^2=b^2={@rs@}\) and the equation of our circle is \[ {@p@}\]  Using \(x'={@ap@}\) and \(y'={@bp@}\) we have the tangent line as \[ {@(x-ac)*(ap-ac)+(y-bc)*(bp-bc)=rs@}\] Expanding and rearranging give \[ {@expand(tas1)@}\]</p>
<p><strong>Normals:</strong> A line through \((x',y')\) with gradient \(m\) has equation \[ y=m(x-x')+y'.\] If a line has gradient \(m\) then the perpendicular line has gradient \(-1/m\). Our tangent line has gradient {@pm@}, so the normal line through \(({@ap@},{@bp@})\) has equation \[ y= {@tas2@} = {@expand(tas2)@}.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>ac:-2+rand(5)
bc:-2+rand(5)
ap:ac+3
bp:bc+4
rs:(ap-ac)^2+(bp-bc)^2
/* A circle through the above points. */
p:(x-ac)^2+(y-bc)^2=rs
q:expand(p)
ta1:(x-ac)*(ap-ac)+(y-bc)*(bp-bc)=rs
tas1:first(solve(ta1,y))
pm:coeff(expand(rhs(tas1)),x)
tas2:-1/pm*(x-ap)+bp
ta2:y=-1/pm*(x-ap)+bp</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>Tangent: <span data-mce-mark="1">[[feedback:prt1]]</span></p>
<p>Normal: <span>[[feedback:prt2]]</span></p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@q@} through \( ({@ap@},{@bp@}) \) gives tangent: {@ta1@}, normal {@ta2@}.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span>The equation for your tangent is correct.</span></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><span>The equation for your tangent is incorrect.</span></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The equation for your normal is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><span>The equation for your normal is incorrect.</span></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2093153158</deployedseed>
    <deployedseed>841483598</deployedseed>
    <deployedseed>1084169326</deployedseed>
    <deployedseed>562057308</deployedseed>
    <deployedseed>1011273153</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
