<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Standard_Conic_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The standard equation of a conic takes one of the following five formats:</p>
<p> \[ (1) \frac{x^2}{a^2} + \frac{y^2}{b^2} = 1, \mbox{(ellipse)},\]</p>
<p> \[ (2) \frac{x^2}{a^2} - \frac{y^2}{b^2} = 1,\quad (3) \frac{y^2}{a^2} - \frac{x^2}{b^2} = 1, \mbox{(hyperbola)},\]</p>
<p> \[ (4) y = a x^2,\quad (5) x = a y^2, \mbox{(parabola)}.\]</p>
<p>Reduce the following quadratic equation of a conic to one of these standard forms:</p>
<p> \[ {@p@} . \]</p>
<p>Give your answer as a list containing (in this order) the format number, the value of \(a\) and the value of \(b\).  In the case of a parabola, enter 0 for \(b\).</p>
<p>\([format,a,b]\) = [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Completing the squares in \(x\) and \(y\), the equation \({@p@}\) can be written as \[ {@eq@}.\] Hence this is in format \(({@form@})\) with \(a = {@a@}\) and \(b = {@b@}\).</p>
<p><br /> The correct answer is therefore \({@correctanswer@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>xc:rand(7)+1;
 yc:-(rand(7)+1);
 a:rand(6)+2;
 b:rand(6)+2;
 form:3;
 eq:a^2*b^2*((y-yc)^2/a^2-(x-xc)^2/b^2-1)=0;
 p:expand(eq);
 correctanswer:[form,a,b];
 wronganswer1:[form-1,a,b];
 wronganswer2:[form,a-10,b];
 wronganswer3:[form,a,b-10];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@p@}, {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>40</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>factor</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1272579496</deployedseed>
    <deployedseed>348645896</deployedseed>
    <deployedseed>561537106</deployedseed>
    <deployedseed>2124918667</deployedseed>
    <deployedseed>214926692</deployedseed>
    <deployedseed>1821035000</deployedseed>
    <deployedseed>1909254224</deployedseed>
    <deployedseed>222408848</deployedseed>
    <deployedseed>804601870</deployedseed>
    <deployedseed>510543426</deployedseed>
    <deployedseed>374029553</deployedseed>
    <deployedseed>1567759341</deployedseed>
    <deployedseed>224958515</deployedseed>
    <deployedseed>1846576043</deployedseed>
    <deployedseed>635451057</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer3</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
