<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>HELM 32.2 E2 Performing a time step</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**32.2 E2 Performing a Time Step**

A numerical scheme has been used to approximate the solution of

\\\[ \displaystyle {@'diff(y,t)=t^2-y^2@},\quad y(0)=0\\\]


and has given the following estimates, to 6 decimal places,

\\\({@'y(tn)@} \approx {@yn@}\\\)<br>
\\\({@'y(tn1)@} \approx {@yn1@}\\\)

Now use the 2-step, explicit linear multistep scheme

\\\(y_{n+2} − 1.6\,y_{n+1} + 0.6\,y_n = h \,(5\,f_{n+1} − 4.6\,f_n) \\\)

to approximate {@'y(tn2)@}.

\\\({@'y(tn2)@}\approx\\\) [[input:S_yn2]] [[validation:S_yn2]] [[feedback:yn2]]]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024012900</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[f(t,y):= t^2-y^2;
h: 0.02;
tn: 0.3;
tn1: tn+h;
tn2: tn+2*h;
yn: 1.471433;
yn1: 1.447892;
yn2: 1.6*yn1-0.6*yn + h*( 5*f(tn1,yn1)-4.6*f(tn,yn) );


stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_value_EN.mac");]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="markdown">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens/>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>S_yn2</name>
      <type>numerical</type>
      <tans>yn2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>7 s.f.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>yn2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>S_yn2</sans>
        <tans>yn2</tans>
        <testoptions>5.0e-7</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>yn2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>yn2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_unit(S_yn2,yn2,5.0e-7)@}</text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_yn2</name>
        <value>yn2</value>
      </testinput>
      <expected>
        <name>yn2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>yn2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
