<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.1.4 (Forming permutations of words)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>How many different {@k@}-letter “words” can be made by rearranging letters from \({@letters@}\)?</p><p>[[input:ans2]][[feedback:prt1]][[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>There are {@n@} choices for the 1st letter, but then only {@ev(n-1,simp)@} for the 2nd, etc. Overall there are \( \frac{{@n@}!}{{@ev(n-k,simp)@}!} = {@prod2@} = {@ta2@} \) possibilities.<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[k: 3+rand(6);
n: rand_range(k+1,8,1);
a: rand(18);
letters: substring("ABCDEFGHIJKLMNOPQRSTUVWXYZ",1+a,1+a+n); /*selection of n consecutive letters from A-Z starting at position a*/
ta1: n^k;
ta2: n! / (n-k)!;
terms2:makelist(n-i,i,0,k-1);
simp:false;
prod2:apply("*",terms2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@k@}-letter words from {@letters@}:  {@ta2@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>cross</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1364799779</deployedseed>
    <deployedseed>1366008838</deployedseed>
    <deployedseed>377400871</deployedseed>
    <deployedseed>1252206913</deployedseed>
    <deployedseed>552978627</deployedseed>
    <deployedseed>777900397</deployedseed>
    <deployedseed>1108932538</deployedseed>
    <deployedseed>1318821051</deployedseed>
    <deployedseed>283870946</deployedseed>
    <deployedseed>1768465723</deployedseed>
    <deployedseed>536341560</deployedseed>
    <deployedseed>628097899</deployedseed>
    <deployedseed>361300308</deployedseed>
    <deployedseed>875589006</deployedseed>
    <deployedseed>1278059859</deployedseed>
    <deployedseed>1982566353</deployedseed>
    <deployedseed>1077184370</deployedseed>
    <deployedseed>680166919</deployedseed>
    <deployedseed>1855627726</deployedseed>
    <deployedseed>1254973910</deployedseed>
    <deployedseed>295850622</deployedseed>
    <deployedseed>562263159</deployedseed>
    <deployedseed>1555356699</deployedseed>
    <deployedseed>382578307</deployedseed>
    <deployedseed>154606000</deployedseed>
    <deployedseed>202266835</deployedseed>
    <deployedseed>1862982960</deployedseed>
    <deployedseed>1717158963</deployedseed>
    <deployedseed>324063186</deployedseed>
    <deployedseed>641581130</deployedseed>
    <deployedseed>142748649</deployedseed>
    <deployedseed>292542638</deployedseed>
    <deployedseed>1663040994</deployedseed>
    <deployedseed>1169944892</deployedseed>
    <deployedseed>1284985731</deployedseed>
    <deployedseed>879534344</deployedseed>
    <deployedseed>336285391</deployedseed>
    <deployedseed>1870910545</deployedseed>
    <deployedseed>741414404</deployedseed>
    <deployedseed>2043209784</deployedseed>
    <deployedseed>172984201</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>ta2*2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>n!/(n-k)!</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>prod2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
