<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.2.20 (Find a if know coef of x^k)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the value of \( a&gt;0\) if the coefficient of \( x^{@k@} \) in the expansion of \(\displaystyle \left({@t1+t2@}\right)^{@n@}\) is {@coef@}.</p><p>\( a = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>By the binomial theorem, the general term in the expansion of&nbsp;\(\displaystyle \left({@t1+t2@}\right)^{@n@}\) is</p><p>\[ \binom{{@n@}}{r}\left({@t1@}\right)^{{@n@}-r}\left({@t2@}\right)^{r} = \binom{{@n@}}{r}{@t1^(n-r)*t2^r@} \]</p><p>We need to find the term \(x^{@k@}\) in the expansion, so we need</p><p>\[\begin{align*} {@expand(2*(n-r)-r)@}&amp;={@k@} \\ r&amp;= {@j_val@}.\end{align*}\]</p><p>Therefore this gives us the term in the expansion containing \(x^{@k@}\)&nbsp;as \[ \binom{@n@}{@j_val@}\times {@ev(t1^(n-r)*t2^r,r=j_val)@} = {@term@}. \]</p><p></p><p>We know that the coefficient of the \( x^{@k@} \)&nbsp;term is {@coef@} so</p><p>\[\begin{align*} {@binomial(n,j_val)/a^(j_val)@} &amp;= {@coef@} \\ a &amp;= {@ta@}.\end{align*}\]</p><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>assume_pos:true
v:x
a_val:rand([2,3,4,5])
n:rand_range(8,12);
j_val:rand_range(2,3);
k:2*(n-j_val)-j_val;
t1: x^2;
t2: 1/(a*x)
p: (x^2+1/(a*x))^10;
/*k:11 ; 
j_val: ev(r,solve(2*(n-r)-r=k,r));*/
term: binomial(n,j_val)*ev(t1^(n-j)*t2^j,j=j_val);
coef: binomial(n,j_val)/a_val^j_val;

ta :a_val;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Coeff of  \( x^{@k@} \) in \(\displaystyle \left({@t1+t2@}\right)^{@n@}\) is {@coef@} so \(a={@a_val@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>expand,coeff,!</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>binomial(n,k)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You must calculate more than just the binomial coefficient!</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>374747074</deployedseed>
    <deployedseed>218380213</deployedseed>
    <deployedseed>1591494410</deployedseed>
    <deployedseed>986200891</deployedseed>
    <deployedseed>1003956222</deployedseed>
    <deployedseed>1002657187</deployedseed>
    <deployedseed>716056132</deployedseed>
    <deployedseed>188707385</deployedseed>
    <deployedseed>1729820154</deployedseed>
    <deployedseed>208807887</deployedseed>
    <deployedseed>28255931</deployedseed>
    <deployedseed>850146336</deployedseed>
    <deployedseed>1918910133</deployedseed>
    <deployedseed>672939971</deployedseed>
    <deployedseed>717783462</deployedseed>
    <deployedseed>709603280</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(binomial(n,k),simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
