<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.2.4 (Coefficient in expansion of 2 brackets)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>What is the coefficient of \(x\) in the expansion of {@p@}?<br> <br>[[input:ans]] [[validation:ans]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We expand the expression:<br></p><p>\[ {@p@} = {@expand(p)@}\]</p>
<p>So the coefficient of \(x\) is {@ta@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a: rand_with_prohib(-4,4,[0]);
b: rand_with_prohib(-4,4,[0]);
c: rand_with_prohib(-4,4,[0]);
d: rand_with_prohib(-4,4,[0]);
p1: a*x+b;
p2: c*x+d;
p: p1*p2;

ta: a*d+b*c;
w1: a*c</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@p@} = {@expand(p)@},\) coefficient of \(x\) is {@ta@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>expand</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>1-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>w1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is the coefficient of \(x^2\)&nbsp; and not of \(x\).&nbsp;</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1003936662</deployedseed>
    <deployedseed>1644831134</deployedseed>
    <deployedseed>813125368</deployedseed>
    <deployedseed>929554157</deployedseed>
    <deployedseed>175447042</deployedseed>
    <deployedseed>1469172468</deployedseed>
    <deployedseed>547934557</deployedseed>
    <deployedseed>192959312</deployedseed>
    <deployedseed>1240918341</deployedseed>
    <deployedseed>536866164</deployedseed>
    <deployedseed>471239201</deployedseed>
    <deployedseed>93068326</deployedseed>
    <deployedseed>1432717534</deployedseed>
    <deployedseed>1300386403</deployedseed>
    <deployedseed>1171622629</deployedseed>
    <deployedseed>723207000</deployedseed>
    <deployedseed>1056873851</deployedseed>
    <deployedseed>2034915454</deployedseed>
    <deployedseed>1653141458</deployedseed>
    <deployedseed>1682355464</deployedseed>
    <deployedseed>2050062737</deployedseed>
    <deployedseed>374298732</deployedseed>
    <deployedseed>58465050</deployedseed>
    <deployedseed>374819108</deployedseed>
    <deployedseed>742596892</deployedseed>
    <deployedseed>554805030</deployedseed>
    <deployedseed>2077836560</deployedseed>
    <deployedseed>1665917976</deployedseed>
    <deployedseed>701292372</deployedseed>
    <deployedseed>1023262962</deployedseed>
    <deployedseed>1751113776</deployedseed>
    <deployedseed>1501089374</deployedseed>
    <deployedseed>1965533925</deployedseed>
    <deployedseed>521628378</deployedseed>
    <deployedseed>540643694</deployedseed>
    <deployedseed>104982347</deployedseed>
    <deployedseed>780084410</deployedseed>
    <deployedseed>1920771592</deployedseed>
    <deployedseed>1883607013</deployedseed>
    <deployedseed>1226235069</deployedseed>
    <deployedseed>1407532560</deployedseed>
    <deployedseed>265835763</deployedseed>
    <deployedseed>1947175337</deployedseed>
    <deployedseed>1673034708</deployedseed>
    <deployedseed>1273269112</deployedseed>
    <deployedseed>129377304</deployedseed>
    <deployedseed>415202839</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>w1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
