<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.9.2 (Evaluate binomial (scaffolded))</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Complete the following calculation:</p><table style="display:inline">
<tbody>
<tr>
<td rowspan="2" valign="bottom">\(\displaystyle\binom{{@n@}}{{@ev(n-k,simp)@}}=\) </td><td style="text-align: center; border-bottom: 1pt solid black;" valign="bottom">\({@n@}!\)</td><td style="text-align: center;" rowspan="2">&nbsp;\(=\)</td><td style="text-align:center; border-bottom: 1pt solid black;">[[input:ans1]][[validation:ans1]][[feedback:prt1]]</td><td style="text-align: center;" rowspan="2">&nbsp;\(=\)</td><td style="text-align:left;" rowspan="2">[[input:ans2]][[validation:ans2]][[feedback:prt2]]</td>
</tr>
<tr>
<td style="text-align: center;">\({@ev(n-k,simp)@}!{@k@}! \)</td><td style="text-align:center;">\({@k@}!\)</td>
</tr>
</tbody>
</table><p><br></p><p>Note: you can enter your answers as a product of integers if you wish, e.g. <code>12*11</code> rather than <code>132</code><br></p><p><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\(\displaystyle&nbsp;\binom{{@n@}}{{@ev(n-k,simp)@}}&nbsp;= \frac{{@n@}!}{{@ev(n-k,simp)@}!{@k@}!} = \frac{{@ta1@}}{{@k@}!} = {@ta2@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[n:rand_range(100,120);
k:rand([2,3]);
kfac:k!;
ta2:binomial(n,n-k);
terms2:makelist(n-i,i,0,k-1);
ta1:(simp:false,apply("*",terms2));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\displaystyle\binom{{@n@}}{{@ev(n-k,simp)@}}=\frac{{@ta1@}}{{@k@}!}={@ta2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>cross</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>!,binomial</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1453399700</deployedseed>
    <deployedseed>1534328660</deployedseed>
    <deployedseed>1791595385</deployedseed>
    <deployedseed>999670451</deployedseed>
    <deployedseed>978887856</deployedseed>
    <deployedseed>39279400</deployedseed>
    <deployedseed>32778423</deployedseed>
    <deployedseed>1856360042</deployedseed>
    <deployedseed>190255706</deployedseed>
    <deployedseed>887563320</deployedseed>
    <deployedseed>179765324</deployedseed>
    <deployedseed>868639223</deployedseed>
    <deployedseed>901665343</deployedseed>
    <deployedseed>412474317</deployedseed>
    <deployedseed>944359532</deployedseed>
    <deployedseed>229406776</deployedseed>
    <deployedseed>1849183242</deployedseed>
    <deployedseed>2011235321</deployedseed>
    <deployedseed>1546562492</deployedseed>
    <deployedseed>150966447</deployedseed>
    <deployedseed>617637355</deployedseed>
    <deployedseed>377460171</deployedseed>
    <deployedseed>1653233553</deployedseed>
    <deployedseed>1327775427</deployedseed>
    <deployedseed>1083275793</deployedseed>
    <deployedseed>261758918</deployedseed>
    <deployedseed>1544255789</deployedseed>
    <deployedseed>1167412885</deployedseed>
    <deployedseed>1578213434</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2*kfac</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1/kfac</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
