<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>(not in quiz) 37.2.3.11 Seven batches of articles are manufactured</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>Seven batches of articles are manufactured. Each batch contains {@n_art_each_batch@} articles. Each article has, independently, a probability of {@p_art_being_def@} of being defective. Find the probability that there is at least one defective article:</p>
<ol class="HELM_parts">

<li><p>In exactly 4 of the batches [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

<li><p>In 4 or more of the batches [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

</ol>
<p><b>Note:</b> round your answer up to \(4\) decimal places.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>SOLUTION:</b></p><p>First, let's compute the probability of having at least 1 defective article on each batch:</p><p>Let \(Z\) be the number of defective articles in a batch, this follows a binomial distribution with \(n={@n_art_each_batch@}\) and \(p = {@p_art_being_def@} \)</p><p>\(\begin{align} \text{P}(Z \geq 1) &amp; = 1 - \text{P}(Z &lt; 1) = 1 - \text{P}(Z=0) \\ &amp; = 1 - \ ^{@n_art_each_batch@}\!C_0p^0(1-p)^{{@n_art_each_batch@}} = 1 - (1 - {@p_art_being_def@})^{{@n_art_each_batch@}} \approx {@p_batch@}&nbsp;\end{align}&nbsp;\)</p><p>With the latter result, we can compute now what the exercise is asking us. Let \(X\) be the number of batches with at least one defective article, so \(X\) follows a binomial distribution with \(n = 7\) and \(p = {@p_batch@}\):</p>
<ol class="HELM_parts">

<li><p>\(\text{P} (X=4) =&nbsp;\ ^7\!C_4({@p_batch@})^4(1-{{@p_batch@}})^3 = {@ta_1@}&nbsp; \)</p></li>

<li><p>\(\begin{align} \text{P} (X\geq 4)&nbsp; = &amp; \ ^7\!C_4({@p_batch@})^4(1-{{@p_batch@}})^3 +&nbsp;^7\!C_5({@p_batch@})^5(1-{{@p_batch@}})^2 \\ &amp; +&nbsp;^7\!C_6({@p_batch@})^6(1-{{@p_batch@}}) +&nbsp;^7\!C_7({@p_batch@})^7&nbsp;= {@ta_2@}&nbsp;\end{align} \)</p></li>

</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>n_art_each_batch: 8+rand(4)
p_art_being_def: 0.1 + 0.1*rand(3)
p_batch: decimalplaces(1 - (1 - p_art_being_def)^n_art_each_batch,4)
ta_1: decimalplaces((factorial(7)/(factorial(4)*factorial(3)))*(p_batch^4)*((1-p_batch)^3),4)
ta_2: decimalplaces((factorial(7)/(factorial(4)*factorial(3)))*(p_batch^4)*((1-p_batch)^3) + 0.5*7*6*(p_batch^5)*((1-p_batch)^2) + 7*(p_batch^6)*(1-p_batch) + p_batch^7,4)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\text{P} (X=4) = {@ta_1@}\), \(\text{P} (X \geq 4) = {@ta_2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:4</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:4</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>0.0002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions>0.0002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>443847003</deployedseed>
    <deployedseed>1676943325</deployedseed>
    <deployedseed>89757571</deployedseed>
    <deployedseed>728618754</deployedseed>
    <deployedseed>1635846331</deployedseed>
    <deployedseed>970443320</deployedseed>
    <deployedseed>661902089</deployedseed>
    <deployedseed>925273533</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>10</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
