<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>37.2.3.1 Tyre burst-mountain-bike rider</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>The probability that a mountain-bike rider travelling along a certain track will have a tyre burst
is {@p_tyre@}. Find the probability that among {@n_riders@} riders:</p>

<ol class="HELM_parts">

<li><p>exactly one has a burst tyre: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

<li><p>at most three have a burst tyre: [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

<li><p>two or more have burst tyres: [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p></li>

</ol>

<p><b>Note:</b> please round your answer to the nearest \(3\) decimal places.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>SOLUTION:</b></p><p>Binomial distribution \(\text{P} (X=r)=\&nbsp; ^n\!C_rp^r(1-p)^{n-r}\) where \(p\) is the probability of single `success' which is `tyre burst'.<br></p><p><br></p><ol class="HELM_parts">
  <li>\(\text{P} (X=1)=\  ^{{@n_riders@}}\!C_1({@p_tyre@})^1({@1-p_tyre@})^{{@n_riders -1@}}={@ta_1@}\)</li>
  <li><p><span class="math display">\(\begin{aligned}
\text{P} (X\le 3)&amp;=\text{P} (X=0)+\text{P} (X=1)+\text{P} (X=2)+\text{P} (X=3)\\
&amp;=({@1-p_tyre@})^{{@n_riders@}}+{@n_riders@}({@p_tyre@})({@1-p_tyre@})^{{@n_riders - 1@}} \\ &amp; \qquad+ \frac{{@n_riders@}\times {@n_riders-1@}}{2\times 1}({@p_tyre@})^2({@1-p_tyre@})^{{@n_riders - 2@}}\\
&amp;\qquad +\frac{{@n_riders@}\times {@n_riders-1@}\times {@n_riders-2@}}{3\times 2\times 1}({@p_tyre@})^3({@1-p_tyre@})^{{@n_riders-3@}} \\ &amp; ={@ta_2@}\end{aligned}\)</span></p></li>
  <li>\(\begin{aligned}&nbsp;\text{P} (X\ge 2)&amp;=1-\text{P} [(X=0)\cup (X=1)] \\ &amp; = 1-({@1-p_tyre@})^{{@n_riders@}}-{@n_riders@}({@p_tyre@})({@1-p_tyre@})^{{@n_riders - 1@}} \\ &amp;= {@ta_3@}\end{aligned}\)</li></ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>p_tyre: 0.03 + 0.01*rand(5);
n_riders: 14 + rand(6);
ta_1: decimalplaces(n_riders*p_tyre*(1-p_tyre)^(n_riders-1),3);
ta_2: decimalplaces((1-p_tyre)^n_riders + n_riders*p_tyre*(1-p_tyre)^(n_riders-1) + 0.5*n_riders*(n_riders -1)*(p_tyre^2)*((1-p_tyre)^(n_riders-2)) + n_riders*(n_riders -1)*(n_riders -2)*(p_tyre^3)*((1-p_tyre)^(n_riders-3))/6,3);
ta_3: decimalplaces(1 - (1-p_tyre)^n_riders - n_riders*p_tyre*(1-p_tyre)^(n_riders-1),3);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>The probability that exactly one has a burst tyre: {@ta_1@}, 
The probability that at most three have a burst tyre: {@ta_2@}, 
The probability that two or more have burst tyres: {@ta_3@},</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 3 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:3, mindp:3</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 3 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:3, mindp:3</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta_3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 3 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:3, mindp:3</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>0.002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions>0.002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions>0.002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1739996348</deployedseed>
    <deployedseed>1415728846</deployedseed>
    <deployedseed>1945974954</deployedseed>
    <deployedseed>414060726</deployedseed>
    <deployedseed>1423325327</deployedseed>
    <deployedseed>1978541256</deployedseed>
    <deployedseed>807176114</deployedseed>
    <deployedseed>1136305975</deployedseed>
    <deployedseed>1072565195</deployedseed>
    <deployedseed>895749751</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta_1+1.000,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta_2+1.000,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta_3+1.000,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
