<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>37.2.3.5  Probability that a machine develops a fault (NOT IN QUIZ)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p><p>The probability that a machine develops a fault within the first 3 years of use is {@p@}. If {@n@} machines are selected at random, calculate the probability that {@n-2@} or more will <b>not</b> develop any faults within the first 3 years of use.</p><p>[[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]</p><p><b>Note:</b> round your answer up to \(4\) decimal places, if required.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>SOLUTION:</b></p><p></p><p>Let <span class="math inline">\(X\)</span> be a random variable ‘number of machines that will not develop any faults within the first 3 years of use’ then it is clear that <span class="math inline">\(X\)</span> follows a binomial distribution with <span class="math inline">\(n={@n@}\)</span> and <span class="math inline">\(p={@1-p@}\)</span>.</p><p><span class="math display">\[\begin{aligned}&nbsp;\text{P} (X \geq {@n-2@}) &amp; = \text{P} [(X={@n-2@})\cup (X={@n-1@})\cup (X={@n@})] \\ &amp;=&nbsp; \  ^{{@n@}}\!C_{@n-2@}({@1-p@})^{@n-2@}({@p@})^2 +\  ^{{@n@}}\!C_{@n-1@}({@1-p@})^{@n-1@}({@p@}) +\  ^{{@n@}}\!C_{{@n@}}({@1-p@})^{{@n@}} \\ &amp;= {@ta@}\end{aligned}\]</span></p><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>n: 36 + rand(5);
p: 0.002 + 0.001*rand(5);
ta: decimalplaces(0.5*n*(n-1)*((1-p)^(n-2))*(p^2) + n*((1-p)^(n-1))*p + (1-p)^n,4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\text{P} (X \geq {@n-2@}) = {@ta@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>maxdp:4, mindp:4</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.0002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>764237321</deployedseed>
    <deployedseed>1944010320</deployedseed>
    <deployedseed>1841422752</deployedseed>
    <deployedseed>1653504969</deployedseed>
    <deployedseed>1376629497</deployedseed>
    <deployedseed>1425993691</deployedseed>
    <deployedseed>959389822</deployedseed>
    <deployedseed>1250447826</deployedseed>
    <deployedseed>1246740454</deployedseed>
    <deployedseed>95334597</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ev(ta+1.034,simp),4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
