<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>AlgMap-15.6</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the values of \(x\) and \(y\), from the following expressions. Give your answer as a set containing the pairs of points which satisfy the equation, e.g.  \(\left( \left\{[x_1, y_1],[x_2, y_2] \right\}\right)\): <br /> \[ \left. \begin{array}{rcl} {@cfx2*x^2+cfy2*y^2@} &amp;=&amp;{@k2@} \\ {@cfx*x+cfy*y@} &amp;=&amp;{@k1@} \end{array} \right\} \] <br /> \([x,y] \in \) [[input:ans]] [[validation:ans]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>By rearranging the linear equation, \({@cfx*x+cfy*y@} = {@k1@} \) we obtain: \[ y = {@(k1-cfx*x)/cfy@}, \] which we can substitute into the quadratic to give: \[ {@cfx2*x^2+cfy2*((k1-cfx*x)/cfy)^2@} = {@k2@}. \] This can be rearranged into the more familiar quadratic form: \[ {@(cfx2+(cfx^2/cfy^2))*x^2-(2*cfx*k1*x)/cfy^2+k1^2/cfy^2-k2@} = 0. \] Solving this gives \( x = 3\) or \(x = -1 \). By substituting these back into the linear equation you can find the corresponding values of y. ( \(y = -1\) and \(y = 3\) respectively). Therefore, the solution set is: \[ (x,y) \in \left\{ [3,-1] , [-1,3] \right\}. \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>cfx2:1;
cfy2:1;
cfx:1;
cfy:1;
k1:2;
k2:10;
ta:{[3,-1],[-1,3]};</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[\[ \left. \begin{array}{rcl}   {@cfx2*x^2+cfy2*y^2@} &=&{@k2@} \\
{@cfx*x+cfy*y@} &=&{@k1@} \end{array} \right\}  \]
\[ (x,y) \in \left\{ [3,-1]  , [-1,3]  \right\}. \]]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{[x,y],...}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>1-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>1-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
