<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Diff_Chain_2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \[\frac{dy}{dx}\] when \[y=\cosh^{-1}\left(\sqrt{{@n1@} {@v@} + {@n2@}}\right) .\]</p>
<p> [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To find \(\frac{dy}{dx}\) when \[y=\cosh^{-1}\left(\sqrt{{@n1@} {@v@} + {@n2@}}\right) .\]</p>
<p>In order to solve this problem we will have to use the chain rule \(\frac{df}{d{@v@}}=\frac{df}{dg} \cdot\frac{dg}{dh}\cdot\frac{dh}{d{@v@}}\)  as follows:</p>
<p>Define \(h({@v@})={@n1@}{@v@}+{@n2@}\), \(g(h) = \sqrt{h} = h^{\frac{1}{2}}\) and \(f(g)=\cosh^{-1}\left(g\right)\).</p>
<p>So, \[\begin{array}{rcl}\displaystyle \frac{d}{d{@v@}}\left(\cosh^{-1}\left(\sqrt{{@n1@} {@v@} + {@n2@}}\right)\right) &amp;\displaystyle =&amp;\displaystyle \frac{df}{dg}\cdot \frac{dg}{dh}\cdot \frac{dh}{d{@v@}} \\ &amp;\displaystyle =&amp;\displaystyle \frac{d}{dg}(\cosh^{-1}\left(g\right))\cdot\frac{d}{dh}(h^{\frac{1}{2}})\cdot \frac{d}{d{@v@}}({@n1@} {@v@}+{@n2@}) \\ &amp;\displaystyle =&amp;\displaystyle \frac{1}{ \sqrt{g^2-1}}\cdot\frac{1}{2}(h^{-\frac{1}{2}})\cdot {@n1@} \\ &amp;\displaystyle =&amp;\displaystyle \frac{1}{ \sqrt{h-1}}\cdot\frac{@n1@}{2 \sqrt{h}} \\&amp;\displaystyle =&amp;\displaystyle {@correctanswer@}\end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
 n1:rand(8)+2;
 n2:rand(5)+2;
 
 f : acosh(sqrt(n1*v+n2)); 
 g: sqrt(n1*v+n2);
 correctanswer : (n1/2)/(sqrt(n1*v+n2)*sqrt(n1*v+n2-1));
 wronganswer1:(1/2)/(sqrt(n1*v+n2)*sqrt(n1*v+n2-1));
 wronganswer2:diff(g,v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\frac{dy}{dx}\) of \(y=\cosh^{-1}\left(\sqrt{{@n1@} {@v@} + {@n2@}}\right)\)  = \({@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>To differentiate a function of a function, you must use the chain rule! It is \[\frac{df}{dg} \cdot\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have applied the chain rule incorrectly. It is \[\frac{df}{dg} \cdot\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>895810538</deployedseed>
    <deployedseed>716563093</deployedseed>
    <deployedseed>340193990</deployedseed>
    <deployedseed>1958239785</deployedseed>
    <deployedseed>554484770</deployedseed>
    <deployedseed>121967129</deployedseed>
    <deployedseed>152348815</deployedseed>
    <deployedseed>1386225229</deployedseed>
    <deployedseed>570359358</deployedseed>
    <deployedseed>1312074877</deployedseed>
    <deployedseed>531956334</deployedseed>
    <deployedseed>2076096142</deployedseed>
    <deployedseed>1832070658</deployedseed>
    <deployedseed>1053274494</deployedseed>
    <deployedseed>1283934680</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
