<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Function_Implicit_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use implicit differentiation to find  \[\frac{dy}{dx}\] <br />for the graph of  \[x^{@n1@}+y^{@n2@}={@n3@}xy\]</p>
<p> [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To find \[\frac{dy}{dx}\]</p>
<p>for the graph of \[x^{@n1@}+y^{@n2@}={@n3@}xy\]<br />rearrange the equation to give</p>
<p>\[x^{@n1@}+y^{@n2@}-{@n3@}xy =0\]</p>
<p>Thus</p>
<p>\[\frac{d}{dx}\left(x^{@n1@}+y^{@n2@}-{@n3@}xy\right) =0\]</p>
<p>\[{@n1@}{@x^(n1-1)@}+{@n2@}{@y^(n2-1)@}\frac{dy}{dx}-{@n3@}(y+x\frac{dy}{dx})=0\]</p>
<p>\[\frac{dy}{dx}({@n2@}{@y^(n2-1)@}-{@n3@}x) = {@n3@}y-{@n1@}{@x^(n1-1)@}\]</p>
<p>\[\frac{dy}{dx} = \frac{{@n3@}y-{@n1@}{@x^(n1-1)@}}{{@n2@}{@y^(n2-1)@}-{@n3@}x}\] </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(4)+2;
 n2:rand(4)+2;
 n3:rand(8)+2;
 p : x^n1+y^n2=n3*x*y; 
 diffp : expand(n3*y-n1*x^(n1-1))/(n2*y^(n2-1)-n3*x);
 correctanswer:  diffp;
 wronganswer1:n1*x^(n1-1)-n3*y;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\frac{dy}{dx} = {@diffp@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1972866120</deployedseed>
    <deployedseed>1464040384</deployedseed>
    <deployedseed>150841961</deployedseed>
    <deployedseed>1572433701</deployedseed>
    <deployedseed>1368593220</deployedseed>
    <deployedseed>1840037134</deployedseed>
    <deployedseed>271345436</deployedseed>
    <deployedseed>734097231</deployedseed>
    <deployedseed>225473005</deployedseed>
    <deployedseed>1454371294</deployedseed>
    <deployedseed>1744460672</deployedseed>
    <deployedseed>1568402037</deployedseed>
    <deployedseed>272119849</deployedseed>
    <deployedseed>836553183</deployedseed>
    <deployedseed>2037587730</deployedseed>
    <deployedseed>1122139974</deployedseed>
    <deployedseed>2013832955</deployedseed>
    <deployedseed>1787062147</deployedseed>
    <deployedseed>134022094</deployedseed>
    <deployedseed>1517983897</deployedseed>
    <deployedseed>514227539</deployedseed>
    <deployedseed>1407979764</deployedseed>
    <deployedseed>709062961</deployedseed>
    <deployedseed>53207446</deployedseed>
    <deployedseed>435139229</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
