<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inequality_6</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve the following inequality: \[ {@p@} .\] <br /> Write your answer as a logical sequence of inequalities, such as \(x&gt;5\) and/or \(x&lt;=13\). <br /> <br /> [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>As \[{@p@}.\]</p>
<p>Then \[{@2*v^4-(1+n1^4)@} \le {@n1^4 -1@}\mbox{     and     }{@2*v^4-(1+n1^4)@} \ge {@-(n1^4 -1)@}.\]</p>
<p>Hence \[{@2*v^4@} \le {@n1^4 -1@}+{@(1+n1^4)@}\mbox{     and     }{@2*v^4@} \ge {@-(n1^4 -1)@}+{@(1+n1^4)@}.\]</p>
<p>\[{@2*v^4@} \le {@n1^4 -1+(1+n1^4)@}\mbox{     and     }{@2*v^4@} \ge {@-(n1^4 -1)+(1+n1^4)@}.\]</p>
<p>Dividing through by 2 gives</p>
<p>\[{@v^4@} \le {@(n1^4 -1+(1+n1^4))/2@}\mbox{     and     }{@v^4@} \ge {@(-(n1^4 -1)+(1+n1^4))/2@}.\]</p>
<p>Taking the 4<sup>th</sup> root gives</p>
<p>\[{@v@}\ge -{@sqrt(sqrt((n1^4 -1+(1+n1^4))/2))@}\mbox{   and  }{@v@}\le {@sqrt(sqrt((n1^4 -1+(1+n1^4))/2))@}\mbox{   and  }{@v@} \ge {@sqrt(sqrt((-(n1^4 -1)+(1+n1^4))/2))@}\mbox{   or   }{@v@} \le -{@sqrt(sqrt((-(n1^4 -1)+(1+n1^4))/2))@}\]</p>
<p> This is entered as \({@v@}&gt;={@-n1@} \mbox{   and   } {@v@}&lt;=-1 \mbox{   or    }{@v@}&gt;= 1 \mbox{   and   } {@v@}&lt;={@n1@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[v:x;
 n1:rand(4)+2;
 p:abs(1+n1^4 -2*v^4)<= n1^4 -1;
 correctanswer:v>=-n1 and v<=-1 or v>= 1 and v<=n1;
 wronganswer1:v>=-n1 and v<=-1 and v>= 1 and v<=n1;
 wronganswer2:v>-n1 and v<-1 or v> 1 and v<n1;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@p@} \le 0, {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>40</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>948004101</deployedseed>
    <deployedseed>1080165636</deployedseed>
    <deployedseed>1226586660</deployedseed>
    <deployedseed>775403031</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
