<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>20.2.1.7 The linearity property of the Laplace transformation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Determine the Laplace transform of the following functions.</p>
<ol class="HELM_parts">
    <li>\(\text{e}^{-{@aa@}t}u(t)=\) [[input:ans1]] [[validation:ans1]]
        <p>[[feedback:prt1]]</p>
    </li>
    <li>\(u(t-{@bb@})=\) [[input:ans2]] [[validation:ans2]]
        <p>[[feedback:prt2]]</p>
    </li>
    <li>\(\text{e}^{-{@cc@}t}\sin {@dd@}t.u(t)=\) [[input:ans3]] [[validation:ans3]]
        <p>[[feedback:prt3]]</p>
    </li>
    <li>\(({@ee@}\cos {@ff@}t-{@gg@}t^{@hh@}).u(t)=\) [[input:ans4]] [[validation:ans4]]
        <p>[[feedback:prt4]]</p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">

    <li>\({\cal L}\{\text{e}^{-{@aa@}t}u(t)\}= \dfrac{1}{s+{@aa@}}\qquad \) (Table 1, Rule 4 with \(a={@aa@}\)) </li>

    <li>You should obtain \({\cal L}\{u(t-{@bb@})\}= \displaystyle \int_{@bb@}^\infty \displaystyle \text{e}^{-st} \, \mathrm{d}t = \dfrac{\text{e}^{-{@bb@}s}}{s}\) (note the lower limit&nbsp; of the integration is \({@bb@}\)) since:
        <p>\({\cal L}\{u(t-{@bb@})\} = \displaystyle \int_0^\infty \displaystyle \text{e}^{-st}u(t-{@bb@}) \, dt = \displaystyle \int_0^{@bb@} \text{e}^{-st}u(t-{@bb@}) \, dt + \displaystyle \int_{@bb@}^\infty \displaystyle \text{e}^{-st}u(t-{@bb@}) \, dt\)
            </p>
        <p>In the first integral \(0 &lt; t &lt; {@bb@}\) and so \((t − {@bb@}) &lt; 0\), therefore \(u(t − {@bb@}) = 0\).</p>
        <p>In the second integral \({@bb@} &lt; t &lt; \infty\) and so \((t − {@bb@}) &gt; 0\), therefore \(u(t − {@bb@}) = 1\). Hence </p>
        <p>\(\qquad {\cal L}\{u(t-{@bb@})\} = 0 + \displaystyle \int_{@bb@}^\infty \text{e}^{-st} \, \mathrm{d}t = \left[\dfrac{\text{e}^{-st}}{(-s)} \right]_{@bb@}^\infty = \dfrac{\text{e}^{-{@bb@}s}}{s}\)
        </p>
    </li>

    <li>\({\cal L}\{\text{e}^{-{@cc@}t}\sin {@dd@}t.u(t)\} = \displaystyle {{@dd@}\over (s+{@cc@})^2+{@dd@}^2}\qquad\) (Table 1, Rule 7 with \(a={@cc@}\) and \(b={@dd@}\))</li>

    <li>By the Linearity Property of the Laplace Transform
        <p>\({\cal L}\{{@ee@} \cos {@ff@}t\,.\,u(t)-{@gg@}t^{@hh@}u(t)\} = {@ee@}{\cal L}\{\cos {@ff@}t\,.\, u(t)\}-{@gg@}{\cal L}\{t^{@hh@}u(t)\}\)
        </p>
        <p>\(= {@ee@}\left(\dfrac{s}{s^2+{@ff@}^2}\right) - {@gg@} \left(\dfrac{{@hh@}!}{{s}^{{@hh@}+1}}\right) = {@ee@}\left(\dfrac{s}{s^2+{@ff@}^2}\right)-\left(\dfrac{{@gg@} \cdot {@hh@}!}{{s}^{{@hh@}+1}}\right)\)</p>
        <p>(Table 1, Rules 3,6)</p>
    </li>

</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(8)+2;
bb: rand(8)+1;
cc: rand(5)+1;
dd: rand(5)+2;
ee: rand(10)+1;
ff: rand(6)+2;
gg: rand(5)+2;
hh: rand([3,4]);
ta1: 1/(s+aa);
ta2: e^(-bb*s)/s;
ta3: dd/((s+cc)^2+dd^2);
ta4: gg*s/(s^2+gg^2)-gg*hh!/s^(hh+1)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1){@ta1@}; 2){@ta2@}; 3){@ta3@}; 4){@ta4@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>13</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2067250273</deployedseed>
    <deployedseed>293724167</deployedseed>
    <deployedseed>550448764</deployedseed>
    <deployedseed>1908280232</deployedseed>
    <deployedseed>1103931174</deployedseed>
    <deployedseed>568877807</deployedseed>
    <deployedseed>47569311</deployedseed>
    <deployedseed>366145382</deployedseed>
    <deployedseed>1562228804</deployedseed>
    <deployedseed>1639924261</deployedseed>
    <deployedseed>926089855</deployedseed>
    <deployedseed>1219447500</deployedseed>
    <deployedseed>712799837</deployedseed>
    <deployedseed>710642363</deployedseed>
    <deployedseed>1014512370</deployedseed>
    <deployedseed>1760622811</deployedseed>
    <deployedseed>916735219</deployedseed>
    <deployedseed>1798680819</deployedseed>
    <deployedseed>1106365503</deployedseed>
    <deployedseed>634915770</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>s</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>e/s</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>3/s</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>3!</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
