<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limits_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>You should use theorems on limits and the result that \[ \lim_{{@v@}\rightarrow 0}{\left(\frac{\sin({@v@})}{@v@}\right)}=1. \] <br /> You should not use a Calculator or L'Hopital's rule. <br /> <br /> Evaluate the following limit: \[ \lim_{{@v@}\rightarrow 0}{\left(\frac{\sin({@n1@} {@v@})}{\sin({@n2@}{@v@})}\right)} .\] <br /> The limit = [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We have \[ \lim_{{@v@}\rightarrow 0}{\left(\frac{\sin({@n1@} {@v@})}{\sin({@n2@}{@v@})}\right)} = \lim_{{@v@}\rightarrow 0}{\left(\frac{\sin({@n1@} {@v@})}{{@n1@} {@v@}} \frac{{@n2@} {@v@}}{\sin({@n2@}{@v@})} \frac{{@n1@}{@v@}}{{@n2@}{@v@}}\right)} = 1\cdot1\cdot\frac{@n1@}{@n2@}= {@correctanswer@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
 n1:rand_with_step(2,9,1);
 n2:rand_with_step(11,19,1);
 f:sin(n1*v)/sin(n2*v);
 correctanswer:n1/n2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\lim_{{@v@}\rightarrow 0}{\left(\frac{\sin({@n1@} {@v@})}{\sin({@n2@}{@v@})}\right)}\) = {@correctanswer@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>limit</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>755191200</deployedseed>
    <deployedseed>1987251928</deployedseed>
    <deployedseed>361610374</deployedseed>
    <deployedseed>548169882</deployedseed>
    <deployedseed>1504426525</deployedseed>
    <deployedseed>1753941928</deployedseed>
    <deployedseed>1021896733</deployedseed>
    <deployedseed>1000836762</deployedseed>
    <deployedseed>1893832123</deployedseed>
    <deployedseed>162057896</deployedseed>
    <deployedseed>812124702</deployedseed>
    <deployedseed>2077406899</deployedseed>
    <deployedseed>1666608699</deployedseed>
    <deployedseed>1688793199</deployedseed>
    <deployedseed>2011396505</deployedseed>
    <deployedseed>1706274855</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>n1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>n2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
