<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limits_2a</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A \({\bf rational}\mbox{ } {\bf function}\) is a quotient of two polynomials. Evaluate the following limit of a rational function: \[ \lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p1@}{@p2@}\right)}. \] If the limit does not exist, enter the value 1111 as your answer.</p>
<p>The limit is: [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The limit of the function \[ \lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p1@}{@p2@}\right)}. \] <br /> is the same as the limit found by keeping only the highest power terms in the numerator and the denominator. Hence the limit is undefined.</p>
<p>The limit of the function can be found by dividing the numerator and denominator by {@v^(degr1)@}.</p>
<p>The function can thus be expressed as \[ \lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p3@}{@p4@}\right)}. \]</p>
<p>The limit of the numerator is {@constantp3@} because all the other terms tend to zero. The limit of the denominator is zero. Hence the limit is undefined.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[v:x;
 degr1:rand_with_step(3, 9, 1);
 q:0;
 p1:block(for i:1 while i<=degr1 do (q:q+v^i*(-(rand(9)+1))),return(q));
 q1:0;
 p2:block(for i:1 while i<=degr1-1 do (q1:q1+v^i*(rand(9)+1)),return(q1));
 set1:[inf, minf];
 ind1:rand(2)+1;
 f:p1/p2;
 answerlimit:limit(p1/p2,v,set1[ind1]);
 correctanswer: 1111;
 wronganswer: expand(coeff(p1,v,degr1)/coeff(p2,v,degr1-1));
 p3 : expand(p1/(v^(degr1)));
 p4 : expand(p2/(v^(degr1)));
 constantp3 : coeff(p3,v,0);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\lim_{{@v@} \rightarrow {@set1[ind1]@}}{\left(\frac{@p1@}{@p2@}\right)}={@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords> limit</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1835687184</deployedseed>
    <deployedseed>1710026659</deployedseed>
    <deployedseed>801368638</deployedseed>
    <deployedseed>1476070656</deployedseed>
    <deployedseed>110328219</deployedseed>
    <deployedseed>1300990986</deployedseed>
    <deployedseed>888860805</deployedseed>
    <deployedseed>21401879</deployedseed>
    <deployedseed>293421934</deployedseed>
    <deployedseed>1274887834</deployedseed>
    <deployedseed>938830140</deployedseed>
    <deployedseed>2041030236</deployedseed>
    <deployedseed>884757123</deployedseed>
    <deployedseed>1167154567</deployedseed>
    <deployedseed>1376569075</deployedseed>
    <deployedseed>1332207800</deployedseed>
    <deployedseed>871486051</deployedseed>
    <deployedseed>203278787</deployedseed>
    <deployedseed>1650815474</deployedseed>
    <deployedseed>1894684532</deployedseed>
    <deployedseed>1693653439</deployedseed>
    <deployedseed>422461280</deployedseed>
    <deployedseed>246445264</deployedseed>
    <deployedseed>1231207568</deployedseed>
    <deployedseed>246188241</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>set1[ind1]</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
