<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.1.4.6 Integration Over Rectangular Areas Exercise</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the following double integral.</p>
<p> \(I=\displaystyle \int_{@aa@}^{@bb@} \displaystyle \int_{@cc@}^{@dd@}st^{3}\, \mathrm{d}s \mathrm{d}t \) </p>
<p> \(I=\) [[input:ans1]] [[validation:ans1]]</p>
<p>If your answer is not an integer, leave it in rational form e.g. <code>23/5</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p> The integration takes place over the rectangle shown below. </p>
[[jsxgraph]] 
        var board = JXG.JSXGraph.initBoard(divid, {
            boundingbox: [{#smin#}, {#tmax#}, {#smax#}, {#tmin#}],
            keepaspectratio: true,
            axis: true,
            showCopyright: false,
			showNavigation: false,
			grid: false
		});
		
		// Hide tick labels and ticks on x-axis but label it
		board.defaultAxes.x.setAttribute({
		  ticks: {
			drawLabels: false, // hide numbers
			majorHeight: 0,
			minorHeight: 0
		  },
		  name: 's',
		  withLabel: true,
		  label: {
			position: 'rt',
			offset: [1, -10]
		  }
		});

		// Hide tick labels and ticks on y-axis but label it
		board.defaultAxes.y.setAttribute({
		  ticks: {
			drawLabels: false,
			majorHeight: 0,
			minorHeight: 0
		  },
		  name: 't',
		  withLabel: true,
		  label: {
			position: 'rt',
			offset: [-15, 0]
		  }
		});
		
		
		// Create points and rectangle
		var bl = board.create('point', [{#cc#},{#aa#}], {fixed: true, withLabel:false, visible: true, showInfobox: false, size: 3, face: 'o', fillColor: 'black', strokeColor: 'black'});
		var br = board.create('point', [{#dd#},{#aa#}], {fixed: true, withLabel:false, visible: true, showInfobox: false, size: 3, face: 'o', fillColor: 'black', strokeColor: 'black'})
		var tr = board.create('point', [{#dd#},{#bb#}], {fixed: true, withLabel:false, visible: true, showInfobox: false, size: 3, face: 'o', fillColor: 'black', strokeColor: 'black'});
		var tl = board.create('point', [{#cc#},{#bb#}], {fixed: true, withLabel:false, visible: true, showInfobox: false, size: 3, face: 'o', fillColor: 'black', strokeColor: 'black'});
		var rectangle = board.create('polygon', [bl, br, tr, tl],{name:"Rectangle",withLabel:false, color: 'white'});
		
		// Set all edges (borders) to black
		rectangle.borders.forEach(line => {
			line.setAttribute({strokeColor: 'darkblue'});
		});
		
		// Label coordinates dynamically
		[bl, br, tr, tl].forEach((pt) => {
			board.create('text', [
				() => pt.X() + 0.1,
				() => pt.Y() + 0.1,
				() => `(${pt.X()}, ${pt.Y()})`
			], {
				fontSize: 12,
				fixed: true,
				highlight: false
			});
		});
[[/jsxgraph]]
<p>We have </p>
<p>\(\qquad \displaystyle I=\displaystyle \int_{@aa@}^{@bb@} \displaystyle \int_{@cc@}^{@dd@}st^{3}\, \mathrm{d}s \mathrm{d}t \) </p>
<p>Computing the inner integral, we get </p>
<p>\(\qquad \displaystyle \int_{@cc@}^{@dd@}st^{3}\, \mathrm{d}s \) </p>
<p> \(\qquad \displaystyle = \left[\dfrac{s^2}{2}t^3\right]_{@cc@}^{@dd@}\) </p>
<p> \(\qquad \displaystyle = \dfrac{{@dd^2@}}{2}t^3 - \dfrac{{@cc^2@}}{2}t^3 \) </p>
<p> \(\qquad \displaystyle = {@t^3*ee@} \) </p>

<p> This can then be put into the outer integral to give </p>
<p>\(\qquad \displaystyle I = \displaystyle \int_{@aa@}^{@bb@} {@t^3*ee@} \, \mathrm{d}t \) </p>
<p> \(\qquad \displaystyle = \displaystyle \left[{@ee*t^4/4@}\right]_{@aa@}^{@bb@} \) </p>
<p> \(\qquad \displaystyle = {@(bb^4/4) * ee@} - \left({@(aa^4/4) * ee@}\right) \) </p>
<p> \(\qquad \displaystyle = {@(bb^4/4) * ee - ((aa^4/4) * ee)@} \) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand_with_step(-3,0,1);
bb: rand_with_prohib(0,3,[aa]);
cc: rand_with_step(-3,0,1);
dd: rand_with_prohib(0,3,[cc]);
tmax: bb + 1;
tmin: aa - 1;
smax: dd + 1;
smin: cc - 1;
ee: dd^2/2 - cc^2/2
exp: s*t^3;
inn: int(exp, s, cc, dd);
ta1: int(inn, t, aa, bb);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\[I=\int_{@aa@}^{@bb@}\int_{@cc@}^{@dd@}\left( st^{3}\right) \,\mathrm{d}s \mathrm{d}t = {@ev(ta1, simp)@}\]</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1444477294</deployedseed>
    <deployedseed>793744087</deployedseed>
    <deployedseed>690769773</deployedseed>
    <deployedseed>1282571454</deployedseed>
    <deployedseed>1950688769</deployedseed>
    <deployedseed>1025643859</deployedseed>
    <deployedseed>1161887571</deployedseed>
    <deployedseed>52415259</deployedseed>
    <deployedseed>744534512</deployedseed>
    <deployedseed>1617689175</deployedseed>
    <deployedseed>1666376427</deployedseed>
    <deployedseed>1574833663</deployedseed>
    <deployedseed>898743923</deployedseed>
    <deployedseed>1871019621</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1 + 1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>jsxgraph_solution</text>
</tag>
    </tags>
  </question>
</quiz>
