<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.1.6.5 Areas and Moments Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the following integral</p>
<p>\( I= \displaystyle \int_{0}^{@aa@}  \displaystyle \int_{0}^{@bb@} (s+1)^4 \, \mathrm{d}t \mathrm{d}s \) </p>
<p> \(I=\) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> We have </p>
<p>\(I=\displaystyle \int_{0}^{@aa@} \displaystyle \int_{0}^{@bb@} (s+1)^4 \, \mathrm{d}t \mathrm{d}s \) </p>
<p> Which, since the integrand is independent of \(t\), can be rewritten as </p>
<p> \( I=\left({@bb@} - 0\right) \displaystyle \int_{0}^{@aa@} (s+1)^4 \, \mathrm{d}s \) </p>
<p> Evaluating, we get </p>
<p> \(\begin{align}
I &= {@bb@}\left[\dfrac{(s+1)^5}{5}\right]_{0}^{@aa@} \\
&= {@bb@}\left( \dfrac{({@aa@}+1)^5}{5} - \dfrac{(0+1)^5}{5} \right) \\
&= {@bb@}\left( {@((aa+1)^5)/5 - 1/5@}\right) \\
&= {@ bb*(((aa+1)^5)/5 - 1/5)@}\end{align} \) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand(5) + 1;
bb: rand(5) + 1;
exp: (s+1)^4;
inn: int(exp, t, 0, bb);
ta1: int(inn, s, 0, aa);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p> \[I= \int_{0}^{@aa@} \int_{0}^{@bb@} (s+1)^4 \, \mathrm{d}t \mathrm{d}s =  {@ev(ta1, simp)@} \] </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>929746879</deployedseed>
    <deployedseed>1204422029</deployedseed>
    <deployedseed>901628157</deployedseed>
    <deployedseed>256837594</deployedseed>
    <deployedseed>1739568865</deployedseed>
    <deployedseed>1970764811</deployedseed>
    <deployedseed>57502665</deployedseed>
    <deployedseed>1120823315</deployedseed>
    <deployedseed>1650436756</deployedseed>
    <deployedseed>523793706</deployedseed>
    <deployedseed>333944028</deployedseed>
    <deployedseed>1616167289</deployedseed>
    <deployedseed>838528782</deployedseed>
    <deployedseed>1803652720</deployedseed>
    <deployedseed>1599132079</deployedseed>
    <deployedseed>1558905005</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1 + 1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
