<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.1.6.3 Areas and Moments Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the following integral</p>
<p>(Note: If you are struggling to see the upper integration limit of the outer integral on your device, it is: \(\displaystyle {@aa@}\)) </p>

<p>\(\displaystyle I=\displaystyle\int_{0}^{@aa@} \displaystyle \int_{0}^{@bb@} y\cos^2(x)\, \mathrm{d}y \mathrm{d}x \) </p>
<p> \(I=\) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> We have </p>
<p> \(\displaystyle I=\displaystyle\int_{0}^{@aa@} \displaystyle \int_{0}^{@bb@} y\cos^2(x)\, \mathrm{d}y \mathrm{d}x \) </p>
<p> Which we can rewrite as </p>
<p>  \(\displaystyle I=\displaystyle\int_{0}^{@aa@} \cos^2(x)\, \mathrm{d}x \times \displaystyle \int_{0}^{@bb@} y\, \mathrm{d}y \) </p>
<p> Recalling the identity \(\cos^2(x) = \dfrac{1+\cos(2x)}{2}\), we can write </p>
<p> \(\displaystyle I=\displaystyle\int_{0}^{@aa@} \dfrac{1+\cos(2x)}{2}\, \mathrm{d}x \times \displaystyle \int_{0}^{@bb@} y\, \mathrm{d}y \) </p>
<p> \(\displaystyle =\displaystyle\dfrac{1}{2}\int_{0}^{@aa@} (1+\cos(2x))\, \mathrm{d}x \times \displaystyle \int_{0}^{@bb@} y\, \mathrm{d}y \) </p>


<p> Evaluating, we get </p>
<p>\(\begin{align} I &= \dfrac{1}{2} \left(\left[x + \dfrac{1}{2}\sin(2x)\right]_{0}^{@aa@} \times \left[\dfrac{y^2}{2}\right]_{0}^{@bb@}\right) \\ 
&=  \dfrac{1}{2}\left({@aa@} + \dfrac{1}{2} \sin\left({@2*aa@}\right) - 0\right) \times \left(\dfrac{{@bb@}^2}{2} - 0\right) \\ 
&=  {@(1/4)*aa*bb^2 + (1/8)*sin(2*aa)*bb^2@} \end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>cc: rand(4) + 1;
dd: rand(4) + 1;
aa: cc*(pi)/dd;
bb: rand(2) + 1;
exp: y*(cos(x))^2;
inn: int(exp, y, 0, bb);
ta1: int(inn, x, 0, aa);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\[I=\int_{0}^{@aa@}\int_{0}^{@bb@} y\cos^2(x) \,\mathrm{d}y \mathrm{d}x= {@ev(ta1, simp)@}\]</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>293195028</deployedseed>
    <deployedseed>277994561</deployedseed>
    <deployedseed>969994041</deployedseed>
    <deployedseed>1758952091</deployedseed>
    <deployedseed>618181985</deployedseed>
    <deployedseed>1775982666</deployedseed>
    <deployedseed>1005967625</deployedseed>
    <deployedseed>1107731817</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1 + 1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
