<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>27.1.6.4 Areas and Moments Exercise</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the following integral</p>
<p>\(\displaystyle I=\displaystyle\int_{@aa@}^{@bb@} \displaystyle \int_{@cc@}^{@dd@} y^2 \, \mathrm{d}y \mathrm{d}x \) </p>
<p> \(I=\) [[input:ans1]] [[validation:ans1]]</p>
<p>If your answer is not an integer, leave it in rational form e.g. <code>23/5</code>.<br></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> We have </p>
<p>\(\displaystyle I=\displaystyle\int_{@aa@}^{@bb@} \displaystyle \int_{@cc@}^{@dd@} y^2 \, \mathrm{d}y \mathrm{d}x \) </p>
<p> Which, since the integrand is independent of \(x\), can be rewritten as </p>
<p> \(\displaystyle I = \left({@bb@}-({@aa@})\right) \times \displaystyle \int_{@cc@}^{@dd@} y^2 \, \mathrm{d}y \) </p>
<p> Evaluating, we get </p>
<p> \(\begin{align}
I &= {@bb-aa@} \times \left[\dfrac{y^3}{3}\right]_{@cc@}^{@dd@} \\
&= {@bb-aa@} \times \left(\dfrac{{@dd@}^3}{3} - \dfrac{({@cc@})^3}{3}\right) \\
&= {@(bb-aa)*(dd^3/3 - cc^3/3)@} \end{align} \) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand_with_step(-8,-1,1);
bb: rand_with_step(1,3,1);
cc: rand_with_step(-3,-1,1);
dd: rand_with_step(1,3,1);
exp: y^2
inn: int(exp, y, cc, dd)
ta1: int(inn, x, aa, bb);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\[I=\int_{@aa@}^{@bb@}\int_{@cc@}^{@dd@} y^2 \,\mathrm{d}y \mathrm{d}x = {@ev(ta1, simp)@}\]</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>506432800</deployedseed>
    <deployedseed>636383036</deployedseed>
    <deployedseed>1422955700</deployedseed>
    <deployedseed>213043628</deployedseed>
    <deployedseed>712466090</deployedseed>
    <deployedseed>89985552</deployedseed>
    <deployedseed>1437869506</deployedseed>
    <deployedseed>406134300</deployedseed>
    <deployedseed>2145147675</deployedseed>
    <deployedseed>715369682</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
