<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.2.4.16 Vector field basic volume integral</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate \(\displaystyle\iiint_V \underline{\nabla} \cdot \underline{F} dV\) when \(\underline{F}\) is the vector field \({@aa@}yz\underline{i} + {@bb@}xy\underline{j}\) and \(V\) is the unit cube \(0 \leq x \leq 1\), \(0 \leq y \leq 1\), \(0 \leq z \leq 1\)</p>

<p>\(\displaystyle \iiint_V \underline{\nabla} \cdot \underline{F} dV =\) [[input:ans1]][[validation:ans1]].</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We aim to evaluate \(\displaystyle \iiint_V \underline{\nabla} \cdot \underline{F} \mathrm{d}V\) when \(\underline{F}\) is the vector field \({@aa@}yz\underline{i} + {@bb@}xy\underline{j}\) and \(V\) is the unit cube \(0 \leq x \leq 1\), \(0 \leq y \leq 1\), \(0 \leq z \leq 1\).</p>

<p>We will first calculate \(\underline{\nabla} \cdot \underline{F}\):
\[
\begin{aligned}
\underline{\nabla} \cdot \underline{F} &= \frac{\partial}{\partial x}({@aa@}yz) + \frac{\partial}{\partial y}({@bb@}xy) + \frac{\partial}{\partial z}(0)\\
&= 0 + {@bb@}x + 0\\
&={@bb@}x
\end{aligned}
\]
Now Calculating the integral with the bounds defined in the question:
\[
\begin{aligned}
\iiint_V \underline{\nabla} \cdot \underline{F} \mathrm{d}V &= \int^{z=1}_{z=0} \int^{y=1}_{y=0} \int^{x=1}_{x=0} {@bb@}x\ \mathrm{d}x\ \mathrm{d}y\ \mathrm{d}z\\
&= \int^{z=1}_{z=0} \int^{y=1}_{y=0} \left[{@First_int_indef@}\right]^{x=1}_{x=0}\ \mathrm{d}y\ \mathrm{d}z\\
&= \int^{z=1}_{z=0} \int^{y=1}_{y=0}{@First_int_def@} \ \mathrm{d}y\ \mathrm{d}z\\
&= \int^{z=1}_{z=0} \left[{@Second_int_indef@}\right]^{y=1}_{y=0}\ \mathrm{d}z\\
&= \int^{z=1}_{z=0} {@Second_int_def@}\ \mathrm{d}z\\\\
&=\left[{@Third_int_indef@}\right]^{z=1}_{z=0}\\
&={@Third_int_def@}
\end{aligned}
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(5)+2;
bb:rand(5)+2;

expression:bb*x;
First_int_indef:integrate(expression, x);
First_int_def:integrate(expression, x, 0, 1);
Second_int_indef:integrate(First_int_def, y);
Second_int_def:integrate(First_int_def, y, 0, 1);
Third_int_indef:integrate(First_int_def, z);
Third_int_def:integrate(First_int_def, z, 0, 1);

ta:Third_int_def;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>\( \underline{F} = {@aa@}yz\underline{i} + {@bb@}xy\underline{j}\ \text{ so } \displaystyle \iiint_V \underline{\nabla} \cdot \underline{F} dV = {@ta@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>onum</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1129005304</deployedseed>
    <deployedseed>196871889</deployedseed>
    <deployedseed>1724274584</deployedseed>
    <deployedseed>1210774788</deployedseed>
    <deployedseed>402000434</deployedseed>
    <deployedseed>1819730375</deployedseed>
    <deployedseed>1175590412</deployedseed>
    <deployedseed>463433037</deployedseed>
    <deployedseed>249236812</deployedseed>
    <deployedseed>1330115717</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>fail</description>
      <testinput>
        <name>ans1</name>
        <value>ta+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
