<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>29.2.3.9 Vector field dot prod surface Integral</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>For \(\underline{F} = ({@x^aa@} + {@y^bb@})\underline{i} + (x^2 + z^2)\underline{j} + {@cc@}xz\underline{k}\) and \(S\) the square bounded by \((1, 0, 1)\), \((1, 0, -1)\), \((-1, 0, -1)\) and \((-1, 0, 1)\) find the integral \(\displaystyle \int_S \underline{F} \cdot \underline{\mathrm{d}S}\)</p>

<p>\(\displaystyle \int_S \underline{F} \cdot \underline{\mathrm{d}S} =\) [[input:ans1]][[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We aim to evaluate the integral \(\displaystyle \int_S \underline{F} \cdot \underline{\mathrm{d}S}\) where \(\underline{F} = ({@x^aa + y^bb@})\underline{i} + (x^2 + z^2)\underline{j} + {@cc@}xz\underline{k}\) and \(S\) is the square bounded by \((1, 0, 1)\), \((1, 0, -1)\), \((-1, 0, -1)\) and \((-1, 0, 1)\).</p>

<p>We first aim to find the vector \(\underline{\mathrm{d}S}\). As the square is in the \(x-z\) plane, we know that \(\underline{\hat{n}} = \underline{j}\), so we have \(\underline{\mathrm{d}S} =\underline{j} \mathrm{d}x\mathrm{d}z\). We can then easily determine the outer bounds to be \(z=1\) as an upper bound, and \(z=-1\) as a lower bound. The inner bounds are \(x=1\) and \(x=-1\) for upper and lower bounds respectively. We can now evaluate the integral as follows:
\[
\begin{aligned}
\int_S \underline{F} \cdot \underline{\mathrm{d}S} &= \int^{z=1}_{z=-1} \int^{x=1}_{x=-1} \left[(({@x^aa + y^bb@})\underline{i} + (x^2 + z^2)\underline{j} + {@cc@}xz\underline{k}) \cdot \underline{j} \right]\ \mathrm{d}x\ \mathrm{d}z\\
&= \int^{z=1}_{z=-1} \int^{x=1}_{x=-1} {@expression@}\ \mathrm{d}x\ \mathrm{d}z\\
&= \int^{z=1}_{z=-1} \left[{@First_int_indef@}\right]^{x=1}_{x=-1}\ \mathrm{d}z\\
&= \int^{z=1}_{z=-1} {@First_int_def@}\ \mathrm{d}z\\
&= \left[{@Second_int_indef@}\right]^{z=1}_{z=-1}\\
&= {@Second_int_def@}
\end{aligned}
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(3)+1;
bb:rand(3)+1;
cc:rand(5)+1;


expression:x^2+z^2;
First_int_indef:integrate(expression, x);
First_int_def:integrate(expression, x, -1, 1);
Second_int_indef:integrate(First_int_def, z);
Second_int_def:integrate(First_int_def, z, -1, 1);

ta:Second_int_def;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>\(\underline{F} = (x^{@aa@} + y^{@bb@})\underline{i} + (x^2 + z^2)\underline{j} + {@cc@}xz\underline{k}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>onum</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>-ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>423273792</deployedseed>
    <deployedseed>810763299</deployedseed>
    <deployedseed>418610162</deployedseed>
    <deployedseed>1049948677</deployedseed>
    <deployedseed>301922819</deployedseed>
    <deployedseed>27155764</deployedseed>
    <deployedseed>884418728</deployedseed>
    <deployedseed>693796679</deployedseed>
    <deployedseed>493882883</deployedseed>
    <deployedseed>1501221768</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>fail</description>
      <testinput>
        <name>ans1</name>
        <value>ta+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
