<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>HELM 33.1 E3 BVP System of Equations, Natural boundary conditions</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[**33.1 E3 1D Boundary Value Problem, System of Equations, Natural Boundary Conditions**

<p>Let \( y = y(x) \) be a solution to the boundary value problem</p>
 
<p>\(\displaystyle {@p(x)*'diff('y,x,2) + q(x)*'diff('y,x) + r(x)*'y(x) = s(x)@} \quad 0 < x < {@L@} \)</p>

 <p>with boundary conditions:</p>
<ul>
    <li>\( y'(0) = {@dy0@}\)</li>
    <li>\( y({@L@}) = {@yL@} \)</li>
</ul>
<p>(Note the derivative boundary condition at \( x = 0 \).)</p>
<p>Using a mesh width of \( h = {@h@} \), obtain a central difference approximation to the differential equation and hence find the system of equations satisfied by \( y_j \approx y(jh) \), \( j = 0, 1, 2 \).</p>

<p>Specify the values to 3 d.p. (trailing zeros can be ommitted).</p>

<table border="0">
 <tbody>
  <tr>
   <td>[[input:S_A]]</td>
   <td>[[validation:S_A]]</td>
   <td>\( \left[ \begin{array}{c} y_0 \\ y_1 \\ y_2 \end{array} \right] = \) </td>
   <td>[[input:S_b]]</td>
   <td>[[validation:S_b]]</td>
  </tr>
 </tbody>
</table>
<p>[[feedback:GLS]]</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024072400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Define parameters */
h: 0.5;
nx: 3;
L: nx*h;
dy0: 2;
yL: 1;
dy0: rand_with_step(-2,2,1);
yL: rand_with_step(-2,2,1);
p0: rand_with_step(1,3,1);
r0: rand_with_step(1,3,1);
p(x):=p0;
q(x):= x^2;
r(x):= r0;
s(x):= x;

/* boudary conditions */
yg:  y1-2*h*dy0;
y3: yL;

/* Define the coefficients of the ODE at the various points */
X: makelist(j*h,j,-1,nx);
Y: ratsimp([yg, y0, y1, y2, yL]);
P: makelist(p(X[j]),j,1,nx+2);
Q: makelist(q(X[j]),j,1,nx+2);
R: makelist(r(X[j]),j,1,nx+2);
S: makelist(s(X[j]),j,1,nx+2);

/* Define the equations using central difference approximations */
eq: expand(makelist(Y[j-1]*(P[j]-h/2*Q[j])+Y[j]*(h^2*R[j]-2*P[j])+Y[j+1]*(P[j]+h/2*Q[j])-
    h^2*S[j],
    j,2,nx+1));
A11: diff(eq[1], y1);

A:jacobian(eq,[y0,y1,y2]);
b:radcan(-transpose(matrix(eq))+A.transpose(matrix([y0,y1,y2])));


sol: invert(A).b;

stack_include("https://raw.githubusercontent.com/mkraska/meclib/main/Maxima/fb_value_EN.mac");

toList(m):= args(transpose(m))[1];]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[\(\displaystyle {@p(x)*'diff('y,x,2) + q(x)*'diff('y,x) + r(x)*'y(x) = s(x)@} \quad 0 < x < {@L@} \), \(y'(0) = {@dy0@}\), \( y({@L@}) = {@yL@}\)  {@float(sol)@}]]></text>
    </questionnote>
    <questiondescription format="markdown">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>S_A</name>
      <type>matrix</type>
      <tans>A</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>S_b</name>
      <type>matrix</type>
      <tans>b</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>GLS</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text>SA: float(round(S_A*1000)/1000);
TA: float(round(A*1000)/1000);
Sb: float(round(S_b*1000)/1000);
Tb: float(round(b*1000)/1000);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Richtige Lösung?</description>
        <answertest>NumRelative</answertest>
        <sans>toList(invert(S_A).S_b)</sans>
        <tans>toList(sol)</tans>
        <testoptions>0.005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>GLS-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[Your system provides the correct solution:
<br>
 {@float(round(1000*invert(S_A).S_b)/1000)@}<br>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>GLS-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[Your system has the solution
<br>
 {@float(round(1000*invert(S_A).S_b)/1000)@}. 
<br>
The correct solution is 
<br>
{@float(round(1000*sol)/1000)@}.<br>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Fehler in A</description>
        <answertest>AlgEquiv</answertest>
        <sans>SA</sans>
        <tans>TA</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>GLS-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>GLS-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Fehler in b</description>
        <answertest>AlgEquiv</answertest>
        <sans>Sb</sans>
        <tans>Tb</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>GLS-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>GLS-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>579624181</deployedseed>
    <deployedseed>1343186453</deployedseed>
    <deployedseed>1134831928</deployedseed>
    <deployedseed>2076650385</deployedseed>
    <deployedseed>1366787899</deployedseed>
    <deployedseed>1910156339</deployedseed>
    <deployedseed>336030133</deployedseed>
    <deployedseed>487474868</deployedseed>
    <deployedseed>1130909704</deployedseed>
    <deployedseed>624011609</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_A</name>
        <value>A</value>
      </testinput>
      <testinput>
        <name>S_b</name>
        <value>b</value>
      </testinput>
      <expected>
        <name>GLS</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>GLS-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_A</name>
        <value>A</value>
      </testinput>
      <testinput>
        <name>S_b</name>
        <value>b</value>
      </testinput>
      <expected>
        <name>GLS</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>GLS-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
