<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text><![CDATA[Discrete mathematics & graph theory - 2 - JSXGraphs]]></text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[[[jsxgraph]]
    /* boundingbox, defined by limits:[left, top, right, bottom] */
    var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-1.2, 1.2, 1.2, -1.2], axis: false, showCopyright: false});

   
    var pts = {#pts#}; /*set up points as defined in question variables: could be randomised! 
                                   Notice the syntax STACK uses for putting the _value_ of a variable into the text before display. */
    var boardpts = new Array();
    var arrayLength = pts.length;
    for (var i = 0; i < arrayLength; i++) {
        boardpts.push(board.create('point', pts[i],{fixed:false,showInfobox:false}));/*Defines traits of the points*/
    }
/*Regular graph*************************************/
    var edges = {#edges#}; /*set up edges as defined in question variables: could be randomised!*/
    var boardedges = new Array();
    var arrayLength = edges.length;
    for (var i = 0; i < arrayLength; i++) {
        boardedges.push(board.create('line', [boardpts[edges[i][0]], boardpts[edges[i][1]]],{straightFirst:false, straightLast:false}));
/***********************************************************/
/*Directed graph, commented out*********************/
/*    var edges = {#edges#};
    var boardedges = new Array();
    var arrayLength = edges.length;
    for (var i = 0; i < arrayLength; i++) {
        boardedges.push(board.create('arrow', [boardpts[edges[i][0]], boardpts[edges[i][1]]]));*/
/***********************************************************/
    }
[[/jsxgraph]]
</p><br><p>What is the chromatic number of the above graph? Try rearranging the points!
[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>To calculate chromatic number we can first check if the graph is empty, bipartite or a known graph. We can simply see that the graph is not empty. We can rearrange the graph into a clear bipartite configuration, thus the chromatic number is 2.</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024092500</text>
    </stackversion>
    <questionvariables>
      <text>/* A list of points, which are lists of coordinates.*/
pts:[[1.0,1.0],[1.0,0.0],[-1.0,0.0],[0.0,-1.0],[-1.0,-1.0],[-1.0,1.0]];

/* A list of edge connections. */
edges:[[0,5],[0,1],[1,2],[1,3],[4,3],[4,0],[2,4],[2,5]];

/*Answer to question*/
ta1:2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Hint: Try rearranging the points.</text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
