<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>14.1.1.6 Integrals</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>Write down, but do not evaluate, the integral defined by the limit as \(\delta x \rightarrow 0\), or \(\delta t \rightarrow 0\) of the following sums.</p>

<p>Recall that by \(\displaystyle{ \sum_{x=a}^b y(x) \delta x}\) we mean to break the interval \([a,b]\) into subintervals of width at most \(\delta x\) and then sum the area of suitable rectangles of width \(\delta x\). Review an earlier key point if you
    are unsure of the meaning of the notation.</p>

<ol class="HELM_parts">

    <li>
        <p>\(\displaystyle{\lim_{\delta x \rightarrow 0}{@sum1@} \delta x}\).</p>
        <p>Write out your answer in the form \(\displaystyle{@gen_int@}\). Then </p>
        <p> \(f(x) =\) [[input:ans1]] [[validation:ans1]]   \(a = \) [[input:ans3]] [[validation:ans3]]   \(b = \) [[input:ans2]] [[validation:ans2]][[feedback:prtA]]</p>
    </li>

    <li>
        <p>\(\displaystyle{\lim_{\delta x \rightarrow 0}{@sum2@} \delta x}\).</p>
        <p>Write out your answer in the form \(\;\displaystyle{@gen_int@}\). Then </p>
        <p>\(f(x) =\) [[input:ans5]] [[validation:ans5]]   \(a = \) [[input:ans7]] [[validation:ans7]]   \(b = \) [[input:ans6]] [[validation:ans6]][[feedback:prtB]]</p>
    </li>

    <li>
        <p>\(\displaystyle{\lim_{\delta t \rightarrow 0}{@sum3@} \delta t}\).</p>
        <p>Write out your answer in the form \(\displaystyle{@gent_int@}\). Then </p>
        <p> \(f(t) =\) [[input:ans8]] [[validation:ans8]]   \(a = \) [[input:ans10]] [[validation:ans10]]   \(b = \) [[input:ans9]] [[validation:ans9]][[feedback:prtC]]</p>
    </li>

    <li>
        <p>\(\displaystyle{\lim_{\delta x \rightarrow 0}{@sum4@} \delta x}\).</p>
        <p>Write out your answer in the form \(\:\displaystyle{@gen_int@}\). Then </p>
        <p>\(f(x) =\) [[input:ans12]] [[validation:ans12]]   \(a = \) [[input:ans14]] [[validation:ans14]]   \(b = \) [[input:ans13]] [[validation:ans13]] [[feedback:prtD]]</p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">

    <li>
        <p>For \(\displaystyle{\lim_{\delta x \rightarrow 0}{@sum1@} \delta x},\) you should get the integral \(\displaystyle\int_{{@lb1@}}^{{@ub1@}} {@f1@} \) \(\mathrm{d}x.\)</p>
    </li>
    <li>
        <p>For \(\displaystyle{\lim_{\delta x \rightarrow 0}{@sum2@} \delta x},\) you should get the integral \(\displaystyle\int_{{@lb2@}}^{{@ub2@}}{@f2@} \) \(\mathrm{d}x.\)</p>
    </li>

    <li>
        <p>For \(\displaystyle{\lim_{\delta t \rightarrow 0}{@sum3@} \delta t},\) you should get the integral \(\displaystyle\int_{{@lb3@}}^{{@ub3@}} {@f3@} \) \(\mathrm{d}t.\)</p>
    </li>

    <li>
        <p>For \(\displaystyle{\lim_{\delta x \rightarrow 0}{@sum4@} \delta x},\) you should get the integral \(\displaystyle \int_{{@lb4@}}^{{@ub4@}} {@f4@} \) \(\mathrm{d}x.\)</p>
    </li>


</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>gen_int: int(f(x),x,a,b) 
gent_int: int(f(t),t,a,b)


nn1: 2 + rand(9)
f1: x^nn1
lb1: 0
ub1: 1
simp:false
sum1: sum(f1,x,lb1,ub1)

simp:true
nn2: 2 + rand(9)
k2: nn2*%pi
f2: k2*x^2
lb2: 0
ub2: 4
simp:false
sum2: sum(f2,x,lb2,ub2)

simp:true
nn3: 2 + rand(9)
f3: t^nn3
lb3: 0
ub3: 1
simp:false
sum3: sum(f3,t,lb3,ub3)

simp:true
nn4: 2 + rand(9)
k4: nn4*m
f4: k4*x^2
lb4: 0
ub4: 1
simp:false
sum4: sum(f4,x,lb4,ub4)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@sum1@}, {@sum2@}, {@sum3@}, {@sum4@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>f1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>lb3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>algebraic</type>
      <tans>f4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>algebraic</type>
      <tans>ub4</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans14</name>
      <type>algebraic</type>
      <tans>lb4</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ub1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>lb1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>f2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ub2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>lb2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>f3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>ub3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prtA</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>f1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3400000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtA-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtA-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>lb1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtA-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtA-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ub1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtA-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtA-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtB</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>f2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3400000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtB-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtB-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>lb2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtB-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtB-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ub2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtB-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtB-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtC</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>f3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3400000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtC-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtC-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>lb3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtC-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtC-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>ub3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtD</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>f4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3400000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtD-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtD-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans14</sans>
        <tans>lb4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtD-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtD-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans13</sans>
        <tans>ub4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtD-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtD-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>930633319</deployedseed>
    <deployedseed>849376623</deployedseed>
    <deployedseed>975489173</deployedseed>
    <deployedseed>1026865830</deployedseed>
    <deployedseed>1114389735</deployedseed>
    <deployedseed>845255576</deployedseed>
    <deployedseed>902500094</deployedseed>
    <deployedseed>771184950</deployedseed>
    <deployedseed>1847386830</deployedseed>
    <deployedseed>1242637180</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>f1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>lb3</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>f4</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>ub4</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>lb4</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ub1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>lb1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>f2</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ub2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>lb2</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>f3</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ub3</value>
      </testinput>
      <expected>
        <name>prtA</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtA-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtB</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtB-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtC-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtD</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtD-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(f1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ev(lb3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>ev(f4+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>ev(ub4+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>ev(lb4+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ub1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(lb1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ev(f2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ev(ub2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ev(lb2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ev(f3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ev(ub3+1,simp)</value>
      </testinput>
      <expected>
        <name>prtA</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtA-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtB</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtB-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtC</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtC-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtD</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtD-3-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
