<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>14.1.1.5 Area Using Limit of Sum 2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>Find the area under \(y = {@yy@} \) from \(x = 0\) to \(x = 2\) using the limit of a sum.</p>
<p> (Give a rational answer).</p>
<p>Area = [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<figure class="HELM_figure">
    [[jsxgraph]]
/* Variables from Maxima */
var a1 = {#a1#};
var yy = '{#yy#}';
var BB = [-0.5, 4 * a1 + 3, 2.5, -3];

/* JSX code */
var board = JXG.JSXGraph.initBoard(divid,
  {boundingbox: BB, axis: true, showCopyright: false, showNavigation: false});
var f = board.jc.snippet(yy, true, 'x', true);

/* Change the code below to suit your needs. */
board.defaultAxes.x.removeTicks(board.defaultAxes.x.defaultTicks);
board.defaultAxes.y.removeTicks(board.defaultAxes.y.defaultTicks);
board.create('ticks', [board.defaultAxes.x, [0, 2]], {drawLabels: true});
board.create('ticks', [board.defaultAxes.y, [f(0), f(2)]], {drawLabels: true});
var plot = board.create('functiongraph', [f]);
var gf = board.create('functiongraph', [f]);
var i = board.create('integral', [[0.0, 2.0], gf],
  {curveLeft: {visible: false}, curveRight: {visible: false}, withLabel: false});
    [[/jsxgraph]]
</figure>
<p>Let \(A\) denote the area under the curve that we aim to find. Using integration as the limit of a sum, we have
    \[A = \int^{2}_0 {@yy@} \, \mathrm{d}x = \left[{@ta1@} \right]^{2}_0 = {@ta@}.\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>a1: 2 + rand(15)
yy: a1*x^2
ta1: int(yy,x)
ta: int(yy,x,0,2)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ \int {@yy@} \mathrm{d}x = {@ta1@} = {@ta@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>Input fraction</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords>[[BASIC-CALCULUS]]</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1559380857</deployedseed>
    <deployedseed>1742114250</deployedseed>
    <deployedseed>115178421</deployedseed>
    <deployedseed>1407176868</deployedseed>
    <deployedseed>213461664</deployedseed>
    <deployedseed>478593565</deployedseed>
    <deployedseed>1700707789</deployedseed>
    <deployedseed>536570703</deployedseed>
    <deployedseed>1946887905</deployedseed>
    <deployedseed>1844186570</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
