<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9.2.1.11 Points, position vectors and vectors between points (2D)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise<br></p>
<p>Given that \(A\) is the point \(({@a_x@},\,{@a_y@})\) and that \(B\) is the point \(({@b_x@},\,{@b_y@})\) find:<br></p>
<ol class="HELM_parts">
    <li>
        <p> \(\vec{OA} =\) [[input:ans1]] [[validation:ans1]] \(\underline{i}\) \(+\) [[input:ans2]][[validation:ans2]] \(\underline{j}\)</p>
      <p>[[feedback:Resulta1]]</p>
    </li>
    <li>
        <p>\(\vec{OB}\) = [[input:ans4]] [[validation:ans4]] \(\underline{i}\) \(+\) [[input:ans5]] [[validation:ans5]] \(\underline{j}\)</p><p>
      </p><p>[[feedback:Resulta2]]</p>
    </li>
    <li>
        <p>\( \vec{AB} = \) [[input:ans7]] [[validation:ans7]] \(\underline{i}\) \(+\) [[input:ans8]][[validation:ans8]] \(\underline{j}\)</p>
        <p>[[feedback:Result]]</p>
    </li>
    <li>
        <p> \( |\vec{AB}| = \) [[input:ans10]] [[validation:ans10]] [[feedback:Resultc]]</p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">
<li>The position vectors of the points \(A\) and \(B\) are \({@a_vec@}\) and \({@b_vec@}\) respectively.</li>
<li>Given that \(\vec{A}\) is \({@a_vec@}\) and that \(\vec{B}\) is \({@b_vec@}\) then the vector from A to B is <br><p></p><p>\[\begin{array}{rcl}\displaystyle \mathbf{r}&amp;\displaystyle =&amp;\displaystyle \left({@b_x@} - ({@a_x@})\right)\underline{i} + \left({@b_y@} - ({@a_y@})\right) \underline{j},\\ &amp;\displaystyle =&amp;\displaystyle {@AB_vec@}. \end{array}\]</p></li>
<li>The modulus of \(\vec{AB}\) is calculated using Pythagoras' theorem as \[ |\vec{AB}| = \sqrt{({@b_x_minus_a_x@})^2 + ({@b_y_minus_a_y@})^2} = {@modab@}\]
</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021120900</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[stack_reset_vars(true);
stackvectortex(ex):= block(sconcat("{\\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j);
ijk_vec: matrix([i],[j]);
a_sel:[2,3,4,5,6,7,8,9,10,11,12,13];
magnitude: rand_selection(a_sel,4);
a_x_magnitude: magnitude[1];
a_x_sign:2*rand(2)-1;
a_x:a_x_magnitude;
a_y_magnitude: magnitude[2];
a_y_sign:2*rand(2)-1;
a_y:a_y_magnitude*a_y_sign;
a_mat:matrix([a_x],[a_y])
a_vec: texboldatoms(dotproduct(ijk_vec,a_mat));


 b_x_magnitude:magnitude[3];
 b_x_sign:2*rand(2)-1;
 b_x:b_x_magnitude;
 b_y_magnitude: magnitude[4];
 b_y_sign:2*rand(2)-1;
 b_y:b_y_magnitude*b_y_sign;

b_mat:matrix([b_x],[b_y])
b_vec: texboldatoms(dotproduct(ijk_vec,b_mat));

 b_x_minus_a_x:b_x-a_x;
 b_y_minus_a_y:b_y-a_y;
AB_mat: matrix([b_x_minus_a_x],[b_y_minus_a_y]);
AB_vec: texboldatoms(dotproduct(ijk_vec, AB_mat));

r:[b_x_minus_a_x,b_y_minus_a_y];
r1:[a_x + b_x,b_y + a_y];
midpoint: [b_x_minus_a_x/2,b_y_minus_a_y/2];
a_ans:[a_x,a_y];
b_ans:[b_x,b_y];

modab: sqrt(b_x_minus_a_x^2 +  b_y_minus_a_y^2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Given \(A\) is \(({@a_x@},\,{@a_y@})\) and \(B\) is \(({@b_x@},\,{@b_y@})\) then
\[ \vec{AB} = {@AB_vec@}\] 
\[ |\vec{AB}| = {@modab@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>a_x</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>modab</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>a_y</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>b_x</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>b_y</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>b_x_minus_a_x</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans> b_y_minus_a_y</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>vector_answer:[ans7,ans8];
follow_thr:[ans4-ans1,ans5-ans2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answer</sans>
        <tans>r</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>VEC-PNT-TRUE</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>VEC-PNT-FALSE</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answer</sans>
        <tans>-r</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>NEGATIVE</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is the vector from B to A. Each component is the additive inverse of that in the vector from A to B.</p>
 <p>Your answer is \[\left({@a_x@} - {@b_x@}\right)\mathbf{i} + \left({@a_y@} - {@b_y@}\right)\mathbf{j} + \left({@a_z@} - {@b_z@}\right)\mathbf{k}\mbox{.}\] The vector from A to B is \[ \mathbf{r} = \left({@b_x@} - {@a_x@}\right)\mathbf{i} + \left({@b_y@} - {@a_y@}\right) \mathbf{j} + \left({@b_z@} - {@a_z@}\right)\mathbf{k}\mbox{.}\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>VEC-PNT-FALSE-1</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answer</sans>
        <tans>r1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ADD_COMPONENTS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is \[\left({@a_x@} + {@b_x@}\right)\mathbf{i} + \left({@a_y@} + {@b_y@}\right)\mathbf{j} + \left({@a_z@} + {@b_z@}\right)\mathbf{k}\mbox{.}\] The vector from A to B is \[ \mathbf{r} = \left({@b_x@} - {@a_x@}\right)\mathbf{i} + \left({@b_y@} - {@a_y@}\right) \mathbf{j} + \left({@b_z@} - {@a_z@}\right)\mathbf{k}\mbox{.}\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>VEC-PNT-FALSE-2</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answer</sans>
        <tans>midpoint</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>MIDPOINT</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is the vector from the origin to the midpoint between A and B. You may have used the wrong method.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>VEC-PNT-FALSE-3</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answer</sans>
        <tans>follow_thr</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer to this part is correct based on the answers you gave to part (a), however there is a mistake in one of those vectors.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>Resulta1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>vector_answera1:[ans1,ans2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answera1</sans>
        <tans>a_ans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Resulta1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Resulta1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>Resulta2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>vector_answera2:[ans4,ans5];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>vector_answera2</sans>
        <tans>b_ans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Resulta2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Resulta2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>Resultc</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>mod_fol_thr: sqrt(ans7^2+ans8^2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>modab</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Resultc-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Resultc-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>mod_fol_thr</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Resultc-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer here is correct for the value of \(\vec{AB}\) that you gave in part (b), however there is a mistake in that previous part.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Resultc-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2015228898</deployedseed>
    <deployedseed>1963238636</deployedseed>
    <deployedseed>1183764266</deployedseed>
    <deployedseed>16672289</deployedseed>
    <deployedseed>1869733464</deployedseed>
    <deployedseed>1074889779</deployedseed>
    <deployedseed>1860513936</deployedseed>
    <deployedseed>1527512526</deployedseed>
    <deployedseed>725835434</deployedseed>
    <deployedseed>2094234466</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>a_x</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>modab</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a_y</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>b_x</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>b_y</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>b_x_minus_a_x</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>b_y_minus_a_y</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>VEC-PNT-TRUE</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resultc</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resultc-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>a_x</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>modab</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a_y</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>b_x</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>b_y</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>-b_x_minus_a_x</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>-b_y_minus_a_y</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>NEGATIVE</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resultc</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resultc-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>a_x</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>sqrt((a_x + b_x)^2 + (a_y + b_y)^2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a_y</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>b_x</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>b_y</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>a_x + b_x</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>a_y + b_y</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>ADD_COMPONENTS</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resultc</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resultc-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>a_x</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>modab/2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a_y</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>b_x</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>b_y</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>b_x_minus_a_x/2</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>b_y_minus_a_y/2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>MIDPOINT</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resultc</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resultc-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>a_x</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>sqrt((-2*b_x - a_x)^2 + (-1*b_y - a_y)^2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a_y</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>-2*b_x</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>-1*b_y</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>-2*b_x - a_x</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>-1*b_y - a_y</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-5-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resulta1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>Resulta2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Resulta2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>Resultc</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Resultc-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
