<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Differential_Equation_9</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the most general solution to the differential equation \[{@n1@} \frac{d^2y}{dx^2} + {@bc@} \frac{dy}{dx} + {@cc@} y = e^{{@n4@}x}\] <br /> Use the letters \({\bf A}\) and \({\bf B}\) to denote the constants in your solution. <br /> y(x) = [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To find the complementary function, solve the differential equation \[{@n1@} \frac{d^2y}{dx^2} + {@bc@} \frac{dy}{dx} + {@cc@} y = 0\] Let \(y = e^{mx}\) and hence \(\frac{dy}{dx}=me^{mx}\) and \(\frac{d^2y}{dx^2} = m^2e^{mx}\) <br />Substituting for \(y\) gives \[{@n1@}m^2e^{mx} + {@bc@}me^{mx} + {@cc@}e^{mx} = 0\] and dividing throughout by \(e^{mx}\) and {@n1@} gives \[m^2 + {@bc/n1@}m + {@cc/n1@} = 0\] which can be solved to give {@m@} <br />The solution is then of the form \[y = A{@e^(m1*x)@}+Bx{@e^(m2*x)@}\] <br />To find the particular integral, let \(y = ke^{{@n4@}x}\) and hence \(\frac{dy}{dx}={@n4@}ke^{{@n4@}x}\) and \(\frac{d^2y}{dx^2} = {@n4*n4@}ke^{{@n4@}x}\) giving \[{@n1*n4*n4@}ke^{{@n4@}x} + {@bc*n4@}ke^{{@n4@}x} + {@cc@}ke^{{@n4@}x} = e^{{@n4@}x}\] Therefore \[{@n1*n4*n4+bc*n4+cc@}k = 1\] and \[k= \frac{1}{@n1*n4*n4+bc*n4+cc@}\] Thus the solution to the differential equation is \[y = A{@e^(m1*x)@}+Bx{@e^(m2*x)@}+\frac{1}{@n1*n4*n4+bc*n4+cc@}e^{@n4*x@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(8)+2;
 n2:rand(6)+1;
 n3:rand(6)+1;
 n4:rand(5)+2;
 bc:n1*2*n2;
 cc:n1*n2^2;
 msol:solve(n1*m^2 + bc*m + cc = 0,m);
 m1:rhs(msol[1]);
 m2:m1;
 correctanswer:A*e^(m1*x)+B*x*e^(m2*x) + e^(n4*x)/(n1*n4*n4+bc*n4+cc);
 altanswer:B*e^(m1*x)+A*x*e^(m2*x) + e^(n4*x)/(n1*n4*n4+bc*n4+cc);
 ode:n1*'diff(y,x,2) + bc*'diff(y,x) + cc*y - exp(n4*x);
 wronganswer1:A*e^(m1*x);
 wronganswer2:A*e^(m1*x)+e^(n4*x)/(n1*n4*n4+bc*n4+cc);
 wronganswer3:B*x*e^(m2*x) + e^(n4*x)/(n1*n4*n4+bc*n4+cc);
 wronganswer4:e^(n4*x)/(n1*n4*n4+bc*n4+cc);
 wronganswer5:A*e^(m1*x)+B*x*e^(m2*x) + e^(n4*x)/(n1*n4*n4+bc*n4+cc-1);
 wronganswer6:a*e^(m1*x)+b*x*e^(m2*x) + e^(n4*x)/(n1*n4*n4+bc*n4+cc);
 wronganswer7:5*x*e^(m2*x) + 3*e^(m1*x) + e^(n4*x)/(n1*n4*n4+bc*n4+cc);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@n1@} \frac{d^2y}{dx^2} + {@bc@} \frac{dy}{dx} + {@cc@} y = e^{{@n4@}x}\) gives \(y = A{@e^(m1*x)@}+Bx{@e^(m2*x)@}+\frac{1}{@n1*n4*n4+bc*n4+cc@}e^{@n4*x@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa1:ev(subst(y=ans1,ode),simp);
 sa2:ev(sa1,nouns);
 sa3:fullratsimp(expand(sa2));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>altanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa3</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer satisfies the ODE but either is not the most general one or it does not use the prescribed constants "A" and "B".</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer does not satisfy the ODE</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1828121236</deployedseed>
    <deployedseed>180640144</deployedseed>
    <deployedseed>1752024964</deployedseed>
    <deployedseed>1122659436</deployedseed>
    <deployedseed>917849661</deployedseed>
    <deployedseed>1544141057</deployedseed>
    <deployedseed>1672474179</deployedseed>
    <deployedseed>223910460</deployedseed>
    <deployedseed>539677662</deployedseed>
    <deployedseed>1495050678</deployedseed>
    <deployedseed>708126127</deployedseed>
    <deployedseed>1949957589</deployedseed>
    <deployedseed>1642950791</deployedseed>
    <deployedseed>1207529474</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>altanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer - c</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer3</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer4</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>8</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer5</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>9</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer6</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>10</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer7</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
