<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-1.g</text>
    </name>
    <questiontext format="html">
      <text>Differentiate \[{@p@}\] with respect to {@v@}.
 
 [[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>
 
 \[\begin{array}{rcl}\displaystyle
   \frac{d}{d{@v@}}\left[ {@p@}\right]  &\displaystyle =&\displaystyle  \frac{d}{d{@v@}}\left[{@pp1@}\right] +\frac{d}{d{@v@}}\left[{@pp2@}\right]+\frac{d}{d{@v@}}\left[{@pp3@}\right] \\
    &\displaystyle =&\displaystyle  {@rdm1@}\cdot \frac{d}{d{@v@}}[{@ppp1@}]+{@rdm2@}\cdot \frac{d}{d{@v@}}\left[{@ppp2@}\right] +{@rdm3@}\cdot \frac{d}{d{@v@}}\left[{@ppp3@}\right] \\
    &\displaystyle =&\displaystyle  {@rdm1@}\cdot [{@diff(ppp1,v)@}]+{@rdm2@}\cdot [{@diff(ppp2,v)@}] +{@rdm3@}\cdot \left[{@diff(ppp3,v)@}\right] \\
    &\displaystyle =&\displaystyle  {@diff(p,v)@}
  \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
 rdm1:rand(4)+2;
 rdm2:-(2+rand(5));
 rdm3:-2-rand(2);
 rdm4:rdm1+1;
 rdm5:5+rand(4);
 rdm6:3+rand(4);
 ppp1:v^rdm4;
 ppp2:tan(rdm5*v);
 ppp3:exp(rdm6*v);
 pp1:rdm1*ppp1;
 pp2:rdm2*ppp2;
 pp3:rdm3*ppp3;
 p:pp1+pp2+pp3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\frac{d}{d{@v@}}({@p@}) = {@diff(p,v)@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>855122509</deployedseed>
    <deployedseed>556245265</deployedseed>
    <deployedseed>736957274</deployedseed>
    <deployedseed>2113092636</deployedseed>
    <deployedseed>1155609450</deployedseed>
    <deployedseed>1721903221</deployedseed>
    <deployedseed>1085137580</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(diff(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
