<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-4.e</text>
    </name>
    <questiontext format="html">
      <text>Expand the following expresion and differentiate it with respect to {@v@}:\[ \frac{@pp1@}{@pp2@}. \]
 
 [[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to solve this problem we will have to use the linearity rules for differentition. <hint>calc_diff_linearity_rule</hint>
 
 Since
 \[\begin{array}{rcl}\displaystyle
 \frac{@pp1@}{@pp2@} &\displaystyle =&\displaystyle  \frac{1}{@pp2@}-\frac{@ppp2@}{@pp2@} \\
    &\displaystyle =&\displaystyle  {@q@}
  \end{array}\]
 
 \[\begin{array}{rcl}\displaystyle
 \frac{d}{d{@v@}}\left[\frac{@pp1@}{@pp2@}\right]  &\displaystyle =&\displaystyle  \frac{d}{d{@v@}}\left( {@q@}\right)  \\
    &\displaystyle =&\displaystyle  \frac{d}{d{@v@}}\left({@1/pp2@}\right) - \frac{d}{d{@v@}}\left( {@ppp2/pp2@}\right) \\
    &\displaystyle =&\displaystyle  {@diff(q,v)@}
  \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
 rdm1:-2-rand(6);
 rdm2:-2-rand(5);
 ppp1:1;
 ppp2:exp(rdm1*v);
 pp1:ppp1-ppp2;
 pp2:exp(rdm2*v);
 p:pp1/pp2;
 q:expand(p);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\frac{d}{d{@v@}}({@p@}) = {@diff(q,v)@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(q,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(q,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1327355069</deployedseed>
    <deployedseed>480892886</deployedseed>
    <deployedseed>1345090948</deployedseed>
    <deployedseed>1692989655</deployedseed>
    <deployedseed>1247343301</deployedseed>
    <deployedseed>1717482419</deployedseed>
    <deployedseed>1117657294</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(diff(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
