<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Diff-02-linearity-5.b</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use the laws of logarithms to find the following derivative: \[ \frac{d}{d{@v@}} \left[{@p@}\right]. \] [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Using the laws of logarithms alg_logarithms we get \[\begin{array}{rcl}\displaystyle {@p@} &amp;\displaystyle =&amp;\displaystyle \ln(1) - \ln({@pp@}) \\ &amp;\displaystyle =&amp;\displaystyle -\ln({@ppp@}) - \ln({@rdm2@}) \\ &amp;\displaystyle =&amp;\displaystyle {@-rdm1@}{@pq@} - \ln({@rdm2@}) \end{array}\] So, by linearity calc_diff_linearity_rule, \[\begin{array}{rcl}\displaystyle \frac{d}{d{@v@}}\left[{@p@}\right] &amp;\displaystyle =&amp;\displaystyle \frac{d}{d{@v@}}\left[{@-rdm1@}{@pq@} - \ln({@rdm2@})\right] \\ &amp;\displaystyle =&amp;\displaystyle \frac{d}{d{@v@}}\left[{@-rdm1@}{@pq@}\right] + \frac{d}{d{@v@}}\left[-\ln({@rdm2@})\right] \\ &amp;\displaystyle =&amp;\displaystyle {@-rdm1@}\frac{d}{d{@v@}}\left[{@pq@}\right] + 0 \\ &amp;\displaystyle =&amp;\displaystyle {@-rdm1@}{@diff(pq,v)@} \\ &amp;\displaystyle =&amp;\displaystyle {@-rdm1*diff(pq,v)@} \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
 rdm1:1/(2+rand(2));
 rdm2:(rand([2,3,5,6,7]))**rdm1;
 ppp:v^(rdm1);
 pp:rdm2*ppp;
 p:log(1/pp);
 pq:log(v);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\frac{d}{d{@v@}}({@p@}) = {@diff(p,v)@} \]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(p,v)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>diff(p,v)</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1172383578</deployedseed>
    <deployedseed>1242803943</deployedseed>
    <deployedseed>975122571</deployedseed>
    <deployedseed>387124762</deployedseed>
    <deployedseed>1442914739</deployedseed>
    <deployedseed>2002730459</deployedseed>
    <deployedseed>1035953248</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(diff(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
