<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Diff_Product_3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \[ \frac{dy}{dx}\] when \[ y = \sqrt{{@n1@}{@v@}}\ln({@n2@}{@v@})\] [[input:ans1]] [[validation:ans1]]</p>
<p> </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In order to solve this problem we will have to use the product rule \(\frac{d}{dx}(fg) = \frac{df}{dx}\cdot g+f\cdot\frac{dg}{dx}\) as follows:</p>
<p>Let \(f(x) =\displaystyle \sqrt{{@n1@}{@v@}}\mbox{ and }g(x) = \ln({@n2@}{@v@})\) <br /> \[\begin{array}{rcl}\displaystyle \frac{d}{d{@v@}}\left(\sqrt{{@n1@}{@v@}}\ln({@n2@}{@v@})\right) &amp; =\displaystyle &amp; \left(\frac{d}{d{@v@}}(\sqrt{{@n1@}{@v@}})\right)\cdot \ln({@n2@}{@v@}) + (\sqrt{{@n1@}{@v@}})\cdot \left(\frac{d}{d{@v@}}\left(\ln({@n2@}{@v@})\right)\right) \\ &amp; =&amp; \left(\frac{\sqrt{@n1@}}{2\sqrt {@v@}}\right)\cdot \ln({@n2@}{@v@}) + (\sqrt{{@n1@}{@v@}})\cdot \frac{1}{@v@} \\ &amp; =&amp; {@correctanswer@} \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
 n1:rand(8) + 2;
 n2:rand(8) + 2;
 pp1:sqrt(n1*v);
 pp2:ln(n2*v);
 p:sqrt(n1*v)*ln(n2*v);
 correctanswer:diff(p,v);
 wronganswer1:diff(pp1,v)*diff(pp2,v);
 wronganswer2:(diff(pp1,v)*pp1)+(pp2*diff(pp2,v));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\frac{d}{d{@v@}}(\sqrt{{@n1@}{@v@}}\ln({@n2@}{@v@})) = \displaystyle {@correctanswer@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>To differentiate a product use the rule \[\frac{d}{dx}(fg) = \frac{df}{dx}\cdot g+f\cdot\frac{dg}{dx}.\] You cannot get the derivative of a product by multiplying the derivatives.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have applied the product rule incorrectly. It is \[\frac{df}{dx}\cdot g+f\cdot\frac{dg}{dx},\] not \[f\cdot\frac{df}{dx}+g\cdot\frac{dg}{dx}.\]</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1502280057</deployedseed>
    <deployedseed>1364662697</deployedseed>
    <deployedseed>607957207</deployedseed>
    <deployedseed>1907285151</deployedseed>
    <deployedseed>1649058069</deployedseed>
    <deployedseed>2063137622</deployedseed>
    <deployedseed>774455681</deployedseed>
    <deployedseed>1721168155</deployedseed>
    <deployedseed>1484763029</deployedseed>
    <deployedseed>1194267091</deployedseed>
    <deployedseed>334810998</deployedseed>
    <deployedseed>1981492109</deployedseed>
    <deployedseed>1056305095</deployedseed>
    <deployedseed>1903757384</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
