<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>6.3.7 (Spot the parametric curve 2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div class="center">            <p class="noindent">            </p>Which graph shows the graph of the function given parametrically by \(x = {@xvar@}\) and \(y = {@yvar@}\) for \(-2 \leq t \leq 2\)?<br><p class="noindent">[[input:ans1]] [[validation:ans1]]&nbsp;<br></p></div><table>
  <tbody><tr>
    <td style="text-align: center;"><span><b>Graph A</b></span></td>
    <td style="text-align: center;"><span><b>Graph B</b></span></td>
  </tr>
  <tr>
    <td style="text-align: center;">{@plot(plist[1],[box, false],[yx_ratio, 1],[axes, solid],[alt,"Graph A"])@}</td>
    <td style="text-align: center;">{@plot(plist[2],[box, false],[yx_ratio, 1],[axes, solid],[alt,"Graph B"])@}</td>
  </tr><tr>
    <td style="text-align: center;">&nbsp;</td>
    <td style="text-align: center;">&nbsp;</td>
  </tr>
  <tr>
    <td style="text-align: center;"><b>Graph C</b></td>
    <td style="text-align: center;"><b>Graph D</b></td>
  </tr>
  <tr>
    <td style="text-align: center;">{@plot(plist[3],[box, false],[yx_ratio, 1],[axes, solid],[alt,"Graph C"])@}</td>
    <td style="text-align: center;">{@plot(plist[4],[box, false],[yx_ratio, 1],[axes, solid],[alt,"Graph D"])@}</td>
  </tr>
</tbody></table><p><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p class="noindent"><br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Make sure m<>1, or adjust the list of options. */
m:2;
/* Make sure a<>b. */
amag:2+rand(2);
pm1:rand([-1,1]);
pm2:rand([-1,1]);
bmag:1+rand(2);
a:pm1*amag;
b:pm2*bmag;
xvar:t^2+t
yvar:t^2-t
q:[parametric,xvar,yvar, [t,-2,2]];
flistpre:[[parametric, cos(t), sin(t), [t,0,2*%pi]],[parametric, t^2, t^3, [t,-2,2]],[parametric, t^3, t, [t,-2,2]],[parametric, t^2, t,[t,-2,2]],[parametric, t^2, 1-t, [t,-2,2]],[parametric, t, t^2, [t,-2,2]]];
flist:rand_selection(flistpre,3);
flist:append([[q,true]],maplist(lambda([ex],[ex,false]),flist));
flist:random_permutation(flist);
/* Define axes */
YL:[y,-5,5]
XL:[x,-5,5]
/* Create a list of just the plots, to plot. */
plist:map(first,flist);
/* Identify both the correct answer, and the position of the correct answer. */
tac:first(mcq_correct(flist));
n:first(sublist_indices(flist, lambda([ex],second(ex))));
corletter:[A,B,C,D][n];
ta1:maplist(lambda([letter],[letter,is(letter=corletter),sconcat("Graph ",letter)]),[A,B,C,D]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]<br></p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@flist@} f is element {#n#}, \((x,y) = ({@xvar@}, {@yvar@})\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>corletter</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>242494260</deployedseed>
    <deployedseed>184845467</deployedseed>
    <deployedseed>1319450551</deployedseed>
    <deployedseed>1574809556</deployedseed>
    <deployedseed>181031542</deployedseed>
    <deployedseed>505593724</deployedseed>
    <deployedseed>733720214</deployedseed>
    <deployedseed>918484921</deployedseed>
    <deployedseed>751401084</deployedseed>
    <deployedseed>1032166348</deployedseed>
    <deployedseed>397001937</deployedseed>
    <deployedseed>1474811369</deployedseed>
    <deployedseed>2103269880</deployedseed>
    <deployedseed>1467069906</deployedseed>
    <deployedseed>577842372</deployedseed>
    <deployedseed>506636225</deployedseed>
    <deployedseed>1931641897</deployedseed>
    <deployedseed>149926436</deployedseed>
    <deployedseed>1256183921</deployedseed>
    <deployedseed>533607612</deployedseed>
    <deployedseed>127564621</deployedseed>
    <deployedseed>1597471399</deployedseed>
    <deployedseed>1031396936</deployedseed>
    <deployedseed>1608579895</deployedseed>
    <deployedseed>1970972000</deployedseed>
    <deployedseed>1986783828</deployedseed>
    <deployedseed>541002085</deployedseed>
    <deployedseed>600939451</deployedseed>
    <deployedseed>855948777</deployedseed>
    <deployedseed>719422011</deployedseed>
    <deployedseed>1274773316</deployedseed>
    <deployedseed>1727419943</deployedseed>
    <deployedseed>2013361335</deployedseed>
    <deployedseed>382982534</deployedseed>
    <deployedseed>2023980294</deployedseed>
    <deployedseed>431618979</deployedseed>
    <deployedseed>735486928</deployedseed>
    <deployedseed>258388205</deployedseed>
    <deployedseed>640408468</deployedseed>
    <deployedseed>10147505</deployedseed>
    <deployedseed>135457050</deployedseed>
    <deployedseed>11746610</deployedseed>
    <deployedseed>1981453532</deployedseed>
    <deployedseed>1513104969</deployedseed>
    <deployedseed>1356204728</deployedseed>
    <deployedseed>91413412</deployedseed>
    <deployedseed>1190183674</deployedseed>
    <deployedseed>1859514210</deployedseed>
    <deployedseed>988676197</deployedseed>
    <deployedseed>239537578</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>corletter</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>delete(corletter,[A,B,C,D])[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
