<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>6.3.19 (Parametric to parabola)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the values of \(A\) and \(B\) so that the curve given in parametric form by \(x = {@xvar@}\) and \(y = {@yvar@}\) can be written as a function \(x = (y+A)^2 +B\).</p><p>\(A = \) [[input:ans1]] [[validation:ans1]]</p><p>\(B = \) [[input:ans2]] [[validation:ans2]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We have that \(t = {@tsol@}\), so substituting this into the expression for \(x\) and completing the square we have that\begin{align*}<br>x = {@xvar@}<br>&amp;={@xsub@} \\<br>&amp; = y^2 - {@2*(b-a)@}y + {@b*(b-2*a)@}\\<br>&amp; = (y - {@(b-a)@})^2 - {@(b-a)^2@} + {@b*(b-2*a)@}\\<br>&amp; = (y - {@(b-a)@})^2 + {@-((b-a)^2)+b*(b-2*a)@}.\end{align*}So \(A= {@ta1@}\) and \(B = {@ta2@}.\)<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a:rand_range(-10,10,1);
b:rand_with_prohib(-10,10,[a,0]);
xvar:t*(t+2*a);
yvar:t+b;
tsol:y-b;
xsub:(y-b)*(y-b+2*a);
ta1:-(b-a);
ta2:b*(b-2*a) - (b-a)^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(x= {@xvar@}\) and \(y = {@yvar@}\) gives \(A= {@ta1@}\) and \(B = {@ta2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your value for \(A\) is incorrect<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your value for \(B\) is incorrect.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1671732852</deployedseed>
    <deployedseed>706136884</deployedseed>
    <deployedseed>1823082710</deployedseed>
    <deployedseed>64724470</deployedseed>
    <deployedseed>405707520</deployedseed>
    <deployedseed>1548357030</deployedseed>
    <deployedseed>602601829</deployedseed>
    <deployedseed>78043584</deployedseed>
    <deployedseed>1562035197</deployedseed>
    <deployedseed>190031447</deployedseed>
    <deployedseed>651934254</deployedseed>
    <deployedseed>1069107453</deployedseed>
    <deployedseed>611602521</deployedseed>
    <deployedseed>533364491</deployedseed>
    <deployedseed>511180744</deployedseed>
    <deployedseed>1664202788</deployedseed>
    <deployedseed>1716207125</deployedseed>
    <deployedseed>354978247</deployedseed>
    <deployedseed>1591130202</deployedseed>
    <deployedseed>1864808902</deployedseed>
    <deployedseed>927439205</deployedseed>
    <deployedseed>172877923</deployedseed>
    <deployedseed>693873430</deployedseed>
    <deployedseed>1511643166</deployedseed>
    <deployedseed>2093714977</deployedseed>
    <deployedseed>59388524</deployedseed>
    <deployedseed>908333675</deployedseed>
    <deployedseed>295039659</deployedseed>
    <deployedseed>554495839</deployedseed>
    <deployedseed>2077044871</deployedseed>
    <deployedseed>172049173</deployedseed>
    <deployedseed>881794235</deployedseed>
    <deployedseed>2104016256</deployedseed>
    <deployedseed>1618509532</deployedseed>
    <deployedseed>113952247</deployedseed>
    <deployedseed>1447897957</deployedseed>
    <deployedseed>929615842</deployedseed>
    <deployedseed>1542665212</deployedseed>
    <deployedseed>577697157</deployedseed>
    <deployedseed>1664884007</deployedseed>
    <deployedseed>298726974</deployedseed>
    <deployedseed>164910723</deployedseed>
    <deployedseed>259518158</deployedseed>
    <deployedseed>1343233276</deployedseed>
    <deployedseed>1333579220</deployedseed>
    <deployedseed>1367650209</deployedseed>
    <deployedseed>1841019445</deployedseed>
    <deployedseed>1741702204</deployedseed>
    <deployedseed>1291787061</deployedseed>
    <deployedseed>1647796255</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
