<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>6.4.7 (Find speed and acceleration)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A {@mm@}kg runner is moving around a circular track with a radius of {@rr@}m. The runner makes it a quarter of the way around the circle in {@tt@} seconds. Determine the speed of the runner. Give your answer to the nearest whole number.</p>

<p>Speed = [[input:ans1]] [[validation:ans1]]\(ms^{-1}\)</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We seek to use the equations we have derived above. Here we have \(m = {@mm@}, R = {@rr@}, T = {@tt@}\). 
To find the speed of the runner we can use the speed = distance divided by time equation to give:
\[
d = \frac{\frac{1}{4} \cdot 2 \cdot \pi \cdot {@rr@}}{{@tt@}} = {@round(vv)@}
\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>mm: 90+rand(20);
rr: 10+rand(10);
tt: 2+rand(15);
vv: (1/4*2*pi*rr)/tt;

ta1:round(vv);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>radius = {@rr@}, speed ={@ta1@}.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1045201141</deployedseed>
    <deployedseed>972359265</deployedseed>
    <deployedseed>650046652</deployedseed>
    <deployedseed>1409681700</deployedseed>
    <deployedseed>1365754465</deployedseed>
    <deployedseed>200494598</deployedseed>
    <deployedseed>140172281</deployedseed>
    <deployedseed>1771400531</deployedseed>
    <deployedseed>1996281669</deployedseed>
    <deployedseed>135388288</deployedseed>
    <deployedseed>1192106972</deployedseed>
    <deployedseed>1560538668</deployedseed>
    <deployedseed>1191755409</deployedseed>
    <deployedseed>1192704715</deployedseed>
    <deployedseed>2125724810</deployedseed>
    <deployedseed>540531594</deployedseed>
    <deployedseed>81082346</deployedseed>
    <deployedseed>213850538</deployedseed>
    <deployedseed>756848306</deployedseed>
    <deployedseed>1629489192</deployedseed>
    <deployedseed>666287630</deployedseed>
    <deployedseed>1517588610</deployedseed>
    <deployedseed>809879502</deployedseed>
    <deployedseed>740895209</deployedseed>
    <deployedseed>1741798685</deployedseed>
    <deployedseed>483318715</deployedseed>
    <deployedseed>1482418688</deployedseed>
    <deployedseed>155904742</deployedseed>
    <deployedseed>532444837</deployedseed>
    <deployedseed>1952970770</deployedseed>
    <deployedseed>1929850743</deployedseed>
    <deployedseed>260200161</deployedseed>
    <deployedseed>2001008324</deployedseed>
    <deployedseed>833585634</deployedseed>
    <deployedseed>1864440776</deployedseed>
    <deployedseed>1568983817</deployedseed>
    <deployedseed>663299295</deployedseed>
    <deployedseed>1865710456</deployedseed>
    <deployedseed>1282250936</deployedseed>
    <deployedseed>578534325</deployedseed>
    <deployedseed>885060326</deployedseed>
    <deployedseed>1667021307</deployedseed>
    <deployedseed>2007638344</deployedseed>
    <deployedseed>1845234664</deployedseed>
    <deployedseed>1205928935</deployedseed>
    <deployedseed>716261172</deployedseed>
    <deployedseed>411636053</deployedseed>
    <deployedseed>1626085199</deployedseed>
    <deployedseed>131063797</deployedseed>
    <deployedseed>1052172584</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
