<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Differential_Equation_11</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>For what values of k is the following differential equation exact? \[ ( {@Nk@} ) \frac{dy}{dx} + {@Mk@} = 0.\] <br /> k = [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>As the differential equation\[ ( {@Nk@} ) \frac{d y}{dx} + {@Mk@} = 0\] is of the form \[A(x,y)\frac{dy}{dx} + B(x,y) = 0\] the condition for the differential equation to be exact is \[\frac{A(x,y)}{x} = \frac{B(x,y)}{y}\] Therefore \[{@diffNk@}={@diffMk@}\] \[k={@correctanswer@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>nk:rand(6)+2;
 PO:x^2 + expand(y*((rand(9)+1)*x+(rand(9)+1)*y+ rand(9)+1));
 P1:y + expand((rand(9)+1)*x + rand(9)+1);
 Pk:PO+k*P1;
 Mk:diff(Pk,x);
 Nk:diff(Pk,y)+k*x-nk*x;
 diffMk:diff(Mk,y,1);
 diffNk:diff(Nk,x,1);
 correctanswer:nk;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>DE \( ( {@Nk@} ) \frac{d y}{dx} + {@Mk@} = 0\) exact for \( k = {@nk@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>582256425</deployedseed>
    <deployedseed>561649084</deployedseed>
    <deployedseed>504185631</deployedseed>
    <deployedseed>89920800</deployedseed>
    <deployedseed>1233752633</deployedseed>
    <deployedseed>697746730</deployedseed>
    <deployedseed>868691957</deployedseed>
    <deployedseed>152355841</deployedseed>
    <deployedseed>229799249</deployedseed>
    <deployedseed>490112864</deployedseed>
    <deployedseed>977739671</deployedseed>
    <deployedseed>1859464756</deployedseed>
    <deployedseed>230030739</deployedseed>
    <deployedseed>1659390898</deployedseed>
    <deployedseed>183448632</deployedseed>
    <deployedseed>2079185852</deployedseed>
    <deployedseed>1835761396</deployedseed>
    <deployedseed>137955905</deployedseed>
    <deployedseed>1735333036</deployedseed>
    <deployedseed>1751937062</deployedseed>
    <deployedseed>889273044</deployedseed>
    <deployedseed>1892747123</deployedseed>
    <deployedseed>718646385</deployedseed>
    <deployedseed>1292409585</deployedseed>
    <deployedseed>853654179</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer-1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
