<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>14.4.1.3 Exercise 3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Calculate the length of the curve \(y^2 = {@y^2@}\) between \(x = {@a@}\) and \(x = {@b@}\), in the first quadrant. Give your answer in <b>exact</b> form.</p>
<p>The length is [[input:ans1]]</p>
<p>[[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p> We are given \( y^2 = {@y^2@} \) then \( y = {@y@}\)on the first quadrant.</p>
<p> The derivative \( \dfrac{\mathrm{d}y}{\mathrm{d}x} = {@dy@}.\)</p>

<p>Using the formula \[\text{length of curve}=\int_a^b \sqrt{1+\left(\frac{\mathrm{d}y}{\mathrm{d}x}\right)^2} \mathrm{d}x,\] the length of the line is calculated to be:
\[ \ell = \int_{{@a@}}^{{@b@}} \sqrt{1+{@dy^2@}} \, \mathrm{d}x = \left[ {@I@} \right]_{{@a@}}^{{@b@}} = {@ta1@}.\]</p>
<p>For this integral you can use the substitution \( u= 1+{@dy^2@}.\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021071900</text>
    </stackversion>
    <questionvariables>
      <text>m: rand([2,4,6]);     /* even number so it simplifies with the power 3/2 after finding the derivative*/
y: m*x^(3/2);

a: 0          /*lower limit, pick 0 so it makes calculation easier*/
ps: rand([4,9,16,25,49]);   /* I want the quantity under the sqrt to give a perfect square so the result is a nice number, I pick this first and then find the uper limit I will use */
b: (ps-1)/(9*(m/2)^2);   /* upper limit*/

dy: diff(y, x);
I:  integrate(sqrt(1+(dy)^2), x);
ta1: integrate(sqrt(1+(dy)^2), x, a, b);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(y^2 = {@m@}x^{3}, \,  ps={@ps@}, \, x = [{@a@}, {@b@}], \, l={@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>units</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>166112053</deployedseed>
    <deployedseed>1861953019</deployedseed>
    <deployedseed>1372023521</deployedseed>
    <deployedseed>1043266875</deployedseed>
    <deployedseed>1916827629</deployedseed>
    <deployedseed>1417976487</deployedseed>
    <deployedseed>1644012744</deployedseed>
    <deployedseed>523697406</deployedseed>
    <deployedseed>1907932369</deployedseed>
    <deployedseed>1827725124</deployedseed>
    <deployedseed>1845410382</deployedseed>
    <deployedseed>2122613686</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
