<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>2.1.23 (Fully factorise a quartic with irred quad)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Fully factorise the polynomial \(p(x)={@p@}\), given that {@x=r1@} is a root.</p><p>\(p(x)=\) [[input:ans1]][[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Since {@x=r1@} is a root, the factor theorem gives that \({@x-r1@}\) is a factor. Taking out this factor gives:</p><p>\[ p(x)=({@x-r1@})({@cubicTerm@}). \]</p><p>To factorise the cubic term, we notice that {@x=inspectedRoot@} is a root and so</p><p>\[ p(x)={@ta1@}. \]</p><p>The quadratic term has discriminant \(b^2-4ac=({@b@})^2-4\times{@a@}\times{@c@}={@ev(b^2-4*a*c,simp)@}&lt;0\) so cannot be factorised further.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>r1:rand_with_prohib(-4,4,[0,1,-1]);
rdenom1:1;
r2:rand([1,-1]);
a:1;
b:rand_with_prohib(-3,3,[0]);
c:rand_range(ev(floor(b^2/4)+1,simp),4);
p:ev(expand((rdenom1*x-r1)*(x-r2)*(a*x^2+b*x+c)),simp);
ta1: ev((rdenom1*x-r1)*(x-r2)*(a*x^2+b*x+c),simp);
inspectedRoot: r2;
cubicTerm: ev(expand((x-r2)*(a*x^2+b*x+c)),simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]<br></p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@p=ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>588569017</deployedseed>
    <deployedseed>1610010071</deployedseed>
    <deployedseed>1609622929</deployedseed>
    <deployedseed>315699167</deployedseed>
    <deployedseed>1841758021</deployedseed>
    <deployedseed>876391109</deployedseed>
    <deployedseed>1459314214</deployedseed>
    <deployedseed>2044949618</deployedseed>
    <deployedseed>101926765</deployedseed>
    <deployedseed>1988666339</deployedseed>
    <deployedseed>63352807</deployedseed>
    <deployedseed>112091690</deployedseed>
    <deployedseed>2100915109</deployedseed>
    <deployedseed>327411035</deployedseed>
    <deployedseed>844728855</deployedseed>
    <deployedseed>1593776702</deployedseed>
    <deployedseed>160268542</deployedseed>
    <deployedseed>339450850</deployedseed>
    <deployedseed>1542193988</deployedseed>
    <deployedseed>831334225</deployedseed>
    <deployedseed>1254805115</deployedseed>
    <deployedseed>1534476873</deployedseed>
    <deployedseed>1655706235</deployedseed>
    <deployedseed>1147614520</deployedseed>
    <deployedseed>178804851</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>(x-r1)*cubicTerm</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
