<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>2.3.10 (Finding independent variable)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given \(f(x) = {@func2@}\) and that \(f(c) = -{@aa*bb@}\) then what is the value of \(c\)?</p>
<p>\(c=\)[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>To begin, we substitute \(c\) into the function. We know that \(f(c)\) evaluates to \(-{@aa*bb@}\) therefore we have
\[
c^2 -{@aa+bb@}c = -{@aa*bb@} \implies c^2 -{@aa+bb@}c + {@aa*bb@}  = 0.
\]
Notice that we now have a quadratic in standard form. Therefore we can solve the quadratic using whichever tecnique we would like from week 1. In this case the quadratic factorises as 
\[
c^2 -{@aa+bb@}c + {@aa*bb@} = (c-{@aa@})(c-{@bb@}).
\]
Hence either \(c = {@aa@}\) or \(c = {@bb@}\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1+rand(10);
bb: 1+ rand(10);

func1: (x-aa)*(x-bb);
func2: expand((x-aa)*(x-bb)) - aa*bb;
ta1: aa;
ta2: bb;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>[{@ta1@},{@ta2@}]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>66574346</deployedseed>
    <deployedseed>1002349474</deployedseed>
    <deployedseed>241573610</deployedseed>
    <deployedseed>1332138872</deployedseed>
    <deployedseed>192345497</deployedseed>
    <deployedseed>1105530756</deployedseed>
    <deployedseed>1572316384</deployedseed>
    <deployedseed>2111741426</deployedseed>
    <deployedseed>1259838874</deployedseed>
    <deployedseed>2034410476</deployedseed>
    <deployedseed>1289924941</deployedseed>
    <deployedseed>687918329</deployedseed>
    <deployedseed>1394868554</deployedseed>
    <deployedseed>948807924</deployedseed>
    <deployedseed>866404307</deployedseed>
    <deployedseed>1378310242</deployedseed>
    <deployedseed>901866052</deployedseed>
    <deployedseed>2083502147</deployedseed>
    <deployedseed>1161984942</deployedseed>
    <deployedseed>1407582308</deployedseed>
    <deployedseed>220620408</deployedseed>
    <deployedseed>1534621482</deployedseed>
    <deployedseed>1337155385</deployedseed>
    <deployedseed>655973504</deployedseed>
    <deployedseed>498154483</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
