<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>2.3.21 (Spot the scaled graph)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div class="center">            <p class="noindent">            </p><p class="noindent">The graph of a function \(y = f(x)\) is given below:<br></p><p class="noindent">{@plot(p,XL,YL,[box, false],[yx_ratio, 1],[axes, solid],[xtics, -10, 2,10],[ytics, -10, 2, 10])@}</p><p class="noindent">Which graph shows the graph of the function \(f({@a*x@})\)?<br></p><p class="noindent">[[input:ans1]] [[validation:ans1]]&nbsp;<br></p></div><table>
  <tbody><tr>
    <td style="text-align: center;"><span><b>Graph A</b></span></td>
    <td style="text-align: center;"><span><b>Graph B</b></span></td>
  </tr>
  <tr>
    <td style="text-align: center;">{@plot(plist[1],XL,YL,[box, false],[yx_ratio, 1],[axes, solid],[xtics, -10, 2,10],[ytics, -10, 2, 10],[alt,"Graph A"])@}</td>
    <td style="text-align: center;">{@plot(plist[2],XL,YL,[box, false],[yx_ratio, 1],[axes, solid],[xtics,-10, 2,10],[ytics, -10, 2, 10],[alt,"Graph B"])@}</td>
  </tr><tr>
    <td style="text-align: center;">&nbsp;</td>
    <td style="text-align: center;">&nbsp;</td>
  </tr>
  <tr>
    <td style="text-align: center;"><b>Graph C</b></td>
    <td style="text-align: center;"><b>Graph D</b></td>
  </tr>
  <tr>
    <td style="text-align: center;">{@plot(plist[3],XL,YL,[box, false],[yx_ratio, 1],[axes, solid],[xtics, -10, 2, 10],[ytics, -10, 2, 10],[alt,"Graph C"])@}</td>
    <td style="text-align: center;">{@plot(plist[4],XL,YL,[box, false],[yx_ratio, 1],[axes, solid],[xtics, -10, 2, 10],[ytics, -10, 2, 10],[alt,"Graph D"])@}</td>
  </tr>
</tbody></table><p><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p class="noindent">The graph of \(f(ax)\) is the graph of \(f(x)\) compressed horizontally if \(a&gt;1\) and, if \(a&lt;-1\), reflected across the \(y\)-axis and compressed horizontally.<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Make sure m<>1, or adjust the list of options. */
m:2;
/* Make sure a<>b. */
amag:2+rand(2);
pm1:rand([-1,1]);
pm2:rand([-1,1]);
bmag:1+rand(2);
a:pm1*amag;
b:pm2*bmag;
p:-(x+1)^2 +1;
q:-(a*x+1)^2+1;
flistpre:[-(x/a+1)^2 +1,-a*(x+1)^2 +1,-((x+1)^2)/a +1,-(x+1)^2 +a,-(x+1+a)^2 +1,-(x+1-a)^2 +1];
flist:rand_selection(flistpre,3);
flist:append([[q,true]],maplist(lambda([ex],[ex,false]),flist));
flist:random_permutation(flist);
/* Define axes */
YL:[y,-5,5]
XL:[x,-5,5]
/* Create a list of just the plots, to plot. */
plist:map(first,flist);
/* Identify both the correct answer, and the position of the correct answer. */
tac:first(mcq_correct(flist));
n:first(sublist_indices(flist, lambda([ex],second(ex))));
corletter:[A,B,C,D][n];
ta1:maplist(lambda([letter],[letter,is(letter=corletter),sconcat("Graph ",letter)]),[A,B,C,D]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]<br></p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@flist@} f is element {#n#}, \(y= {@p@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>corletter</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>782206436</deployedseed>
    <deployedseed>2055135021</deployedseed>
    <deployedseed>1923166716</deployedseed>
    <deployedseed>1285628160</deployedseed>
    <deployedseed>121235231</deployedseed>
    <deployedseed>1575511747</deployedseed>
    <deployedseed>1929240938</deployedseed>
    <deployedseed>886992188</deployedseed>
    <deployedseed>1823090588</deployedseed>
    <deployedseed>458498210</deployedseed>
    <deployedseed>412962919</deployedseed>
    <deployedseed>1507795030</deployedseed>
    <deployedseed>128011617</deployedseed>
    <deployedseed>1934834059</deployedseed>
    <deployedseed>330929914</deployedseed>
    <deployedseed>953668978</deployedseed>
    <deployedseed>2075495769</deployedseed>
    <deployedseed>992543717</deployedseed>
    <deployedseed>1024005312</deployedseed>
    <deployedseed>2066244760</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>corletter</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>delete(corletter,[A,B,C,D])[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
