<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>2.3.9 (Evaluate a function algebra 1)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the function \(f(x) = {@func@}\) at \(2a\).</p>
<p>\(f(2a)=\)[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We are required to find \(f(2a)\). This means that we substitute the value of \(x=2a\) into the equation \({@func@}\), in the same way as we would if we wanted to evaluate the function at a given number, Therefore we have
\[
f(2a) = {@aa@}(2a)^2 + {@bb@}(2a) - {@cc@} = {@ta1@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020070100</text>
    </stackversion>
    <questionvariables>
      <text>aa: 1+rand(10);
bb: 1+rand(10);
cc: 1+rand(10);

func: aa*x^2+bb*x-cc;

ta1: ev(func, x=2*a);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>f(2a)={@ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1934474118</deployedseed>
    <deployedseed>2003539669</deployedseed>
    <deployedseed>1243672913</deployedseed>
    <deployedseed>238283298</deployedseed>
    <deployedseed>1560634988</deployedseed>
    <deployedseed>599684439</deployedseed>
    <deployedseed>476772591</deployedseed>
    <deployedseed>1840738942</deployedseed>
    <deployedseed>1479658517</deployedseed>
    <deployedseed>920113481</deployedseed>
    <deployedseed>1551079470</deployedseed>
    <deployedseed>1110166039</deployedseed>
    <deployedseed>423865576</deployedseed>
    <deployedseed>350436458</deployedseed>
    <deployedseed>2087750728</deployedseed>
    <deployedseed>322114271</deployedseed>
    <deployedseed>2057621252</deployedseed>
    <deployedseed>1464533874</deployedseed>
    <deployedseed>1861869301</deployedseed>
    <deployedseed>940799982</deployedseed>
    <deployedseed>1807773501</deployedseed>
    <deployedseed>490800497</deployedseed>
    <deployedseed>1199970517</deployedseed>
    <deployedseed>248847034</deployedseed>
    <deployedseed>1594903031</deployedseed>
    <deployedseed>593253394</deployedseed>
    <deployedseed>1518813857</deployedseed>
    <deployedseed>1494377800</deployedseed>
    <deployedseed>1380403317</deployedseed>
    <deployedseed>1540449652</deployedseed>
    <deployedseed>1539839604</deployedseed>
    <deployedseed>1558866672</deployedseed>
    <deployedseed>1825027951</deployedseed>
    <deployedseed>1769711967</deployedseed>
    <deployedseed>1697918692</deployedseed>
    <deployedseed>980723931</deployedseed>
    <deployedseed>734799640</deployedseed>
    <deployedseed>68665185</deployedseed>
    <deployedseed>1303090547</deployedseed>
    <deployedseed>1813361584</deployedseed>
    <deployedseed>1976196283</deployedseed>
    <deployedseed>1919152627</deployedseed>
    <deployedseed>1479014612</deployedseed>
    <deployedseed>409614297</deployedseed>
    <deployedseed>1245707229</deployedseed>
    <deployedseed>939667748</deployedseed>
    <deployedseed>79736948</deployedseed>
    <deployedseed>1734651757</deployedseed>
    <deployedseed>2072544703</deployedseed>
    <deployedseed>447301280</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>aa*2*a^2+bb*2*x-cc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
