<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 5292  -->
  <question type="stack">
    <name>
      <text>Free-text input with plot</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The free-text input below is pre-populated with a simple ASCII plot block.</p>
<div class="free-text-container">
[[input:ans1]]
[[validation:ans1]]
[[ascii input="ans1"]]
  [[filter type="plot" /]]
  [[filter type="markdown-math" /]]
[[/ascii]]
</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[[[ascii]]
  [[filter type="plot" /]]
  [[filter type="markdown-math" /]]
{@ta1@}
[[/ascii]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026042402</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ta1:"A graph of `y=x^2-1`.

!!p
x: -4..4
y: -2..8
y=x^2-1
!!p

The vertex is at `(0,-1)`.";
sh1:"!!p
x: -4..4
y: -2..8
y=x^2-1
!!p";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>This is an example question showing a free-text input with an ASCII plot block.</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[[[commonstring key="symbolicprtcorrectfeedback"/]] [[commonstring key="defaultprtcorrectfeedback"/]]]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[[[commonstring key="symbolicprtpartiallycorrectfeedback"/]] [[commonstring key="defaultprtpartiallycorrectfeedback"/]]]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[[[commonstring key="symbolicprtincorrectfeedback"/]] [[commonstring key="defaultprtincorrectfeedback"/]]]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>freetext</type>
      <tans>ta1</tans>
      <boxsize>100</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{#sh1#}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>monospace:true</options>
    </input>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
    </qtest>
  </question>

</quiz>
