<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.0.15 (Expand brackets (AlgMap-7.7))</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Expand the following: <br> <br> \( \left( {@cf1*x+c@} \right) \left( {@cf2*x+d@} \right)\) = [[input:ans]] [[validation:ans]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In this example:<br></p>
<p>\[\begin{align} {@(cf1*x+c)*(cf2*x+d)@} &amp;= {@(cf1*cf2)*x^2@}+\left({@(c*cf2*x)@}\right)+\left({@(d*cf1*x)@}\right)+\left({@(c*d)@}\right) \\ &amp;= {@ta@}. \end{align} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a:rand(10)+1;
b:rand(10)+1;
e:rand([-1,1]);
f:rand([-1,1]);
c:a*e;
d:b*f;
cf1:(rand(3)+2)*(-1)^rand([-1,1]);
cf2:(rand(3)+2)*(-1)^rand([-1,1]);
ta:expand((cf1*x+c)*(cf2*x+d));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]<br></p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[{@(cf1*x+c)*(cf2*x+d)@} = {@ta@}.\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>expand</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>1-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-1-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is algebraically equivalent to the correct answer but you need to leave the answer as a polynomial (multiply out the brackets!)</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>1-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>549873436</deployedseed>
    <deployedseed>677646600</deployedseed>
    <deployedseed>98913168</deployedseed>
    <deployedseed>497882031</deployedseed>
    <deployedseed>1242507175</deployedseed>
    <deployedseed>183204854</deployedseed>
    <deployedseed>147859657</deployedseed>
    <deployedseed>1315062015</deployedseed>
    <deployedseed>450972361</deployedseed>
    <deployedseed>1085936573</deployedseed>
    <deployedseed>387232810</deployedseed>
    <deployedseed>1397870877</deployedseed>
    <deployedseed>57632770</deployedseed>
    <deployedseed>303242262</deployedseed>
    <deployedseed>921028055</deployedseed>
    <deployedseed>1881760095</deployedseed>
    <deployedseed>1019228209</deployedseed>
    <deployedseed>1750665316</deployedseed>
    <deployedseed>1799075258</deployedseed>
    <deployedseed>1526682249</deployedseed>
    <deployedseed>716584123</deployedseed>
    <deployedseed>1107655521</deployedseed>
    <deployedseed>1915361667</deployedseed>
    <deployedseed>894412741</deployedseed>
    <deployedseed>150375793</deployedseed>
    <deployedseed>1356564673</deployedseed>
    <deployedseed>824977961</deployedseed>
    <deployedseed>982773529</deployedseed>
    <deployedseed>826264269</deployedseed>
    <deployedseed>238996820</deployedseed>
    <deployedseed>1192554933</deployedseed>
    <deployedseed>435331883</deployedseed>
    <deployedseed>1656301454</deployedseed>
    <deployedseed>1724198274</deployedseed>
    <deployedseed>1732926435</deployedseed>
    <deployedseed>83762538</deployedseed>
    <deployedseed>1177195507</deployedseed>
    <deployedseed>107411993</deployedseed>
    <deployedseed>1962844570</deployedseed>
    <deployedseed>88474117</deployedseed>
    <deployedseed>593967828</deployedseed>
    <deployedseed>1920689591</deployedseed>
    <deployedseed>701361260</deployedseed>
    <deployedseed>1463616118</deployedseed>
    <deployedseed>560669373</deployedseed>
    <deployedseed>1696001827</deployedseed>
    <deployedseed>337352784</deployedseed>
    <deployedseed>1189666167</deployedseed>
    <deployedseed>505445569</deployedseed>
    <deployedseed>424152515</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>(cf1*x+c)*(cf2*x+d)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>1-1-T </expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
