<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.1.10 (Exp laws with e)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Suppose \({@eq1@}\) and \({@eq2@}\).</p><p>Then \({@f@} = \) [[input:ans1]]. [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>By the laws of exponentials we have that\[\begin{align*}{@f@}&amp;= {@f1@}{@f2@}\\ &amp; = (e^x)^{-{@a@}} (e^y)^{1/{@b@}}\\ &amp;= \frac{1}{{@A@}^{@a@}}\cdot {{@B@}}^{1/{@b@}} \\&amp;=\frac{1}{@A^a@}\cdot{@B^(1/b)@} \\ &amp;= {@ta1@}.\end{align*}\]</p><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>A:2+rand(4);
B:2+rand(4);
a:2+rand(4);
b:2+rand(4);
eq1:exp(x)=A;
eq2:exp(y)=B;
f:exp(-a*x+ (y/b));
f1:exp(-a*x);
f2:exp(y/b);
f11:exp(x)^(-a);
f22:(exp(y))^(1/b);
ta1: A^(-a)*B^(1/b);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eq1@} and {@eq2@} imply \({@f@} = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2032492988</deployedseed>
    <deployedseed>743652843</deployedseed>
    <deployedseed>629409604</deployedseed>
    <deployedseed>366264562</deployedseed>
    <deployedseed>1381703111</deployedseed>
    <deployedseed>655377958</deployedseed>
    <deployedseed>1823902545</deployedseed>
    <deployedseed>1055625922</deployedseed>
    <deployedseed>1410503370</deployedseed>
    <deployedseed>1929740601</deployedseed>
    <deployedseed>1681216633</deployedseed>
    <deployedseed>1352663591</deployedseed>
    <deployedseed>1374190159</deployedseed>
    <deployedseed>489823731</deployedseed>
    <deployedseed>879273352</deployedseed>
    <deployedseed>787790506</deployedseed>
    <deployedseed>131944599</deployedseed>
    <deployedseed>1595518453</deployedseed>
    <deployedseed>2055733798</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
