<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.1.14 (Intersection of two exps)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the solutions to the equation \({@p1@} = {@p2@}\).&nbsp; Enter your solutions as a set of values.<br></p><p>\(x = \)[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The equation is equivalent to the equation \begin{align*}<br>1= \frac{@p1@}{@p2@}<br>&amp;=e^{\left({@q1@}\right) - \left({@q2@}\right)}\\<br>&amp;= {@p@}. \end{align*}Therefore, since the only value of \(y\) such that \(e^y=1\) is \(y=0\), the solutions are the solutions to the quadratic equation \[{@q@} = 0.\]This factorizes as \[{@factor(q)@}=0,\] from which we see that the solutions are \(x={@x1@}, {@x2@}\).<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a: rand(3)+2
b: rand(3)+1
c: (rand(3)+1)*(a+b)
d: rand_with_prohib(2,7,[a,b])
p1: %e^(a*x^2+(a+b)*d*x)
p2: %e^(-b*x^2+c*x+c*d)
q1:a*x^2+(a+b)*d*x
q2:-b*x^2+c*x+c*d
p: p1/p2
q: expand(((a+b)*x - c)*(x+d))
x1: c/(a+b)
x2: -d
ta1: {c/(a+b), -d}</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@p1@} = {@p2@}\) so \(x={@x1@}, {@x2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{?,?}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>330146412</deployedseed>
    <deployedseed>1096374420</deployedseed>
    <deployedseed>695408733</deployedseed>
    <deployedseed>2136645536</deployedseed>
    <deployedseed>1359800576</deployedseed>
    <deployedseed>1622461012</deployedseed>
    <deployedseed>1166313841</deployedseed>
    <deployedseed>1813292608</deployedseed>
    <deployedseed>1745978012</deployedseed>
    <deployedseed>647071708</deployedseed>
    <deployedseed>1788036697</deployedseed>
    <deployedseed>1591974373</deployedseed>
    <deployedseed>2095865548</deployedseed>
    <deployedseed>1301581357</deployedseed>
    <deployedseed>1294972569</deployedseed>
    <deployedseed>1325179051</deployedseed>
    <deployedseed>2015565614</deployedseed>
    <deployedseed>477531907</deployedseed>
    <deployedseed>2142276488</deployedseed>
    <deployedseed>567624612</deployedseed>
    <deployedseed>481467742</deployedseed>
    <deployedseed>2096066952</deployedseed>
    <deployedseed>1029695575</deployedseed>
    <deployedseed>278394665</deployedseed>
    <deployedseed>1775968747</deployedseed>
    <deployedseed>398999560</deployedseed>
    <deployedseed>878340386</deployedseed>
    <deployedseed>1132666532</deployedseed>
    <deployedseed>45219618</deployedseed>
    <deployedseed>944011872</deployedseed>
    <deployedseed>1216772772</deployedseed>
    <deployedseed>611845496</deployedseed>
    <deployedseed>1409212006</deployedseed>
    <deployedseed>1457651400</deployedseed>
    <deployedseed>1308149854</deployedseed>
    <deployedseed>1179441476</deployedseed>
    <deployedseed>1286297112</deployedseed>
    <deployedseed>460391491</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{x1+1,x2+2}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
