<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.1.15 (Quadratic natural exponential)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Solve the following equation algebraically.</p>
<p>{@quad1@}=0, \(x=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Note that we can write the given equation as a quadratic in \(e^x\). Therefore we can factorise it to obtain
\[
{@quad2@} = 0.
\]
Note that \(e^x = -{@cc@}\) cannot happen since this is negative. Therefore we have that 
\[
e^x = e^{@bb@},
\]
thus since the base numbers are the same we can equate the indices to get that the only solution is \(x = {@ta1@}\).
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa: 2+rand(5);
bb: 1+rand(2);
cc: 1+rand(10);
dd: 1+rand(5);

fact1: e^x-e^bb;
fact2: e^x+cc;

quad1: expand(fact1*fact2);
quad2: fact1*fact2;
ta1: bb;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@quad2@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1714635636</deployedseed>
    <deployedseed>971958113</deployedseed>
    <deployedseed>577873057</deployedseed>
    <deployedseed>42745857</deployedseed>
    <deployedseed>1911513610</deployedseed>
    <deployedseed>1715476101</deployedseed>
    <deployedseed>215274931</deployedseed>
    <deployedseed>668461016</deployedseed>
    <deployedseed>1927471655</deployedseed>
    <deployedseed>1110285511</deployedseed>
    <deployedseed>337719602</deployedseed>
    <deployedseed>839321213</deployedseed>
    <deployedseed>1938990845</deployedseed>
    <deployedseed>703489613</deployedseed>
    <deployedseed>213941800</deployedseed>
    <deployedseed>1584385144</deployedseed>
    <deployedseed>494131430</deployedseed>
    <deployedseed>382271005</deployedseed>
    <deployedseed>1840505454</deployedseed>
    <deployedseed>748099595</deployedseed>
    <deployedseed>1050463433</deployedseed>
    <deployedseed>664484008</deployedseed>
    <deployedseed>912997213</deployedseed>
    <deployedseed>537976962</deployedseed>
    <deployedseed>1961761175</deployedseed>
    <deployedseed>1237295371</deployedseed>
    <deployedseed>519352420</deployedseed>
    <deployedseed>484984517</deployedseed>
    <deployedseed>1576832002</deployedseed>
    <deployedseed>1631884073</deployedseed>
    <deployedseed>1063565102</deployedseed>
    <deployedseed>82008784</deployedseed>
    <deployedseed>1156614227</deployedseed>
    <deployedseed>1932560808</deployedseed>
    <deployedseed>299332449</deployedseed>
    <deployedseed>421540300</deployedseed>
    <deployedseed>1399890807</deployedseed>
    <deployedseed>806444153</deployedseed>
    <deployedseed>613511255</deployedseed>
    <deployedseed>374016761</deployedseed>
    <deployedseed>1810978591</deployedseed>
    <deployedseed>2146208155</deployedseed>
    <deployedseed>1776444037</deployedseed>
    <deployedseed>1881109270</deployedseed>
    <deployedseed>1474191523</deployedseed>
    <deployedseed>1090938572</deployedseed>
    <deployedseed>187418176</deployedseed>
    <deployedseed>223593983</deployedseed>
    <deployedseed>1255258317</deployedseed>
    <deployedseed>733216697</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
