<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.2.23 (Take natural logs)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Solve the equation {@eq@}.&nbsp; Enter \(\ln(2)\) as <code>ln(2)</code>, for example.<br></p><p>\(x = \) [[input:ans1]]. [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Taking logarithms of both sides of the equation {@eq@} we get \[\ln\left({@a@}^{{@lhsexp@}}\right) = \ln\left({@b@}^{{@rhsexp@}}\right),\]which, using the laws of logarithms, is equivalent to\[{@lhs@} = {@rhs@}.\]Expanding and rearranging we get the solution \[x=&nbsp; {@ta1@}.\]<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a:2+rand(2);
b:rand_with_prohib(2,3,[a]);
alp:1+rand(10);
bet:1+rand(5);
gam:1+rand(10);
del:1+rand(5);
lhsexp:(alp*x + bet);
rhsexp:(gam*x-del)
eq:a^(alp*x + bet) = b^(gam*x-del);
lhs:(alp*x + bet)*ln(a);
rhs:(gam*x - del)*ln(b);
top:-del*ln(b) - bet*ln(a);
bot:alp*ln(a) - gam*(ln(b));
ta1:(-del*ln(b) - bet*ln(a))/(alp*ln(a) - gam*(ln(b)));

lhs1:ln(a^(alp*x + bet));
rhs1:ln(b^(gam*x-del));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eq@} has solution \(x = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2004442682</deployedseed>
    <deployedseed>150589642</deployedseed>
    <deployedseed>584530490</deployedseed>
    <deployedseed>820358528</deployedseed>
    <deployedseed>1639549624</deployedseed>
    <deployedseed>1162394848</deployedseed>
    <deployedseed>1542656019</deployedseed>
    <deployedseed>1779502809</deployedseed>
    <deployedseed>2142122201</deployedseed>
    <deployedseed>1497228571</deployedseed>
    <deployedseed>1695168937</deployedseed>
    <deployedseed>174050020</deployedseed>
    <deployedseed>1397445586</deployedseed>
    <deployedseed>1376974535</deployedseed>
    <deployedseed>1568219950</deployedseed>
    <deployedseed>1261678710</deployedseed>
    <deployedseed>2077734187</deployedseed>
    <deployedseed>1225788055</deployedseed>
    <deployedseed>1155302650</deployedseed>
    <deployedseed>433946144</deployedseed>
    <deployedseed>377752777</deployedseed>
    <deployedseed>1301816289</deployedseed>
    <deployedseed>616195434</deployedseed>
    <deployedseed>1344532842</deployedseed>
    <deployedseed>161509058</deployedseed>
    <deployedseed>1963788153</deployedseed>
    <deployedseed>1616053218</deployedseed>
    <deployedseed>642559176</deployedseed>
    <deployedseed>155347223</deployedseed>
    <deployedseed>507146868</deployedseed>
    <deployedseed>96167258</deployedseed>
    <deployedseed>563600015</deployedseed>
    <deployedseed>845028528</deployedseed>
    <deployedseed>1868314924</deployedseed>
    <deployedseed>1063133027</deployedseed>
    <deployedseed>1929518865</deployedseed>
    <deployedseed>1177212247</deployedseed>
    <deployedseed>1530512098</deployedseed>
    <deployedseed>1831652736</deployedseed>
    <deployedseed>75040475</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
