<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.2.4 (Compute logs any base 2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the values of the following.</p>
<p></p>
<ul style="list-style: none;">
<li> \( \log_{{@aa@}} {@bb@} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li>
<li> \( {@aa@}^{\log_{{@aa@}}{@cc@}} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>
</ul>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>(a) \(\log_{{@aa@}}{@bb@} = \log_{{@aa@}} {@aa@}^{{@nn@}} = {@nn@}\).</p>
<p>(b) Since raising to the power and taking logs are inverse operations (previous key point), we have
\( {@aa@}^{\log_{{@aa@}}{@cc@}} = {@cc@} .\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand([2,3,5]);
nn: rand(3)+2;
bb: (aa)^(nn);
cc: rand(10)+20;
ta1:nn;
ta2:cc;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( \log_{{@aa@}} {@bb@} = {@ta1@}, \quad  {@aa@}^{\log_{{@aa@}}{@cc@}} = {@ta2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>log,lg,ln,[[BASIC-ALGEBRA]]</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>log,lg,ln,[[BASIC-ALGEBRA]]</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>143897248</deployedseed>
    <deployedseed>1070274376</deployedseed>
    <deployedseed>1804120166</deployedseed>
    <deployedseed>2009625480</deployedseed>
    <deployedseed>1252429630</deployedseed>
    <deployedseed>922248166</deployedseed>
    <deployedseed>1980094583</deployedseed>
    <deployedseed>1594891895</deployedseed>
    <deployedseed>1948150938</deployedseed>
    <deployedseed>1678100165</deployedseed>
    <deployedseed>788877741</deployedseed>
    <deployedseed>1951496403</deployedseed>
    <deployedseed>343873770</deployedseed>
    <deployedseed>1507862521</deployedseed>
    <deployedseed>160672460</deployedseed>
    <deployedseed>1786099750</deployedseed>
    <deployedseed>2132130447</deployedseed>
    <deployedseed>1094139431</deployedseed>
    <deployedseed>1806986982</deployedseed>
    <deployedseed>1884731501</deployedseed>
    <deployedseed>729409160</deployedseed>
    <deployedseed>1127907636</deployedseed>
    <deployedseed>72303318</deployedseed>
    <deployedseed>296525330</deployedseed>
    <deployedseed>1418762423</deployedseed>
    <deployedseed>1706323432</deployedseed>
    <deployedseed>1499134051</deployedseed>
    <deployedseed>1515370546</deployedseed>
    <deployedseed>549437879</deployedseed>
    <deployedseed>831374987</deployedseed>
    <deployedseed>1405922044</deployedseed>
    <deployedseed>1648772135</deployedseed>
    <deployedseed>1005771284</deployedseed>
    <deployedseed>235328186</deployedseed>
    <deployedseed>2096634379</deployedseed>
    <deployedseed>2026354270</deployedseed>
    <deployedseed>759606159</deployedseed>
    <deployedseed>1943062464</deployedseed>
    <deployedseed>1082221395</deployedseed>
    <deployedseed>666983850</deployedseed>
    <deployedseed>2003337047</deployedseed>
    <deployedseed>1813928963</deployedseed>
    <deployedseed>1919242231</deployedseed>
    <deployedseed>907782971</deployedseed>
    <deployedseed>1133482744</deployedseed>
    <deployedseed>1461107983</deployedseed>
    <deployedseed>1936017920</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
