<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.3.18 (Find temperature natural log 2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The temperature \(T\), in degrees C, of a chemical reaction is given by the formula  
\[\displaystyle{T={@aa@}e^{{@b2@}t}}\]
where \(t\geq 0\) is the time, in seconds.</p>

<p>Calculate the time taken for the temperature to reach \({@cc@}^\circ \) C  . Please give your answer to the nearest second.</p>

<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text>We solve 
\[{@cc@}={@aa@}e^{{@b1@}t}.\]
Rearranging gives
\[e^{{@b1@}t} = \frac{{@cc@}}{{@aa@}}\]
and taking logarithms
\[t = \frac{1}{{@b1@}} \ln\left(\frac{{@cc@}}{{@aa@}}\right) = {@fl@} .\]
So the time is {@ta@} seconds.</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa : 70 + rand(20)
b1 : 2 + rand(5)
b2 : b1/100
cc : 140 + rand(20)
exp : aa*e^(b2*t)
fl: float(log(cc/aa)/b2)
ta : round(log(cc/aa)/b2)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ {@exp@} = {@cc@}.\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1019368800</deployedseed>
    <deployedseed>1127969181</deployedseed>
    <deployedseed>318278591</deployedseed>
    <deployedseed>2066908374</deployedseed>
    <deployedseed>359978827</deployedseed>
    <deployedseed>2104700052</deployedseed>
    <deployedseed>7324307</deployedseed>
    <deployedseed>619581934</deployedseed>
    <deployedseed>193924942</deployedseed>
    <deployedseed>50548011</deployedseed>
    <deployedseed>1503562936</deployedseed>
    <deployedseed>315899701</deployedseed>
    <deployedseed>780549437</deployedseed>
    <deployedseed>228616746</deployedseed>
    <deployedseed>1012667353</deployedseed>
    <deployedseed>1505874766</deployedseed>
    <deployedseed>1157445123</deployedseed>
    <deployedseed>1657012361</deployedseed>
    <deployedseed>283849722</deployedseed>
    <deployedseed>765395117</deployedseed>
    <deployedseed>409492782</deployedseed>
    <deployedseed>1667934830</deployedseed>
    <deployedseed>984827542</deployedseed>
    <deployedseed>1344759726</deployedseed>
    <deployedseed>1697482366</deployedseed>
    <deployedseed>1084524640</deployedseed>
    <deployedseed>1976607725</deployedseed>
    <deployedseed>2104479709</deployedseed>
    <deployedseed>999660749</deployedseed>
    <deployedseed>122014403</deployedseed>
    <deployedseed>461569016</deployedseed>
    <deployedseed>756051184</deployedseed>
    <deployedseed>1480304120</deployedseed>
    <deployedseed>797645916</deployedseed>
    <deployedseed>1068536637</deployedseed>
    <deployedseed>2090592567</deployedseed>
    <deployedseed>809659195</deployedseed>
    <deployedseed>1626161205</deployedseed>
    <deployedseed>1072430771</deployedseed>
    <deployedseed>527851966</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta+10</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
