<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.4.7 (Compound beats simple interest)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>How much more money will you have earnt after {@tt@} years for a £{@PP@} investment at a {@ii@} using compound interest rather than simple interest? Give your answer to the nearest pound.</p>

<p>£[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>In this case we have that \(P = {@PP@},t={@tt@}, i = {@ii@}\). So substituting these values into the simple interest formula gives 
\[
A_{\mbox{simple}} = {@PP@} \cdot (1 + {@tt@}\cdot {@ii@}) \approx {@simp@} \mbox{ pounds}.
\]
and substituting these values into the compound interest formula gives 
\[
A_{\mbox{compound}} = {@PP@} \cdot (1 + {@ii@})^{{@tt@}} \approx {@comp@} \mbox{ pounds}.
\]
Thus you will earn {@ta1@} more money using compound interest.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>PP: 100+rand([100,200,300,400,500]);
tt: 20+rand(20);
ii: rand([0.05,0.1,0.03, 0.01, 0.02]) ;

AA: PP*(1+ii)^tt;
comp: round(AA);

BB: PP*(1 + tt*ii);
simp: round(BB);
ta1: comp-simp;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@comp@}-{@simp@} = {@ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>278310528</deployedseed>
    <deployedseed>1112949841</deployedseed>
    <deployedseed>317411132</deployedseed>
    <deployedseed>1844867102</deployedseed>
    <deployedseed>1257669404</deployedseed>
    <deployedseed>2078272768</deployedseed>
    <deployedseed>233274282</deployedseed>
    <deployedseed>1000657298</deployedseed>
    <deployedseed>1411492715</deployedseed>
    <deployedseed>179480792</deployedseed>
    <deployedseed>284780277</deployedseed>
    <deployedseed>145736533</deployedseed>
    <deployedseed>1445177848</deployedseed>
    <deployedseed>2102528522</deployedseed>
    <deployedseed>385165845</deployedseed>
    <deployedseed>1233864967</deployedseed>
    <deployedseed>883596964</deployedseed>
    <deployedseed>523909138</deployedseed>
    <deployedseed>139810159</deployedseed>
    <deployedseed>2018393295</deployedseed>
    <deployedseed>1388302434</deployedseed>
    <deployedseed>1006235246</deployedseed>
    <deployedseed>1282068161</deployedseed>
    <deployedseed>609424661</deployedseed>
    <deployedseed>1373928561</deployedseed>
    <deployedseed>219309364</deployedseed>
    <deployedseed>1378348080</deployedseed>
    <deployedseed>401574438</deployedseed>
    <deployedseed>2005143813</deployedseed>
    <deployedseed>169760199</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
