<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.4.8 (Simplify sum of logs 2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Simplify the following expression.&nbsp;</p><p>Syntax hint: Enter <code>lg</code> for logs base \(10\).</p>
<p>\(\displaystyle{{@lg1@} - {@lg2@} + {@lg3@} =}\) [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[We have:<p>\({@lg1@} = \log_{10}10^{{@aa@}} = {@aa@}\),</p>
<p>\({@lg2@} = \log_{10}(\frac{10}{{@bb@}}) = \log_{10}10 - \log_{10}{@bb@} = 1 - \log_{10}{@bb@}\),</p>
<p>\({@lg3@} = \log_{10}10^{{@cc@}} = {@cc@}\).</p>
Therefore, the expression is equal to \( {@aa@} - (1 - \log_{10}{@bb@}) + {@cc@} = {@ta@}\).]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand([-1,-2,-3]);
bb: rand([19,21,23,27,29]);
cc: rand([2,3,4]);
lg1: lg(10^(aa));
lg2: lg(10/(bb));
lg3: lg(10^(cc));
ta: aa+cc-1+lg(bb);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@lg1@} - {@lg2@} + {@lg3@} = {@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>/,^,[[BASIC-ALGEBRA]]</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>743727448</deployedseed>
    <deployedseed>1041751693</deployedseed>
    <deployedseed>73037035</deployedseed>
    <deployedseed>2125850265</deployedseed>
    <deployedseed>1080283687</deployedseed>
    <deployedseed>922031417</deployedseed>
    <deployedseed>141803869</deployedseed>
    <deployedseed>638343859</deployedseed>
    <deployedseed>1727727871</deployedseed>
    <deployedseed>1706478866</deployedseed>
    <deployedseed>1905621130</deployedseed>
    <deployedseed>251586984</deployedseed>
    <deployedseed>1909225284</deployedseed>
    <deployedseed>1626740468</deployedseed>
    <deployedseed>1328225615</deployedseed>
    <deployedseed>1072699817</deployedseed>
    <deployedseed>1710336774</deployedseed>
    <deployedseed>1770220825</deployedseed>
    <deployedseed>768746481</deployedseed>
    <deployedseed>86018711</deployedseed>
    <deployedseed>521864314</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
