<?xml version='1.0' encoding='utf-8'?>
<quiz><question type="stack">
    <name>
      <text>Presenting "strings"</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;STACKs support for "strings" has improved during the last few years and the use of textual values e.g. in MCQ options has led to more use cases becoming apparent. However, as inputting "strings" to Maxima requires some escaping to be done like in most programming languages and as that escaping slightly differs from the escaping used in the other languages used (e.g. (X)HTML, JavaScript) when authoring STACK questions there are certain things one needs to note.&lt;/p&gt;

&lt;p&gt;Presenting how a string should be inputted is the easiest part as it uses the same logic as mathematical expressions. Just use the &lt;code&gt;[[escape value="{#expression#}"/]]&lt;/code&gt;-syntax, note that it will also output the wrapping "-chars for the string:&lt;/p&gt;

&lt;table&gt;
&lt;tr&gt;&lt;th&gt;&lt;code&gt;[[escape value="{@str@}"/]]&lt;/code&gt;&lt;/th&gt;&lt;th&gt;&lt;code&gt;[[escape value="{#str#}"/]]&lt;/code&gt;&lt;/th&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@teststring1@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{#teststring1#}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@teststring2@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{#teststring2#}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@teststring3@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{#teststring3#}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@teststring4@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{#teststring4#}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;/table&gt;

&lt;p&gt;When mixing "strings" with (X)HTML or JavaScript you need to note the special escapes those languages need, STACK provides two functions &lt;code&gt;str_to_html(str)&lt;/code&gt; and &lt;code&gt;str_to_js(str)&lt;/code&gt; that do the escaping and entity conversions so that one can simply output the string where needed.&lt;/p&gt;


&lt;!-- Note. that for display reasons we actually escape the HTML-output twice to show that we escaped it... In practical use (e.g. attribute values for 'style' or such) you probably only want to escape it once. --&gt;
&lt;table&gt;
&lt;tr&gt;&lt;th&gt;&lt;code&gt;[[escape value="{@str_to_html(str)@}"/]]&lt;/code&gt;&lt;/th&gt;&lt;th&gt;&lt;code&gt;[[escape value="{@str_to_js(str)@}"/]]&lt;/code&gt;&lt;/th&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@str_to_html(str_to_html(teststring1))@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{@str_to_js(teststring1)@}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@str_to_html(str_to_html(teststring2))@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{@str_to_js(teststring2)@}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@str_to_html(str_to_html(teststring3))@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{@str_to_js(teststring3)@}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;
&lt;tr&gt;&lt;td&gt;&lt;code&gt;{@str_to_html(str_to_html(teststring4))@}&lt;/code&gt;&lt;/td&gt;&lt;td&gt;&lt;code&gt;{@str_to_js(teststring4)@}&lt;/code&gt;&lt;/td&gt;&lt;/tr&gt;&lt;/table&gt;

&lt;p&gt;Note! Moodles output filtering does often handle &amp;, &lt; and other such characters that have special meanings in HTML automatically in the textual content of HTML-output but when you wish to push those characters  to attributes or more complex code you will need to do these messy things yourself.&lt;/p&gt;

&lt;p&gt;Also remember that while you may construct complex "strings" at the CAS level some things may be better left to the CASText level. While some "string" trickery may be necessary when dealing with randomly generated options for MCQ questions generally one should try to leave the textual content of the question outside of the CAS, both of the following examples generate the same output but the later one is the preferred one:&lt;/p&gt;
&lt;ul&gt;
&lt;li&gt;&lt;code&gt;[[ escape value='{@sconcat("The derivative of \\(f(x)\\) is \\(", tex1(diff(f(x),x)), "\\).")@}' /]]&lt;/code&gt;&lt;/li&gt;
&lt;li&gt;&lt;code&gt;[[ escape value="The derivative of {@'f(x)@} is {@diff(f(x),x)@}." /]]&lt;/code&gt;&lt;/li&gt;
&lt;/ul&gt;

&lt;p&gt;Note that if you need to build the output conditionally or loop over something there are new tools for those tasks and you generally do not need to do that at CAS-level. Check the other examples in the question set for those.&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text />
    </generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <questionvariables>
      <text>teststring1: "abcd...XYZÅÄÖ";
teststring2: "Quotes in \"strings\": 'test'";
teststring3: "\\LaTeX like this \\frac{a}{b}";
teststring4: "&amp; &lt; &gt;";

f(x):=x^2+sin(x)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text />
    </specificfeedback>
    <questionnote>
      <text />
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
  </question>

</quiz>