<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>HELM 30.3 E4 LU-Decomposition with permutation</text>
    </name>
    <questiontext format="markdown">
      <text>**30.3 E4 LU Decomposition**

Calculate an \\\(LU\\\) decomposition for this matrix:

\\\[{@M@}\\\]

\\\(L=\\\)[[input:S_L]] [[validation:S_L]] [[feedback:L]]

\\\(U=\\\)[[input:S_U]] [[validation:S_U]] [[feedback:U]]

Using this decomposition, solve the following system of equations \\\(Mx=r\\\). Note that you have to swap rows in \\\(r\\\) also, if you swap them in \\\(M\\\) before factorisation:

\\\({@M@}\,{@matrix([x1],[x2],[x3])@}={@rhs@}.\\\)

First, calculate the intermediate vector \\\(y\\\) with \\\(Ly=r\\\). 

\\\(y=\\\) [[input:S_y]] [[validation:S_y]] [[feedback:y]]

and then the solution \\\(x\\\) with \\\(Ux=y\\\):

\\\(x=\\\) [[input:S_x]] [[validation:S_x]] [[feedback:x]]</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024012900</text>
    </stackversion>
    <questionvariables>
      <text>/* compute random L and U and M=L.U */

m:5;

L0: matrix([1, 0, 0],
          [rand_with_prohib(-m,m,[0]), 1, 0],
          [rand_with_prohib(-m,m,[0]), rand_with_prohib(-m,m,[0]), 1]);
U0: matrix([rand_with_prohib(-m,m,[0]), rand_with_prohib(-m,m,[0]), rand_with_prohib(-m,m,[0])],
          [0, rand_with_prohib(-m,m,[0]), rand_with_prohib(-m,m,[0])],
          [0, 0, rand_with_prohib(-m,m,[0])] );

M0: L0.U0;

/* apply a row operation with unknown factor and determine k such that the second subdeterminant is zero*/
M: rowop(M0,2,3,k);
[M,k]: ev([M,k], solve(determinant(submatrix(3,M,3)),k));
/* multiply row 2 by a factor such that no fractions remain */
d: lcm(denom(M[2,1]),denom(M[2,2]),denom(M[2,3]));
M: ratsimp(matrix([1,0,0],[0,d,0],[0,0,1]).M); 
/* LU decomposition of M */
MS: rowswap(M,2,3)
LU: lu_factor(MS)[1];
L: LU*matrix([1,0,0], [1,1,0],[1,1,1]);
L[1,1]:1; L[2,2]:1; L[3,3]:1;
U: LU*matrix([1,1,1], [0,1,1],[0,0,1]);
M1:L.U;

x: matrix([rand_with_prohib(-m,m,[0])], [rand_with_prohib(-m,m,[0])], [rand_with_prohib(-m,m,[0])]);
y: U.x;
rhs: M.x;
rhs0:rowswap(rhs,2,3);

rhsc: M.x;</text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@M@}, {@L@}, {@U@} {@y@} {@rhs@} {@x@}</text>
    </questionnote>
    <questiondescription format="markdown">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>S_L</name>
      <type>matrix</type>
      <tans>L</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>S_U</name>
      <type>matrix</type>
      <tans>U</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>S_x</name>
      <type>matrix</type>
      <tans>x</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>S_y</name>
      <type>matrix</type>
      <tans>y</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>L</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>S_L</sans>
        <tans>L</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>L-1-T</trueanswernote>
        <truefeedback format="markdown">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>L-1-F</falseanswernote>
        <falsefeedback format="markdown">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>U</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>S_U</sans>
        <tans>U</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>U-1-T</trueanswernote>
        <truefeedback format="markdown">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>U-1-F</falseanswernote>
        <falsefeedback format="markdown">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>x</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>S_x</sans>
        <tans>x</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>x-1-T</trueanswernote>
        <truefeedback format="markdown">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>x-1-F</falseanswernote>
        <falsefeedback format="markdown">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>y</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>S_y</sans>
        <tans>y</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>y-1-T</trueanswernote>
        <truefeedback format="markdown">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>y-1-F</falseanswernote>
        <falsefeedback format="markdown">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1578832493</deployedseed>
    <deployedseed>325160686</deployedseed>
    <deployedseed>577519026</deployedseed>
    <deployedseed>429458189</deployedseed>
    <deployedseed>596581985</deployedseed>
    <deployedseed>534864407</deployedseed>
    <deployedseed>1015426746</deployedseed>
    <deployedseed>1829364594</deployedseed>
    <deployedseed>2090378208</deployedseed>
    <deployedseed>977998961</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_L</name>
        <value>L</value>
      </testinput>
      <testinput>
        <name>S_U</name>
        <value>U</value>
      </testinput>
      <testinput>
        <name>S_x</name>
        <value>x</value>
      </testinput>
      <testinput>
        <name>S_y</name>
        <value>y</value>
      </testinput>
      <expected>
        <name>L</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>L-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>U</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>U-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>x</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>x-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>y</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>y-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>S_L</name>
        <value>L</value>
      </testinput>
      <testinput>
        <name>S_U</name>
        <value>U</value>
      </testinput>
      <testinput>
        <name>S_x</name>
        <value>x</value>
      </testinput>
      <testinput>
        <name>S_y</name>
        <value>y</value>
      </testinput>
      <expected>
        <name>L</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>L-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>U</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>U-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>x</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>x-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>y</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>y-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
