<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.3.5.3 Factorising quadratics with a*x^2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Factorise:<br></p>

<ol class="HELM_parts">

<li><p> \( {@a1@}x^2 + {@x@} =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

<li><p> \( {@b1@}x^2 + {@b1+1@}x + 1 =\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

<li><p> \( {@c1*c2@}x^2 + {@c1+c2@}x + 1 =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p></li>

<li><p> \( {@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@}=\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p></li>

<li><p> \( {@e1*e2@}x^2 + {@e2*e3 - e1*e4@}x&nbsp; - {@e3*e4@}=\) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</p></li>

<li><p> \( {@-f1@}x^2 + {@f3 - f1@}x + {@f3@}=\) [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</p></li>

</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">
    <li>
        <p>Note that \(x\) is a common factor of both terms. Thus</p>
        <p>\[ {@a1@}x^2 + x = x({@a1@}x + 1). \]</p>
    </li>

    <li>
        <p><span style="font-size: 0.9375rem;">First note the coefficient of \(x^2\) in this case is \({@b1@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;</span></p>
        <p>\[ {@b1@}({@b1@}x^2 + {@b1+1@}x + 1) = {@b1@}({@b1@}x^2) + {@b1@}({@1+b1@})x + {@b1@}(1) = ({@b1@}x)^2 +&nbsp;{@1+b1@}({@b1@}x)&nbsp; + {@b1@}. \]&nbsp;</p>
        <p>We now introduce a new variable \(z\) such that \(z = {@b1@}x\). Thus we can write&nbsp;</p>
        <p>\[ ({@b1@}x)^2 +&nbsp;{@1+b1@}({@b1@}x)&nbsp; + {@b1@} \text{ as } z^2 +&nbsp;{@1+b1@}z + {@b1@}. \]&nbsp;</p>
        <p>This can be factorised to give \( (z + {@b1@})(z + 1\). Returning to the original variable by replacing \(z\) by \({@b1@}x\) we find</p>
        <p>&nbsp;\[ &nbsp;{@b1@}({@b1@}x^2 + {@b1+1@}x + 1) = ({@b1@}x + {@b1@})({@b1@}x + 1) \]&nbsp;</p>
        <p>A factor of \({@b1@}\) can be extracted from the first bracket on the right so that&nbsp;</p>
        <p>\[ {@b1@}({@b1@}x^2 + {@b1+1@}x + 1)&nbsp; = {@b1@}(x + 1)({@b1@}x + 1), \]&nbsp;</p>
        <p>so that&nbsp;</p>
        <p>\[ {@b1@}x^2 + {@b1+1@}x + 1 = (x + 1)({@b1@}x + 1). \]</p>
    </li>

    <li>
        <p><span style="font-size: 0.9375rem;">First note the coefficient of \(x^2\) in this case is \({@c1*c2@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;</span></p>
        <p>\[ {@c1*c2@}({@c1*c2@}x^2 + {@c1+c2@}x + 1) = {@c1*c2@}({@c1*c2@}x^2) + {@c1*c2@}({@c1+c2@})x - {@c1*c2@}(1) = ({@c1*c2@}x)^2 + {@c1+c2@} ({@c1*c2@}x)&nbsp; - {@c1*c2@}. \]&nbsp;</p>
        <p>We now introduce a new variable \(z\) such that \(z = {@c1*c2@}x\). Thus we can write&nbsp;</p>
        <p>\[ ({@c1*c2@}x)^2 + {@c1+c2@} ({@c1*c2@}x)&nbsp; - {@c1*c2@} \text{ as } z^2 +{@c1+c2@} z - {@c1*c2@}. \]</p>
        <p>This can be factorised to give \( (z + {@c1@})(z + {@c2@})\). Returning to the original variable by replacing \(z\) by \({@c@}x\) we find</p>
        <p>\[ {@c1*c2@}({@c1*c2@}x^2 + {@c1+c2@}x + 1) = ({@c1*c2@}x + {@c1@})({@c1*c2@}x + {@c2@}). \]</p>
        <p>A factor of \({@c1@}\) can be extracted from the first bracket on the right, and a factor of \({@c2@}\) can be extracted from the second bracket on the right, so that</p>
        <p>\[ {@c1*c2@}({@c1*c2@}x^2 + {@c1+c2@}x + 1) = {@c1*c2@}({@c2@}x + 1)({@c1@}x + 1), \]</p>
        <p>so that</p>
        <p>\[ {@c1*c2@}x^2 + {@c1+c2@}x + 1 = ({@c2@}x + 1)({@c1@}x + 1). \]</p>
    </li>

    <li>
        <p>First note the coefficient of \(x^2\) in this case is \({@d1*d2@}\). Multiply the whole expression by this number and rearrange as follows:</p>
        <p>\[ {@d1*d2@}({@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@}) = {@d1*d2@}({@d1*d2@}x^2) + {@d1*d2@}({@d1*d3+d2*d4@})x + {@d1*d2@}({@d3*d4@}) = ({@d1*d2@}x)^2 + {@d1*d3+d2*d4@} ({@d1*d2@}x) + {@d1*d2*d3*d4@}. \]</p>
        <p>We now introduce a new variable \(z\) such that \(z = {@d1*d2@}x\). Thus we can write</p>
        <p>\[ ({@d1*d2@}x)^2 + {@d1*d3+d2*d4@} ({@d1*d2@}x) + {@d1*d2*d3*d4@} \text{ as } z^2 +{@d1*d3+d2*d4@} z + {@d1*d2*d3*d4@}. \]</p>
        <p>This can be factorised to give \( (z + {@d1*d3@})(z + {@d2*d4@})\). Returning to the original variable by replacing \(z\) by \({@d1*d2@}x\) we find</p>
        <p>\[ {@d1*d2@}({@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@}) = ({@d1*d2@}x + {@d1*d3@})({@d1*d2@}x + {@d2*d4@}) \]</p>
        <p>A factor of \({@d1@}\) can be extracted from the first bracket on the right, and a factor of \({@d2@}\) can be extracted from the second bracket on the right, so that/p&gt;
        </p><p>\[ {@d1*d2@}({@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@}) =({@d1*d2@}({@d2@}x + {@d3@})({@d1@}x + {@d4@}) \]</p>
        <p>so that</p>
        <p>\[ {@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@} = ({@d2@}x + {@d3@})({@d1@}x + {@d4@}). \]</p>
    </li>

    <li>
        <p>First note the coefficient of \(x^2\) in this case is \({@e1*e2@}\). Multiply the whole expression by this number and rearrange as follows:</p>
        <p>\[ {@e1*e2@}({@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@}) = {@e1*e2@}({@e1*e2@}x^2) + {@e1*e2@}({@-e1*e4+e2*e3@})x - {@e1*e2@}({@e3*e4@}) = ({@e1*e2@}x)^2 + {@-e1*e4+e2*e3@} ({@e1*e2@}x)- {@e1*e2*e3*e4@}. \]</p>
        <p>We now introduce a new variable \(z\) such that \(z = {@e1*e2@}x\). Thus we can write</p>
        <p>\[ ({@e1*e2@}x)^2 + {@-e1*e4+e2*e3@} ({@e1*e2@}x) - {@e1*e2*e3*e4@} \text{ as } z^2 +{@-e1*e4+e2*e3@} z - {@e1*e2*e3*e4@}. \]</p>
        <p>This can be factorised to give \( (z - {@e1*e4@})(z + {@e2*e3@})\). Returning to the original variable by replacing \(z\) by \({@e1*e2@}x\) we find</p>
        <p>\[ {@e1*e2@}({@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@}) = ({@e1*e2@}x - {@e1*e4@})({@e1*e2@}x + {@e2*e3@}) \]</p>
        <p>A factor of \({@e1@}\) can be extracted from the first bracket on the right, and a factor of \({@e2@}\) can be extracted from the second bracket on the right, so that</p>
        <p>\[ {@e1*e2@}({@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@}) =({@e1*e2@}({@e2@}x - {@e4@})({@e1@}x + {@e3@}) \]</p>
        <p>so that</p>
        <p>\[ {@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@} = ({@e2@}x - {@e4@})({@e1@}x + {@e3@}). \]</p>
    </li>

    <li>
        <p>First note the coefficient of \(x^2\) in this case is \({@-f1@}\). Multiply the whole expression by this number and rearrange as follows:</p>
        <p>\[ {@-f1@}({@-f1@}x^2 + {@f3 - f1@}x + {@f3@}) = {@-f1@}({@-f1@}x^2) {@-f1@}({@f3 -f1@})x + {@-f1@}({@f3@} = ({@f1@}x)^2 -{@f3 -f1@}({@f1@}x) - {@f1*f3@}. \]</p>
        <p>We now introduce a new variable \(z\) such that \(z = {@f1@}x\). Thus we can write</p>
        <p>\[ ({@f1@}x)^2 -{@f3 -f1@}({@f1@}x); - {@f1*f3@} \text{ as } z^2 -{@f3 -f1@}z - {@f1*f3@}. \]</p>
        <p>This can be factorised to give \( (z + {@f1@})(z - {@f3@}\). Returning to the original variable by replacing \(z\) by \({@f1@}x\) we find</p>
        <p>\[ {@-f1@}({@-f1@}x^2 + {@f3 - f1@}x + {@f3@}) = ({@f1@}x + {@f1@})({@f1@}x - {@f3@}) \]</p>
        <p>A factor of \({@f1@}\) can be extracted from the first bracket on the right, and a factor of /(-1/) can be extracted from the second bracket on the right so that</p>
        <p>\[ {@-f1@}({@-f1@}x^2 + {@f3 - f1@}x + {@f3@}) = -{@f1@}(x + 1)(-{@f1@}x + {@f3@}), \]</p>
        <p>so that</p>
        <p>\[ {@-f1@}x^2 + {@f3 - f1@}x + {@f3@} = (x + 1)(-{@f1@}x + {@f3@}). \]</p>
    </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>a1: rand(2)+3;
ta1: x*(a1*x + 1);
F1:a1*x^2 + x;

b1: rand([3,4]);
ta2: (b1*x+1)*(x + 1);
F2:b1*x^2 + (b1+1)*x + 1;

c:3;
c1: 2;
c2:c1+1;
ta3: (c1*x + 1)*(c2*x + 1);
F3: c1*c2*x^2 + (c1+c2)*x + 1;

d1: 2;
d2:d1+1;
d3:rand([4,5]);
d4:rand([3,5]);
ta4: (d1*x + d4)*(d2*x + d3);
F4: d1*d2*x^2 + (d1*d3+d2*d4)*x + d3*d4;

e1: 2;
e2:e1+1
e3:rand([5,7])
e4:2
ta5: (e1*x + e3)*(e2*x - e4);
F5:e1*e2*x^2 + (e2*e3 - e1*e4)*x  - e3*e4;

f1: 3;
f3:rand([4,5,7]);
ta6: (-f1*x+ f3)*(x + 1);
F6:-f1*x^2 + (f3 - f1)*x + f3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@ta1@}, {@ta2@}, {@ta3@}, {@ta4@} , {@ta5@}, {@ta6@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><div title="Page 2">
			<div>
				<div>
					<div>
						<p></p></div></div></div></div><p></p>Note that \(x\) is a common factor of both terms. Thus<p>\[ {@a1@}x^2 + x = x({@a1@}x + 1). \]&nbsp;</p><p><br></p><p></p><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p></p><p></p>First note the coefficient of \(x^2\) in this case is \({@b1@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;<p>\[ {@b1@}({@b1@}x^2 + {@b1+1@}x + 1) = {@b1@}({@b1@}x^2) + {@b1@}({@1+b1@})x + {@b1@}(1) = ({@b1@}x)^2 +&nbsp;{@1+b1@}({@b1@}x)&nbsp; + {@b1@}. \]&nbsp;</p><p>We now introduce a new variable \(z\) such that \(z = {@b1@}x\). Thus we can write&nbsp;</p><p>\[ ({@b1@}x)^2 +&nbsp;{@1+b1@}({@b1@}x)&nbsp; + {@b1@} \text{ as } z^2 +&nbsp;{@1+b1@}z + {@b1@}. \]&nbsp;</p><p>This can be factorised to give \( (z + {@b1@})(z + 1\). Returning to the original variable by replacing \(z\) by \({@b1@}x\) we find</p><p>&nbsp;\[ &nbsp;{@b1@}({@b1@}x^2 + {@b1+1@}x + 1) = ({@b1@}x + {@b1@})({@b1@}x + 1) \]&nbsp;</p><p>A factor of \({@b1@}\) can be extracted from the first bracket on the right so that&nbsp;</p><p>\[ {@b1@}({@b1@}x^2 + {@b1+1@}x + 1)&nbsp; = {@b1@}(x + 1)({@b1@}x + 1), \]&nbsp;</p><p>so that&nbsp;</p><p>\[ {@b1@}x^2 + {@b1+1@}x + 1 = (x + 1)({@b1@}x + 1). \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p></p><p></p>First note the coefficient of \(x^2\) in this case is \({@c1*c2@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;<p>\[ {@c1*c2@}({@c1*c2@}x^2 + {@c1+c2@}x + 1) = {@c1*c2@}({@c1*c2@}x^2) + {@c1*c2@}({@c1+c2@})x - {@c1*c2@}(1) = ({@c1*c2@}x)^2 + {@c1+c2@} ({@c1*c2@}x)&nbsp; - {@c1*c2@}. \]&nbsp;</p><p>We now introduce a new variable \(z\) such that \(z = {@c1*c2@}x\). Thus we can write&nbsp;</p><p>\[ ({@c1*c2@}x)^2 + {@c1+c2@} ({@c1*c2@}x)&nbsp; - {@c1*c2@} \text{ as } z^2 +{@c1+c2@} z - {@c1*c2@}. \]&nbsp;</p><p>This can be factorised to give \( (z + {@c1@})(z + {@c2@})\). Returning to the original variable by replacing \(z\) by \({@c@}x\) we find</p><p>&nbsp;\[ {@c1*c2@}({@c1*c2@}x^2 + {@c1+c2@}x + 1)&nbsp; = &nbsp;({@c1*c2@}x + {@c1@})({@c1*c2@}x + {@c2@}). \]&nbsp;</p><p>A factor of \({@c1@}\) can be extracted from the first bracket on the right, and&nbsp; a factor of \({@c2@}\) can be extracted from the second bracket on the right, so that&nbsp;</p><p>\[ {@c1*c2@}({@c1*c2@}x^2 + {@c1+c2@}x + 1)&nbsp;&nbsp; = {@c1*c2@}({@c2@}x + 1)({@c1@}x + 1), \]&nbsp;</p><p>so that&nbsp;</p><p>\[ {@c1*c2@}x^2 + {@c1+c2@}x + 1 = &nbsp;({@c2@}x + 1)({@c1@}x + 1). \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p></p><p></p>First note the coefficient of \(x^2\) in this case is \({@d1*d2@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;<p>\[ {@d1*d2@}({@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@}) = {@d1*d2@}({@d1*d2@}x^2) + {@d1*d2@}({@d1*d3+d2*d4@})x + {@d1*d2@}({@d3*d4@}) = ({@d1*d2@}x)^2 + {@d1*d3+d2*d4@} ({@d1*d2@}x)&nbsp; + {@d1*d2*d3*d4@}. \]&nbsp;</p><p>We now introduce a new variable \(z\) such that \(z = {@d1*d2@}x\). Thus we can write&nbsp;</p><p>\[ ({@d1*d2@}x)^2 + {@d1*d3+d2*d4@} ({@d1*d2@}x)&nbsp; + {@d1*d2*d3*d4@} \text{ as } z^2 +{@d1*d3+d2*d4@} z + {@d1*d2*d3*d4@}. \]&nbsp;</p><p>This can be factorised to give \( (z + {@d1*d3@})(z + {@d2*d4@})\). Returning to the original variable by replacing \(z\) by \({@d1*d2@}x\) we find</p><p>&nbsp;\[ {@d1*d2@}({@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@}) = ({@d1*d2@}x + {@d1*d3@})({@d1*d2@}x + {@d2*d4@}) \]&nbsp;</p><p>A factor of \({@d1@}\) can be extracted from the first bracket on the right, and a&nbsp;factor of \({@d2@}\) can be extracted from the second bracket on the right, so that&nbsp;</p><p>\[ {@d1*d2@}({@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@})&nbsp; =({@d1*d2@}({@d2@}x + {@d3@})({@d1@}x + {@d4@}) \]&nbsp;</p><p>so that&nbsp;</p><p>\[ {@d1*d2@}x^2 + {@d1*d3+d2*d4@}x + {@d3*d4@} = ({@d2@}x + {@d3@})({@d1@}x + {@d4@}). \]</p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p>First note the coefficient of \(x^2\) in this case is \({@e1*e2@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;<p>\[ {@e1*e2@}({@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@}) = {@e1*e2@}({@e1*e2@}x^2) + {@e1*e2@}({@-e1*e4+e2*e3@})x - {@e1*e2@}({@e3*e4@}) = ({@e1*e2@}x)^2 + {@-e1*e4+e2*e3@} ({@e1*e2@}x)&nbsp; - {@e1*e2*e3*e4@}. \]&nbsp;</p><p>We now introduce a new variable \(z\) such that \(z = {@e1*e2@}x\). Thus we can write&nbsp;</p><p>\[ ({@e1*e2@}x)^2 + {@-e1*e4+e2*e3@} ({@e1*e2@}x)&nbsp; - {@e1*e2*e3*e4@} \text{ as } z^2 +{@-e1*e4+e2*e3@} z &nbsp;- {@e1*e2*e3*e4@}. \]&nbsp;</p><p>This can be factorised to give \( (z - {@e1*e4@})(z + {@e2*e3@})\). Returning to the original variable by replacing \(z\) by \({@e1*e2@}x\) we find</p><p>&nbsp;\[ &nbsp;{@e1*e2@}({@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@}) = ({@e1*e2@}x - {@e1*e4@})({@e1*e2@}x + {@e2*e3@}) \]&nbsp;</p><p>A factor of \({@e1@}\) can be extracted from the first bracket on the right, and a&nbsp;factor of \({@e2@}\) can be extracted from the second bracket on the right, so that&nbsp;</p><p>\[ &nbsp;{@e1*e2@}({@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@})&nbsp; =({@e1*e2@}({@e2@}x - {@e4@})({@e1@}x + {@e3@}) \]&nbsp;</p><p>so that&nbsp;</p><p>\[ {@e1*e2@}x^2 + {@-e1*e4+e2*e3@}x - {@e3*e4@} = ({@e2@}x - {@e4@})({@e1@}x + {@e3@}). \]</p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p>First note the coefficient of \(x^2\) in this case is \({@-f1@}\). Multiply the whole expression by this number and rearrange as follows:&nbsp;<p>\[ {@-f1@}({@-f1@}x^2 + {@f3 - f1@}x + {@f3@}) = {@-f1@}({@-f1@}x^2)&nbsp; {@-f1@}({@f3 -f1@})x + {@-f1@}({@f3@} = ({@f1@}x)^2 -&nbsp;{@f3 -f1@}({@f1@}x)&nbsp; - {@f1*f3@}. \]&nbsp;</p><p>We now introduce a new variable \(z\) such that \(z = {@f1@}x\). Thus we can write&nbsp;</p><p>\[ ({@f1@}x)^2 -&nbsp;{@f3 -f1@}({@f1@}x)&nbsp; - {@f1*f3@} \text{ as } z^2 -&nbsp;{@f3 -f1@}z - {@f1*f3@}. \]&nbsp;</p><p>This can be factorised to give \( (z + {@f1@})(z - {@f3@}\). Returning to the original variable by replacing \(z\) by \({@f1@}x\) we find</p><p>&nbsp;\[ &nbsp;{@-f1@}({@-f1@}x^2 + {@f3 - f1@}x + {@f3@}) = ({@f1@}x + {@f1@})({@f1@}x - {@f3@}) \]&nbsp;</p><p>A factor of \({@f1@}\) can be extracted from the first bracket on the right, and a factor of /(-1/) can be extracted from the second bracket on the right so that&nbsp;</p><p>\[ &nbsp;{@-f1@}({@-f1@}x^2 + {@f3 - f1@}x + {@f3@})&nbsp; =&nbsp; -{@f1@}(x + 1)(-{@f1@}x + {@f3@}), \]&nbsp;</p><p>so that&nbsp;</p><p>\[ {@-f1@}x^2 + {@f3 - f1@}x + {@f3@} = (x + 1)(-{@f1@}x + {@f3@}). \]</p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1171329110</deployedseed>
    <deployedseed>1033748056</deployedseed>
    <deployedseed>321709849</deployedseed>
    <deployedseed>1811026098</deployedseed>
    <deployedseed>1772660814</deployedseed>
    <deployedseed>1827785156</deployedseed>
    <deployedseed>2139129640</deployedseed>
    <deployedseed>808785925</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4+1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5+1</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>F1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>f2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>f3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>f4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>f6</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>f6</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
