<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.3.3.3 Removing 2 sets of brackets</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Remove the brackets from each of the following expressions:<br></p>

<ol class="HELM_parts">

<li><p>\( ({@a1@}+x)({@a2@}+x)=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

<li><p>\( (x + {@b1@})(x + {@b2@})=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

<li><p> \((x + {@c1@})(x + {@c2@})=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p></li>

<li><p> \((x + {@d1@})(x - {@d2@})=\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p></li>

<li><p>\( ({@e1@}+x)x=\) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</p></li>

<li><p>\( ({@f1@}x + 1)x=\) [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</p></li>

<li><p> \( ( {@g1@}x + 1)(x + 1)=\) [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</p></li>

<li><p> \(( {@h1@}x + 1)( {@h2@}x + 1)=\) [[input:ans8]] [[validation:ans8]] [[feedback:prt8]]</p></li>

<li><p>\( ({@i1@}x+ {@i2@}) ( {@i3@}x + {@i4@} )=\) [[input:ans9]] [[validation:ans9]] [[feedback:prt9]]</p></li>

<li><p>\( ({@j1@}x+ {@j2@}) ( {@j3@}x - {@j4@} )=\) [[input:ans10]] [[validation:ans10]] [[feedback:prt10]]</p></li>

<li><p> \( ( {@k1@} - {@k2@}x)(x + 1)=\) [[input:ans11]] [[validation:ans11]] [[feedback:prt11]]</p></li>

<li><p> \(( {@l1@} - x)(1 - x)=\) [[input:ans12]] [[validation:ans12]] [[feedback:prt12]]</p></li>

</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol class="HELM_parts">
    <li>
        <p></p>
        <p>\[ ({@a1@}+x) ({@a2@}+x) =&nbsp;({@a1@}+x) ({@a2@}) + ({@a1@}+x)(x) \]</p>
        <p>\[ = ({@a1@})({@a2@}) + (x)({@a2@}) + ({@a1@})(x) + (x)(x) \]</p>
        <p>\[ = {@a1*a2@} + {@a1+a2@}x + x^2. \]&nbsp;</p>
        <p></p>
    </li>
    <li>
        <p><span style="font-size: 0.9375rem;"></span></p>
        <p>\[ (x + {@b1@}) (x + {@b2@})= (x + {@b1@}) (x) +&nbsp;(x + {@b1@})&nbsp;({@b2@}) \]</p>
        <p>\[ = (x)(x) + ({@b1@})(x) + ({@b2@})(x) + ({@b1@})({@b2@}) \]</p>
        <p>\[ = x^2 + {@b1 + b2@}x + {@b1*b2@}. \]&nbsp;</p><br>
        <p></p>
    </li>
    <li>
        <p><span style="font-size: 0.9375rem;"></span></p>
        <p>\[ (x + {@c1@}) (x + {@c2@})= (x + {@c1@}) (x) + (x + {@c1@})({@c2@}) \]<br></p>
        <p></p>
        <p>\[ = (x)(x) + ({@c1@})(x) + ({@c2@})(x) + ({@c1@})({@c2@}) \]</p>
        <p>\[ = x^2 + {@c1 + c2@}x + {@c1*c2@}. \]</p><br>
        <p></p>
    </li>
    <li>
        <p><span style="font-size: 0.9375rem;"></span></p>
        <p>\[ (x + {@d1@}) (x - {@d2@})= (x + {@d1@}) (x) - (x + {@d1@})({@d2@}) \]</p>
        <p>\[ = (x)(x) + ({@d1@})(x) - ({@d2@})(x) - ({@d1@})({@d2@}) \]</p>
        <p>\[ = x^2 + {@d1 - d2@}x - {@d1*d2@}. \]</p>
        <p></p>
    </li>
    <li>
        <p><span style="font-size: 0.9375rem;"></span></p>
        <p>In the expression \( ({@e1@}+x)x \), the \( {@x@} \) outside the bracket must multiply both the \({@e1@}\) and the \(x\) in the brackets. Thus</p>
        <p>\[ ({@e1@}+x)x = {@e1@}x + x \times x \]</p>
        <p>\[ =&nbsp;{@e1@}x + x^2 \]</p>
        <p></p>
    </li>
    <li>
        <p></p>
        <p>In the expression \( ({@f1@}x + 1)x \), the \( {@x@} \) outside the bracket must multiply both the \({@f1@}x\) and the \(1\) in the brackets. Thus</p>
        <p>\[ ({@f1@}x + 1)x = {@f1@}x \times x + 1 \times x \]</p>
        <p>\[ {@ta6@} \]</p>
    </li>
    <li><span style="font-size: 0.9375rem;">\[ ( {@g1@}x + 1)(x + 1)&nbsp;= ( {@g1@}x + 1)(x) + ( {@g1@}x + 1)(1) \]&nbsp;</span>\[ = {@g1@}x(x)+ (1)(x) + {@g1@}x(1) +1(1) \] \[ = {@ta7@} \]</li>
    <li>
        <p><span style="font-size: 0.9375rem;">\[ ( {@h1@}x + 1)( {@h2@}x + 1)= ( {@h1@}x + 1)({@h2@}x) + ( {@h1@}x + 1)(1) \]&nbsp;</span><span style="font-size: 0.9375rem;">\[ = {@h1@}x({@h2@}x)+ (1)({@h2@}x) + {@h1@}x(1) +1(1) \]&nbsp;</span><span style="font-size: 0.9375rem;">\[ = {@ta8@} \]</span></p>
    </li>
    <li>
        <p>\[ ({@i1@}x+ {@i2@}) ( {@i3@}x + {@i4@} )= ({@i1@}x+ {@i2@})({@i3@}x) + ({@i1@}x+ {@i2@})({@i4@}) \]</p>
        <p>\[ = {@i1@}x({@i3@}x)+ ({@i2@})({@h3@}x) + {@i1@}x({@i4@}) +{@i2@}({@i4@}) \]</p>
        <p>\[ = {@ta9@} \]</p>
    </li>
    <li>
        <p>\[ ({@j1@}x+ {@j2@}) ( {@j3@}x - {@j4@} )= ({@j1@}x+ {@j2@})({@j3@}x) + ({@j1@}x+ {@j2@})(-{@j4@}) \]</p>
        <p>\[ = {@j1@}x({@j3@}x)+ ({@j2@})({@j3@}x) - {@j1@}x({@j4@}) - {@j2@}({@j4@}) \]</p>
        <p>\[ = {@ta10@} \]</p>
    </li>
    <li>
        <p>\[ ( {@k1@} - {@k2@}x)(x + 1)= ( {@k1@} - {@k2@}x)(x) + ( {@k1@} - {@k2@}x)(1) \]</p>
        <p>\[ = {@k1@}(x) - {@k2@}(x) + {@k1@}(1) - {@k2@}(1) \]</p>
        <p>\[ = {@ta11@} \]</p>
    </li>
    <li>
        <p>\[ ( {@l1@} - x)(1 - x)= ( {@l1@} - x)(1) + ( {@l1@} - x)(-x) \]</p>
        <p>\[ = {@l1@}(1) - x(1) + {@l1@}(-x) - x(-x) \]</p>
        <p>\[ = {@ta12@} \]</p>
    </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>a1: rand(2)+7;
a2:a1-3;
ta1: a1*a2 + (a1 +a2)*x + x^2;

b1: rand(2)+2;
b2:b1+1;
ta2: x^2 + (b1+b2)*x + b1*b2;
ta2f: x^2 + (b1*x) + b2*x + b1*b2;

c1: rand(3)+3;
c2:c1;
ta3: x^2 + (c1+c2)*x + c1*c2;
ta3f:x^2 + c1*x + c2*x + c1*c2;

d1: rand(2)+6;
d2:d1-2;
ta4: x^2 + (d1 - d2)*x - d1*d2;
ta4f:x^2 + d1*x -d2*x -d1*d2;

e1: rand(2)+2;
ta5: x^2 + e1*x;

f1:e1+1;
ta6:f1*x^2 + x; 

g1:f1;
ta7:g1*x^2 + (g1 + 1)*x + 1;
ta7f:g1*x^2 + g1*x + x + 1;

h1:g1;
h2:e1;
ta8:h1*h2*x^2 + (h1+h2)*x + 1;
ta8f:h1*h2*x^2 + h1*x +h2*x + 1;

i1:rand(2)+3;
i2:i1+2;
i3:i1-1;
i4:i2+1;
ta9:i1*i3*x^2 + (i1*i4 + i2*i3)*x + i2*i4;

j1:i1; 
j2:i2;
 j3:i3; 
j4:1;
ta10:j1*j3*x^2 + (-j1*j4 + j2*j3)*x - j2*j4;

k1:j2;
k2:j1;
k3:-k2;
ta11: (k1 - k2)*x + k1 + k3*x^2;

l1:i3;
ta12:x^2  -(l1+1)*x + l1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@ta1@},{@ta2@},{@ta3@}, {@ta4@}, {@ta10@},{@ta11@}, {@ta12@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>ta10</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>ta11</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>algebraic</type>
      <tans>ta12</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>ta9</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>(, )</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p>\[ ({@a1@}+x) ({@a2@}+x) =&nbsp;({@a1@}+x) ({@a2@}) + ({@a1@}+x)(x) \]</p><p>\[ = ({@a1@})({@a2@}) + (x)({@a2@}) + ({@a1@})(x) + (x)(x) \]</p><p>\[ = {@a1*a2@} + {@a1+a2@}x + x^2. \]&nbsp;</p><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\[ ({@a1@}+x) ({@a2@}+x) =&nbsp;({@a1@}+x) ({@a2@}) + ({@a1@}+x)(x) \]</p>
<p>\[ = ({@a1@})({@a2@}) + (x)({@a2@}) + ({@a1@})(x) + (x)(x) \]</p>
<p>\[ = {@a1*a2@} + {@a1+a2@}x + x^2. \]&nbsp;</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans10</sans>
        <tans>ta10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>ta10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p></p><p>\[ ({@j1@}x+ {@j2@}) ( {@j3@}x - {@j4@} )= ({@j1@}x+ {@j2@})({@j3@}x) + ({@j1@}x+ {@j2@})(-{@j4@}) \]</p><p>\[ = {@j1@}x({@j3@}x)+ ({@j2@})({@j3@}x) - {@j1@}x({@j4@}) - {@j2@}({@j4@}) \]</p><p>\[ = {@ta10@} \]</p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ ({@j1@}x+ {@j2@}) ( {@j3@}x - {@j4@} )= ({@j1@}x+ {@j2@})({@j3@}x) + ({@j1@}x+ {@j2@})(-{@j4@}) \]</p><p>\[ = {@j1@}x({@j3@}x)+ ({@j2@})({@j3@}x) - {@j1@}x({@j4@}) - {@j2@}({@j4@}) \]</p><p>\[ = {@ta10@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans11</sans>
        <tans>ta11</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>ta11</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p>\[ ( {@k1@} - {@k2@}x)(x + 1)= ( {@k1@} - {@k2@}x)(x) + ( {@k1@} - {@k2@}x)(1) \]</p><p>\[ = {@k1@}(x) - {@k2@}(x) + {@k1@}(1) - {@k2@}(1) \]</p><p>\[ = {@ta11@} \]</p><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ ( {@k1@} - {@k2@}x)(x + 1)= ( {@k1@} - {@k2@}x)(x) + ( {@k1@} - {@k2@}x)(1) \]</p><p>\[ = {@k1@}(x) - {@k2@}(x) + {@k1@}(1) - {@k2@}(1) \]</p><p>\[ = {@ta11@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans12</sans>
        <tans>ta12</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>ta12</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p>\[ ( {@l1@} - x)(1 - x)= ( {@l1@} - x)(1) + ( {@l1@} - x)(-x) \]</p><p>\[ = {@l1@}(1) - x(1) + {@l1@}(-x) - x(-x) \]</p><p>\[ = {@ta12@} \]</p><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ ( {@l1@} - x)(1 - x)= ( {@l1@} - x)(1) + ( {@l1@} - x)(-x) \]</p><p>\[ = {@l1@}(1) - x(1) + {@l1@}(-x) - x(-x) \]</p><p>\[ = {@ta12@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p><br></p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have removed the brackets correctly, but you have not simplified your answer fully.</p><p></p><p>\[ (x + {@b1@}) (x + {@b2@})= (x + {@b1@}) (x) +&nbsp;(x + {@b1@})&nbsp;({@b2@}) \]</p><p>\[ = (x)(x) + ({@b1@})(x) + ({@b2@})(x) + ({@b1@})({@b2@}) \]</p><p>\[ = x^2 + {@b1 + b2@}x + {@b1*b2@}. \]&nbsp;</p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ (x + {@b1@}) (x + {@b2@})= (x + {@b1@}) (x) +&nbsp;(x + {@b1@})&nbsp;({@b2@}) \]</p><p>\[ = (x)(x) + ({@b1@})(x) + ({@b2@})(x) + ({@b1@})({@b2@}) \]</p><p>\[ = x^2 + {@b1 + b2@}x + {@b1*b2@}. \]&nbsp;</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p></p><p>\[ (x + {@c1@}) (x + {@c2@})= (x + {@c1@}) (x) + (x + {@c1@})({@c2@}) \]</p><p>\[ = (x)(x) + ({@c1@})(x) + ({@c2@})(x) + ({@c1@})({@c2@}) \]</p><p>\[ = x^2 + {@c1 + c2@}x + {@c1*c2@}. \]</p><p></p><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\[ (x + {@c1@}) (x + {@c2@})= (x + {@c1@}) (x) + (x + {@c1@})({@c2@}) \]<br></p><p></p><p>\[ = (x)(x) + ({@c1@})(x) + ({@c2@})(x) + ({@c1@})({@c2@}) \]</p><p>\[ = x^2 + {@c1 + c2@}x + {@c1*c2@}. \]</p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p><br></p><br><p></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p></p><p>\[ (x + {@d1@}) (x - {@d2@})= (x + {@d1@}) (x) - (x + {@d1@})({@d2@}) \]</p><p>\[ = (x)(x) + ({@d1@})(x) - ({@d2@})(x) - ({@d1@})({@d2@}) \]</p><p>\[ = x^2 + {@d1 - d2@}x - {@d1*d2@}. \]</p><br><p></p><p></p><p><br></p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ (x + {@d1@}) (x - {@d2@})= (x + {@d1@}) (x) - (x + {@d1@})({@d2@}) \]</p><p>\[ = (x)(x) + ({@d1@})(x) - ({@d2@})(x) - ({@d1@})({@d2@}) \]</p><p>\[ = x^2 + {@d1 - d2@}x - {@d1*d2@}. \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>In the expression \( ({@e1@}+x)x \), the \( {@x@} \) outside the bracket must multiply both the \({@e1@}\) and the \(x\) in the brackets. Thus</p><p>\[ ({@e1@}+x)x = {@e1@}x + x \times x \]</p><p>\[ =&nbsp;{@e1@}x + x^2 \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>In the expression \( ({@f1@}x + 1)x \), the \( {@x@} \) outside the bracket must multiply both the \({@f1@}x\) and the \(1\) in the brackets. Thus</p><p>\[ ({@f1@}x + 1)x = {@f1@}x \times x + 1 \times x \]</p><p>\[ {@ta6@} \]</p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p></p><p>\[ ( {@g1@}x + 1)(x + 1) = ( {@g1@}x + 1)(x) + ( {@g1@}x + 1)(1) \]</p><p>\[ = {@g1@}x(x)+ (1)(x) + {@g1@}x(1) +1(1) \]</p><p>\[ = {@ta7@} \]</p><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ ( {@g1@}x + 1)(x + 1)&nbsp;= ( {@g1@}x + 1)(x) + ( {@g1@}x + 1)(1) \]</p><p>\[ = {@g1@}x(x)+ (1)(x) + {@g1@}x(1) +1(1) \]</p><p>\[ = {@ta7@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p></p><p>\[ ( {@h1@}x + 1)( {@h2@}x + 1)= ( {@h1@}x + 1)({@h2@}x) + ( {@h1@}x + 1)(1) \]</p><p>\[ = {@h1@}x({@h2@}x)+ (1)({@h2@}x) + {@h1@}x(1) +1(1) \]</p><p>\[ = {@ta8@} \]</p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ ( {@h1@}x + 1)( {@h2@}x + 1)= ( {@h1@}x + 1)({@h2@}x) + ( {@h1@}x + 1)(1) \]</p><p>\[ = {@h1@}x({@h2@}x)+ (1)({@h2@}x) + {@h1@}x(1) +1(1) \]</p><p>\[ = {@ta8@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p></p><p>You have removed the brackets correctly, but you have not simplified your answer fully.<br></p><p></p><p>\[ ({@i1@}x+ {@i2@}) ( {@i3@}x + {@i4@} )= ({@i1@}x+ {@i2@})({@i3@}x) + ({@i1@}x+ {@i2@})({@i4@}) \]</p><p>\[ = {@i1@}x({@i3@}x)+ ({@i2@})({@h3@}x) + {@i1@}x({@i4@}) +{@i2@}({@i4@}) \]</p><p>\[ = {@ta9@} \]</p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ ({@i1@}x+ {@i2@}) ( {@i3@}x + {@i4@} )= ({@i1@}x+ {@i2@})({@i3@}x) + ({@i1@}x+ {@i2@})({@i4@}) \]</p><p>\[ = {@i1@}x({@i3@}x)+ ({@i2@})({@h3@}x) + {@i1@}x({@i4@}) +{@i2@}({@i4@}) \]</p><p>\[ = {@ta9@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>662068008</deployedseed>
    <deployedseed>827903215</deployedseed>
    <deployedseed>276192335</deployedseed>
    <deployedseed>1490967358</deployedseed>
    <deployedseed>966840595</deployedseed>
    <deployedseed>125434971</deployedseed>
    <deployedseed>1968577152</deployedseed>
    <deployedseed>1325500931</deployedseed>
    <deployedseed>307413387</deployedseed>
    <deployedseed>375336124</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ta10</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ta11</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>ta12</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta9</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>1111</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ta12</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta8</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta9</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta10</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt10-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt11-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt12-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt7-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt8-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt9-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
