<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.3.3.4 Removing 3 sets of brackets</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Remove the brackets from each of the following expression and simplify your answer:<br></p>

<p>\( (s + {@a1@})(s + {@a2@})(s - {@a3@})=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\[ (s + {@a1@})(s + {@a2@})(s - {@a3@})=\]</p>
<p>\[(s + {@a1@})(s + {@a2@})(s) + (s + {@a1@})(s + {@a2@})(- {@a3@})= \]</p>
<p>\[(s + {@a1@})(s(s) + {@a2@}(s)) + (s + {@a1@})(s(- {@a3@}) + {@a2@}(- {@a3@}))= \]</p>
<p>\[(s + {@a1@})(s^2) + (s + {@a1@})({@a2@}s) + (s + {@a1@})(- {@a3@}s) + (s + {@a1@}) ({@a2*(- a3)@})= \]</p>
<p>\[ s(s^2) + {@a1@}(s^2) + s({@a2@}s) + {@a1@}({@a2@}s) + s(- {@a3@}s) + {@a1@}(- {@a3@}s)+ s({@a2*(- a3)@}) + {@a1@} ({@a2*(- a3)@}) \]</p>
<p>\[ s^3 + {@a1@}s^2 + {@a2@}s^2 + {@a1*a2@}s - {@a3@}s^2 - {@a1*a3@}s - {@a2*a3@}s - {@a1*a2* a3@} = \]</p>
<p>\[ {@ta1@} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021072200</text>
    </stackversion>
    <questionvariables>
      <text>a1: rand(2)+3;
a2: a1+5;
a3:a2-(rand(3)+2);
ta1: s^3 + (a1 + a2 - a3)*s^2 + (a1*a2 -a1*a3 - a2*a3)*s - a1*a2*a3</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\((s + {@a1@})(s + {@a2@})(s - {@a3@})={@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>12</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><div title="Page 2">
			<div>
				<div>
					<div>
						<p></p></div></div></div></div><p><br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have removed the brackets correctly, but have not simplified your answer.</p><p></p><p>\[ (s + {@a1@})(s + {@a2@})(s - {@a3@}) \]<br></p><p></p><p>\[=&nbsp; (s + {@a1@})(s + {@a2@})(s)&nbsp; + (s + {@a1@})(s + {@a2@})(- {@a3@}) \]<br></p><p>\[ =&nbsp; (s + {@a1@})(s(s) + {@a2@}(s))&nbsp; + (s + {@a1@})(s(- {@a3@}) + {@a2@}(- {@a3@})) \]<br></p><p>\[ = (s + {@a1@})(s^2) + (s + {@a1@})({@a2@}s)&nbsp; + (s + {@a1@})(- {@a3@}s)&nbsp;+ (s + {@a1@}) ({@a2*(- a3)@}) \]<br></p><p>\[ = s(s^2) + {@a1@}(s^2) + s({@a2@}s)&nbsp; + {@a1@}({@a2@}s)&nbsp; + s(- {@a3@}s) + {@a1@}(- {@a3@}s)&nbsp;+ s({@a2*(- a3)@}) + {@a1@} ({@a2*(- a3)@}) \]<br></p><p>\[=&nbsp; s^3 + {@a1@}s^2 + {@a2@}s^2&nbsp; + {@a1*a2@}s&nbsp; - {@a3@}s^2 - {@a1*a3@}s&nbsp; - {@a2*a3@}s - {@a1*a2* a3@} \]<br></p><p>\[ = {@ta1@} \]</p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>\[ (s + {@a1@})(s + {@a2@})(s - {@a3@}) \]<br></p><p></p><p>\[=&nbsp; (s + {@a1@})(s + {@a2@})(s)&nbsp; + (s + {@a1@})(s + {@a2@})(- {@a3@}) \]<br></p><p>\[ =&nbsp; (s + {@a1@})(s(s) + {@a2@}(s))&nbsp; + (s + {@a1@})(s(- {@a3@}) + {@a2@}(- {@a3@})) \]<br></p><p>\[ = (s + {@a1@})(s^2) + (s + {@a1@})({@a2@}s)&nbsp; + (s + {@a1@})(- {@a3@}s)&nbsp;+ (s + {@a1@}) ({@a2*(- a3)@}) \]<br></p><p>\[ = s(s^2) + {@a1@}(s^2) + s({@a2@}s)&nbsp; + {@a1@}({@a2@}s)&nbsp; + s(- {@a3@}s) + {@a1@}(- {@a3@}s)&nbsp;+ s({@a2*(- a3)@}) + {@a1@} ({@a2*(- a3)@}) \]<br></p><p>\[=&nbsp; s^3 + {@a1@}s^2 + {@a2@}s^2&nbsp; + {@a1*a2@}s&nbsp; - {@a3@}s^2 - {@a1*a3@}s&nbsp; - {@a2*a3@}s - {@a1*a2* a3@} \]<br></p><p>\[ = {@ta1@} \]</p><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>221081226</deployedseed>
    <deployedseed>1344528113</deployedseed>
    <deployedseed>282205874</deployedseed>
    <deployedseed>794116327</deployedseed>
    <deployedseed>1618988710</deployedseed>
    <deployedseed>1585776196</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>555</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>s^3 + a1*s^2 + a2*s^2  + a1*a2*s  - a3*s^2 - a1*a3*s  - a2*a3*s - a1*a2* a3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
