<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>23.2.6.2 fs ex2 (new question version)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise"><span style="font-size: 14px;">Exercise</span><br>
</p>
<p> Write out the terms of the Fourier series for the periodic signal given by \[ f(t) = t^2,\]
 \[ f(t + 2) = f(t).\] You should give the constant term \(\dfrac{a_0}{2}\) and each term corresponding to \( n = 1\) to \( n = 6\):<br></p> 
<p>[[input:ans1]][[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[Since \(T = 2\) the form of the Fourier series is
\[ \dfrac{a_0}{2} + \sum_{n=1}^\infty \left\{ a_n \cos (n\pi t) + b_n \sin (n\pi t) \right\} \]
where
\[
\begin{align}
a_n &amp;= \displaystyle{ \int _{-1}^{1} }f(t) \cos (n\pi t) \hspace{0.15cm}dt \hspace{1cm} n = 0, 1, 2, \dots \\
b_n &amp;=  \displaystyle{ \int _{-1}^{1} } f(t)  \sin (n\pi t) \hspace{0.15cm}dt  \hspace{1cm} n =  1, 2, \dots
\end{align}
\]
Because the function is non-negative there will be a constant term. 
The constant term will be \(\dfrac{a_0}{2}\) where \( a_0 = \displaystyle{ \int_{-1}^{1}} t^2 \hspace{0.15cm}dt  = \dfrac{2}{3} \).
<p></p>
For \(n = 1, 2, 3, \dots , \) we have
\[
\begin{align}
a_n &amp;= \displaystyle{ \int_{-1}^{1} } t^2 \cos (n\pi t) \hspace{0.15cm} dt \\
&amp;= \left[ \dfrac{t^2}{n\pi} \sin (n\pi t) \right] _{-1}^{1} - \dfrac{2}{n\pi} \displaystyle{ \int_{-1}^{1} }  t \sin n\pi t \hspace{0.15cm} dt \\
&amp;= 0 - \dfrac{2}{n\pi} \left\{ \left[ - \dfrac{t}{n\pi} \cos n \pi t \right] _{-1}^1 + \dfrac{1}{n\pi} \displaystyle{ \int_{-1}^{1} }   \cos n\pi t \hspace{0.15cm} dt \right\} \\
&amp;= -\dfrac{2}{n\pi} \left\{ -\dfrac{1}{n\pi} \cos n\pi - \dfrac{1}{n\pi} \cos (-n\pi) + \dfrac{1}{n^2\pi^2} \left[ \sin n\pi t \right]_{-1}^1 \right\} \\
&amp;= \dfrac{4}{n^2\pi^2} \cos n\pi
\end{align}
\]
As the function is even, we have \(b_n = 0\). 
<p></p>
Thus the series is given by
\[ 
\begin{align} f(t) &amp;= \dfrac{1}{3} - \dfrac{4}{\pi^2} \cos \pi t + \dfrac{1}{\pi^2} \cos 2\pi t - \dfrac{4}{9\pi^2} \cos 3\pi t  \\
&amp;+ \dfrac{1}{4\pi^2} \cos 4\pi t - \dfrac{4}{25\pi^2} \cos 5\pi t + \dfrac{1}{9\pi^2} \cos 6\pi t + \dots
\end{align}
\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>ta1: 1/3 - 4/(%pi^2)*cos(pi*t) + 1/(%pi^2)*cos(2*%pi*t) - 4/(9*%pi^2)*cos(3*pi*t) + 1/(4*%pi^2)*cos(4*%pi*t) - 4/(25*%pi^2)*cos(5*%pi*t) + 1/(9*%pi^2)*cos(6*%pi*t);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>100</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
