<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>8.0.24 (Rationalise hence find limit)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><b>(a)</b> Write \(\dfrac{{@numer2@}}{{@a*n+b@}}\) as the sum of a number and a proper rational function.&nbsp;</p><p>\(\dfrac{{@numer2@}}{{@a*n+b@}} = \) [[input:ans2]][[validation:ans2]][[feedback:prt2]]</p><p><b>(b)</b> Hence find \(\displaystyle\lim_{n\to\infty} \dfrac{{@numer2@}}{{@a*n+b@}} = \)&nbsp;[[input:ans3]][[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>(a)</b> \(\displaystyle\dfrac{{@numer2@}}{{@a*n+b@}} = \dfrac{{@(a*n+b)*ev(1/a,simp)+k-b/a@}}{({@a*n+b@})} = \dfrac{{@(a*n+b)*ev(1/a,simp)@}}{({@a*n+b@})}+\dfrac{{@k-b/a@}}{({@a*n+b@})} ={@ta1@} \).</p><p><b>(b) </b>Since the term \(\dfrac{{@ev(k-b/a,simp)@}}{({@a*n+b@})}\to0\) as \(n\to\infty\), we have \(\displaystyle\lim_{n\to\infty} \dfrac{{@numer2@}}{{@a*n+b@}} = {@ta2@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a:rand([2,3,4,5,6]);
b:rand([-4,-3,-2,-1,1,2,3,4]);
k:rand([0,1,2,3,4]);
numer2:ev(expand((a*n+b)*(1/a)-b/a+k),simp);
ta1: unary_minus_sort(ev((1/a),simp)+ev((-b/a+k),simp)/(a*n+b));
ta2: limit(ta1,n,inf);
polypart: 1/a;
fracpart:(-b/a+k)/(a*n+b);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@(numer2)/(a*n+b)=ta1@} \(\lim_{n\to\infty} \frac{n+k}{an+b} = {@ta2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>polynomialp(ans2-fracpart,[x])</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your final answer does not seem to be in the correct form.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1229562750</deployedseed>
    <deployedseed>1417866914</deployedseed>
    <deployedseed>651442217</deployedseed>
    <deployedseed>30449679</deployedseed>
    <deployedseed>176637699</deployedseed>
    <deployedseed>100106792</deployedseed>
    <deployedseed>1119768826</deployedseed>
    <deployedseed>398037143</deployedseed>
    <deployedseed>1198917828</deployedseed>
    <deployedseed>2059267948</deployedseed>
    <deployedseed>1811148284</deployedseed>
    <deployedseed>739264713</deployedseed>
    <deployedseed>513569026</deployedseed>
    <deployedseed>1658518687</deployedseed>
    <deployedseed>289024142</deployedseed>
    <deployedseed>132013495</deployedseed>
    <deployedseed>1691388353</deployedseed>
    <deployedseed>822624930</deployedseed>
    <deployedseed>1691311196</deployedseed>
    <deployedseed>310966352</deployedseed>
    <deployedseed>1144287382</deployedseed>
    <deployedseed>458088478</deployedseed>
    <deployedseed>1441793925</deployedseed>
    <deployedseed>1381393027</deployedseed>
    <deployedseed>1362893146</deployedseed>
    <deployedseed>187558341</deployedseed>
    <deployedseed>1580968057</deployedseed>
    <deployedseed>914238582</deployedseed>
    <deployedseed>1913519633</deployedseed>
    <deployedseed>919466760</deployedseed>
    <deployedseed>326948558</deployedseed>
    <deployedseed>698195705</deployedseed>
    <deployedseed>1464296686</deployedseed>
    <deployedseed>232822956</deployedseed>
    <deployedseed>1196891200</deployedseed>
    <deployedseed>1744867117</deployedseed>
    <deployedseed>1455409247</deployedseed>
    <deployedseed>1109899905</deployedseed>
    <deployedseed>1190875757</deployedseed>
    <deployedseed>467181327</deployedseed>
    <deployedseed>503339668</deployedseed>
    <deployedseed>1147393926</deployedseed>
    <deployedseed>1972290049</deployedseed>
    <deployedseed>1547915843</deployedseed>
    <deployedseed>1522537918</deployedseed>
    <deployedseed>257231122</deployedseed>
    <deployedseed>721081913</deployedseed>
    <deployedseed>2043931820</deployedseed>
    <deployedseed>1844301479</deployedseed>
    <deployedseed>1594908715</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta2+1</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
