<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>8.0.21 (Which are periodic? 2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="noindent">Which of the following sequences are periodic?</p>  
 [[input:ans1]] [[validation:ans1]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p class="noindent">To see if the sequence is periodic, we can evaluate the first few terms and see if a pattern seems to emerge.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>corbase:[u_n=sin(n*%pi),u_n=(-1)^n,u_n=1];
wrongbase:[u_n=n!,u_n=n,u_n=sin(1/n)];
mcq:multiselqnalpha(corbase, 2, wrongbase, 2);
ta1:mcq[1];
version:mcq[2];
corr1:mcq_correct(ta1);
incorr1:mcq_incorrect(ta1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{#version#}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>setify(ans1)</sans>
        <tans>setify(corr1)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>708129464</deployedseed>
    <deployedseed>1235259685</deployedseed>
    <deployedseed>2104290618</deployedseed>
    <deployedseed>390916773</deployedseed>
    <deployedseed>1119679033</deployedseed>
    <deployedseed>1891471945</deployedseed>
    <deployedseed>1792135128</deployedseed>
    <deployedseed>1940749819</deployedseed>
    <deployedseed>346802060</deployedseed>
    <deployedseed>1702964589</deployedseed>
    <deployedseed>1431526693</deployedseed>
    <deployedseed>647169256</deployedseed>
    <deployedseed>1192081214</deployedseed>
    <deployedseed>2055424846</deployedseed>
    <deployedseed>851327548</deployedseed>
    <deployedseed>197393232</deployedseed>
    <deployedseed>569519126</deployedseed>
    <deployedseed>2019322695</deployedseed>
    <deployedseed>1033638056</deployedseed>
    <deployedseed>1240042866</deployedseed>
    <deployedseed>16928144</deployedseed>
    <deployedseed>1121994948</deployedseed>
    <deployedseed>1471904041</deployedseed>
    <deployedseed>1520315621</deployedseed>
    <deployedseed>1853795239</deployedseed>
    <deployedseed>83529273</deployedseed>
    <deployedseed>1885651257</deployedseed>
    <deployedseed>28853225</deployedseed>
    <deployedseed>252068384</deployedseed>
    <deployedseed>1761233248</deployedseed>
    <deployedseed>97493269</deployedseed>
    <deployedseed>324856001</deployedseed>
    <deployedseed>2094253269</deployedseed>
    <deployedseed>819092591</deployedseed>
    <deployedseed>967060602</deployedseed>
    <deployedseed>1197505426</deployedseed>
    <deployedseed>684944281</deployedseed>
    <deployedseed>1215899</deployedseed>
    <deployedseed>565505791</deployedseed>
    <deployedseed>1140442066</deployedseed>
    <deployedseed>1878192904</deployedseed>
    <deployedseed>683060789</deployedseed>
    <deployedseed>1690434191</deployedseed>
    <deployedseed>1852267298</deployedseed>
    <deployedseed>1882008189</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>corr1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>incorr1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>sort([incorr1[1],corr1[1]])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>[corr1[1]]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
