<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>8.1.4 (Arithmetic seq find missing value)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The following is part of an arithmetic sequence \[ \ldots, {@a@}, {@s2@}, x, {@s4@}, {@s5@}, \ldots \]</p><p>What is the value of \(x\)?&nbsp;</p><p>\( x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Looking at the first two (or the last two) given terms we see the common difference is <span style="white-space: nowrap;">{@d@}</span>. The missing term is the second given term plus the difference so \( x = {@ta1@} \).<span style="white-space: nowrap;"></span><br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a: rand_with_step(-19,19,2);
d: rand(5)+4;
s2: a+d;
s4: a+3*d;
s5: a+4*d;
ta1: a+2*d;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>common difference \(d = {@d@}\) first term \(a = {@a@} \) so \( x = {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1537080077</deployedseed>
    <deployedseed>77754188</deployedseed>
    <deployedseed>1335732638</deployedseed>
    <deployedseed>580195529</deployedseed>
    <deployedseed>871109030</deployedseed>
    <deployedseed>500257851</deployedseed>
    <deployedseed>1598056607</deployedseed>
    <deployedseed>1484340056</deployedseed>
    <deployedseed>1022730207</deployedseed>
    <deployedseed>75765274</deployedseed>
    <deployedseed>589537778</deployedseed>
    <deployedseed>1337249975</deployedseed>
    <deployedseed>1574085575</deployedseed>
    <deployedseed>239606137</deployedseed>
    <deployedseed>214376092</deployedseed>
    <deployedseed>842189128</deployedseed>
    <deployedseed>1655511614</deployedseed>
    <deployedseed>1059660327</deployedseed>
    <deployedseed>804889709</deployedseed>
    <deployedseed>2051982154</deployedseed>
    <deployedseed>414908753</deployedseed>
    <deployedseed>1261740434</deployedseed>
    <deployedseed>1110448004</deployedseed>
    <deployedseed>1283422151</deployedseed>
    <deployedseed>386690296</deployedseed>
    <deployedseed>62787330</deployedseed>
    <deployedseed>769996745</deployedseed>
    <deployedseed>2031895994</deployedseed>
    <deployedseed>1152316985</deployedseed>
    <deployedseed>899596542</deployedseed>
    <deployedseed>1908992542</deployedseed>
    <deployedseed>547905272</deployedseed>
    <deployedseed>13969879</deployedseed>
    <deployedseed>2023902735</deployedseed>
    <deployedseed>1527142264</deployedseed>
    <deployedseed>210690428</deployedseed>
    <deployedseed>1584558427</deployedseed>
    <deployedseed>1933833629</deployedseed>
    <deployedseed>247618042</deployedseed>
    <deployedseed>282305211</deployedseed>
    <deployedseed>1943046625</deployedseed>
    <deployedseed>446337169</deployedseed>
    <deployedseed>1037976489</deployedseed>
    <deployedseed>1269698731</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
