<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>8.4.7 (Fibonacci rabbits)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Beginning with a single pair of rabbits, how many pairs of rabbits will there be after {@nn@} months,  assuming that every month each male and female rabbit gives birth to a new pair of rabbits, and the new pair of rabbits itself starts giving birth to additional pairs of rabbits after the first month of their birth?</p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Recall that in this situation, rabbits multiply according to the Fibonacci sequence. So to find the number of rabbits after {@nn@} months we need to find the {@nn@}th term of the Fibonacci sequence. This gives us {@ta1@} pairs of rabbits.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>nn: 1+rand(50);
gg1: (1+sqrt(5))/2; 
gg2:1-gg1;


ta1: simplify(1/sqrt(5)*(gg1^nn - gg2^nn));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@nn@} months, {@ta1@} pairs of rabbits</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1582539321</deployedseed>
    <deployedseed>1547294814</deployedseed>
    <deployedseed>578088287</deployedseed>
    <deployedseed>132304831</deployedseed>
    <deployedseed>1183885492</deployedseed>
    <deployedseed>533184948</deployedseed>
    <deployedseed>229894641</deployedseed>
    <deployedseed>92355225</deployedseed>
    <deployedseed>91582639</deployedseed>
    <deployedseed>510686878</deployedseed>
    <deployedseed>1181578054</deployedseed>
    <deployedseed>1063434349</deployedseed>
    <deployedseed>1018401716</deployedseed>
    <deployedseed>1661722065</deployedseed>
    <deployedseed>1112183598</deployedseed>
    <deployedseed>482405014</deployedseed>
    <deployedseed>808172389</deployedseed>
    <deployedseed>267042904</deployedseed>
    <deployedseed>1081422577</deployedseed>
    <deployedseed>1568960306</deployedseed>
    <deployedseed>1231058881</deployedseed>
    <deployedseed>1976214989</deployedseed>
    <deployedseed>498956305</deployedseed>
    <deployedseed>2010418364</deployedseed>
    <deployedseed>300397366</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
