<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.8.6 (Sum to infinity 2016-LEAPS-W6-GeoSer-2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the sum to infinity of the geometric series \( {@expr@} + \cdots \).</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The first term is \(a={@a@}\) and the common ratio is \(r={@r@}\).</p>
<p>The sum to infinity is given by \(\dfrac{a}{1-r}=\dfrac{{@a@}}{{@(1-r),simp=false@}}={@(simp:true,ta)@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022082900</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[simp:true;
r:rand([1,-1])*rand([1/5, 1/4, 1/3, 2/5, 1/2, 3/5, 2/3, 3/4, 4/5]);
a:rand([2,3,4,5,6]);
ta:a/(1-r);
terms:makelist(a*r^(k-1),k,4);
simp:false;
expr:unary_minus_sort(apply("+",terms));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@expr@}+\cdots={@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>721180385</deployedseed>
    <deployedseed>1541343831</deployedseed>
    <deployedseed>471918485</deployedseed>
    <deployedseed>2048509353</deployedseed>
    <deployedseed>595926307</deployedseed>
    <deployedseed>590729697</deployedseed>
    <deployedseed>750003481</deployedseed>
    <deployedseed>1643818779</deployedseed>
    <deployedseed>1310777811</deployedseed>
    <deployedseed>2052228670</deployedseed>
    <deployedseed>2021476684</deployedseed>
    <deployedseed>82325715</deployedseed>
    <deployedseed>1035496054</deployedseed>
    <deployedseed>1869163634</deployedseed>
    <deployedseed>905378591</deployedseed>
    <deployedseed>1415724850</deployedseed>
    <deployedseed>1984838038</deployedseed>
    <deployedseed>1772931196</deployedseed>
    <deployedseed>112548837</deployedseed>
    <deployedseed>1295886505</deployedseed>
    <deployedseed>1919931123</deployedseed>
    <deployedseed>922424685</deployedseed>
    <deployedseed>59683826</deployedseed>
    <deployedseed>1322344622</deployedseed>
    <deployedseed>522067109</deployedseed>
    <deployedseed>1322673641</deployedseed>
    <deployedseed>117171802</deployedseed>
    <deployedseed>699577621</deployedseed>
    <deployedseed>2040731617</deployedseed>
    <deployedseed>1146798390</deployedseed>
    <deployedseed>1141339317</deployedseed>
    <deployedseed>2071421969</deployedseed>
    <deployedseed>1014906685</deployedseed>
    <deployedseed>694398833</deployedseed>
    <deployedseed>1683738715</deployedseed>
    <deployedseed>595481926</deployedseed>
    <deployedseed>553315954</deployedseed>
    <deployedseed>973789796</deployedseed>
    <deployedseed>914765833</deployedseed>
    <deployedseed>1917467033</deployedseed>
    <deployedseed>175852649</deployedseed>
    <deployedseed>396446574</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
