<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.8.4 (Find recurrence relation)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A sequence is defined by \(u_{n+1} =au_n +b\) with \(u_1 = {@u1@}\), \(u_2 = {@u2@}\) and \(u_3 = {@u3@}\) Find the values of \(a\) and \(b\).
<br/>
<p>a = [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b = [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>To solve this problem we form two simultaneous equations using the values of the sequence that we are given
\[
\begin{align*}
&{@eqn1@}
\\
&{@eqn2@}
\end{align*}
\]
We then solve these equations simultaneously using your preferred method to get \(a = {@ta1@}\) and \(b = {@ta2@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>u1:1+ rand(10);
u2:1+ rand(10);
u3: 1+ rand(20);

eqn1: u2 = u1*a + b;
eqn2: u3 = u2*a + b;

aa: solve([eqn1,eqn2],[a,b])[1];
bb: solve([eqn1,eqn2],[a,b])[1];

ta1: map(rhs,aa)[1];
ta2: map(rhs,bb)[2];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eqn1@}, {@eqn2@}, {@ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1915900462</deployedseed>
    <deployedseed>2034390704</deployedseed>
    <deployedseed>1005999639</deployedseed>
    <deployedseed>1374107836</deployedseed>
    <deployedseed>24434470</deployedseed>
    <deployedseed>164670871</deployedseed>
    <deployedseed>193800271</deployedseed>
    <deployedseed>809374443</deployedseed>
    <deployedseed>527448138</deployedseed>
    <deployedseed>1925245108</deployedseed>
    <deployedseed>1989802540</deployedseed>
    <deployedseed>1394028471</deployedseed>
    <deployedseed>289536078</deployedseed>
    <deployedseed>1259867489</deployedseed>
    <deployedseed>95477378</deployedseed>
    <deployedseed>1881002467</deployedseed>
    <deployedseed>1504900401</deployedseed>
    <deployedseed>2091066996</deployedseed>
    <deployedseed>75067264</deployedseed>
    <deployedseed>1908345712</deployedseed>
    <deployedseed>1153416558</deployedseed>
    <deployedseed>1407844783</deployedseed>
    <deployedseed>748908269</deployedseed>
    <deployedseed>233736264</deployedseed>
    <deployedseed>610472075</deployedseed>
    <deployedseed>592006970</deployedseed>
    <deployedseed>2061980244</deployedseed>
    <deployedseed>1810579670</deployedseed>
    <deployedseed>1343827670</deployedseed>
    <deployedseed>476623071</deployedseed>
    <deployedseed>962495154</deployedseed>
    <deployedseed>397428043</deployedseed>
    <deployedseed>696246319</deployedseed>
    <deployedseed>450491933</deployedseed>
    <deployedseed>2130224444</deployedseed>
    <deployedseed>944533868</deployedseed>
    <deployedseed>1496186228</deployedseed>
    <deployedseed>1802701479</deployedseed>
    <deployedseed>1362954494</deployedseed>
    <deployedseed>802760567</deployedseed>
    <deployedseed>1704375368</deployedseed>
    <deployedseed>1766471057</deployedseed>
    <deployedseed>1513174332</deployedseed>
    <deployedseed>417055721</deployedseed>
    <deployedseed>1635397202</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
