<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>PT.8.9 (Infinite sum 1)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Let \( u_1, u_2, \ldots, u_n, \ldots \) be an arithmetic sequence and \( v_1, v_2, \ldots, v_n, \ldots \) be a geometric sequence. The first terms \( u_1\) and \( v_1\) are both equal to {@t1@}, and the third terms \( u_3\) and \( v_3\) are both equal to {@t3@}.</p><p><b>(a)</b> Find \(u_{@n@} = \) [[input:ans1]] [[validation:ans1]]</p><p><b>(b)</b> Given that \( v_1, v_2, \ldots \) is a sequence of positive numbers, calculate \( \displaystyle \sum_{n=1}^{\infty}v_n .\)</p><p>\( \displaystyle \sum_{n=1}^{\infty}v_n = \) [[input:ans2]] [[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>(a)</b> For the arithmetic sequence we have: \[ \begin{align*} u_n &amp;=u_1+(n-1)d&nbsp; \\ &amp;= {@t1+('n-1)*'d@}.\end{align*} \] Since \(u_3={@t3@}\), this gives \[ \begin{align*} {@t1+(3-1)*'d@} &amp;= {@t3@} \\ 2d&amp;= {@t3@}-{@t1@}&nbsp; \\ d&amp;= {@d@}.\end{align*} \] Hence \( u_{@n@} = {@t1@}+({@n@}-1)({@d@}) = {@ta1@}\).</p><p><b>(b)</b> For the geometric sequence we have \[ v_n= v_1 r^{n-1} \] So \[ {@t3@} = {@t1@} r^2 \] \[ r= {@r@} \] Hence \[ S = \frac{@t1@}{1-{@r@}}= {@ta2@} \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>k: rand([9,4])
r: sqrt(1/k)
t3: rand(5)+5
t1: (t3)*k
d: t1*(r^2-1)/2
n: rand([10,11,12])
ta1: t1+(n-1)*d
ta2: t1/(1-r)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]] [[feedback:prt2]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( r= {@r@} \quad d= {@d@} \quad u_1 = {@t1@} \) so \( u_{@n@} = {@ta1@} \) and \( \displaystyle \sum_{n=1}^{\infty}v_n = {@ta2@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>846298575</deployedseed>
    <deployedseed>1225943540</deployedseed>
    <deployedseed>803018299</deployedseed>
    <deployedseed>581008614</deployedseed>
    <deployedseed>83806261</deployedseed>
    <deployedseed>980272166</deployedseed>
    <deployedseed>1084739466</deployedseed>
    <deployedseed>1125495449</deployedseed>
    <deployedseed>743910215</deployedseed>
    <deployedseed>54183476</deployedseed>
    <deployedseed>1443977880</deployedseed>
    <deployedseed>1675477180</deployedseed>
    <deployedseed>1648679881</deployedseed>
    <deployedseed>785116764</deployedseed>
    <deployedseed>1599279703</deployedseed>
    <deployedseed>984858767</deployedseed>
    <deployedseed>1583671946</deployedseed>
    <deployedseed>375015021</deployedseed>
    <deployedseed>817451087</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
