<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Series_3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>By writing \(x\) as \({@n2@}+(x-{@n2@})\) and using the Maclaurin expansion for \(\ln(1+t)\), or otherwise, find the first four non zero terms in the Taylor series expansion for</p>
<p>\[\ln({@n1@}x)\]</p>
<p>about \(x = {@n2@}\).</p>
<p>Use exact expressions, i.e. enter \(\ln(5)\) rather than \(1.609...\).</p>
<p>Answer: [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\[\ln({@n1@}x)\]</p>
<p>can be written as</p>
<p>\[\ln({@n1@}({@n2@}+(x-{@n2@}))) \]</p>
<p>which equals \[ \ln{@n1@} + \ln({@n2@}+(x-{@n2@})) =\ln{@n1@} + \ln\left({@n2@}\left(1+\frac{(x-{@n2@})}{@n2@}\right)\right)=\ln{@n1@} + \ln{@n2@}+\ln\left(1+\frac{(x-{@n2@})}{@n2@}\right)\]</p>
<p>As \(\ln(1+t) = t-\frac{t^2}{2}+\frac{t^3}{3}\) and \(\ln{@n1@} + \ln{@n2@} =\ln({@n1*n2@})\)</p>
<p>\[\ln({@n1@}x)= \ln({@n1*n2@})+\frac{x-{@n2@}}{@n2@}-\frac{(x-{@n2@})^2}{@2*n2^2@}+\frac{(x-{@n2@})^3}{@3*n2^3@}\]</p>
<p> </p>
<p> </p>
<p> </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(4)+2;
 n2:rand(4)+2;
 f : ln(n1*x);  
 correctanswer:ln(n1*n2)+(x-n2)/n2-((x-n2)/n2)^2/2+((x-n2)/n2)^3/3;
 wronganswer1:correctanswer-ln(n1*n2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\ln({@n1@}x)\) about \(x = {@n2@}\) is \({@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff, taylor,sum,solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1436289473</deployedseed>
    <deployedseed>688554558</deployedseed>
    <deployedseed>1550467957</deployedseed>
    <deployedseed>1797200102</deployedseed>
    <deployedseed>570203821</deployedseed>
    <deployedseed>2083950902</deployedseed>
    <deployedseed>837199278</deployedseed>
    <deployedseed>783032314</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
