<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Series_2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Using the Binomial series, find the first four non zero terms in the Maclaurin series for</p>
<p> \[\sqrt{{@n1@}+{@n2@}x}\]</p>
<p>Answer: [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p> \[\sqrt{{@n1@}+{@n2@}x}\] can be written as \[{@sqrt(n1)@}\sqrt{\left(1+{@(n2/n1)*x@}\right)}.\]As \[ {(1+x)}^{1/2} = 1 +\frac{1}{2} x + \frac{(1/2)(1/2-1)}{2 !} x^2 +\frac{(1/2)(1/2 -1)(1/2 - 2)}{3 ! } x^3 + \ldots ,\]</p>
<p>\[ {(1+ {@n2 *x / n1@} )}^{1/2} = 1 +\frac{1}{2} {@n2 *x / n1@} + \frac{(1/2)(1/2-1)}{2 !} ({@n2 *x / n1@} )^2 +\frac{(1/2)(1/2 -1)(1/2 - 2)}{3 ! } ({@n2 *x / n1@} )^3 + \ldots .\]</p>
<p>Hence</p>
<p>\[\sqrt{{@n1@}+{@n2@}x} = {@sqrt(n1)@}(1 + {@n2/n1 * x / 2@} - {@(n2/n1*x)^2 / 8@} + {@(n2/n1*x)^3 / 16@} + \ldots) \]</p>
<p>or\[\sqrt{{@n1@}+{@n2@}x} = {@sqrt(n1)@} +{@sqrt(n1)*(n2/n1)/2@} x-{@sqrt(n1)*(n2/n1)^2/8@} x^2 + {@sqrt(n1)*(n2/n1)^3/16@} x^3 + \ldots \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(8)+2;
 n2:rand(8)+2;
 f : sqrt(n1+n2*x);
 correctanswer : expand(sqrt(n1)*(1 +n2/n1*x/2-(n2/n1*x)^2/8+(n2/n1*x)^3/16));
 wronganswer1:expand(correctanswer/sqrt(n1));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\sqrt{({@n1@}+{@n2@}x)}. = {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>40</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff,taylor,sum,solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>219104182</deployedseed>
    <deployedseed>1936299538</deployedseed>
    <deployedseed>1979211555</deployedseed>
    <deployedseed>737939164</deployedseed>
    <deployedseed>1559423826</deployedseed>
    <deployedseed>404399832</deployedseed>
    <deployedseed>882348324</deployedseed>
    <deployedseed>1358258656</deployedseed>
    <deployedseed>1771074974</deployedseed>
    <deployedseed>2040213504</deployedseed>
    <deployedseed>680019207</deployedseed>
    <deployedseed>2024221618</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
