<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>7.3.29 (Reduced row echelon form 2 by 3)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Write the matrix 
\[
{@mat1@}
\]
in reduced row echelon form.</p>

<p>[[input:ans1]][[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>In this case we perform row operations to get 
\[
 {@ta1@}
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>a1: 1+rand(10);
a2: 1+rand(10);
a3: 1+rand(10);
a4: 1+rand(10);
a5: 1+rand(10);
a6: 1+rand(10);

mat1: matrix([a1,a2,a3],[a4,a5,a6]);

ta1: rref(mat1)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@mat1@} \implies {@ta1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2124857511</deployedseed>
    <deployedseed>287228164</deployedseed>
    <deployedseed>1318254978</deployedseed>
    <deployedseed>1364597709</deployedseed>
    <deployedseed>1488250187</deployedseed>
    <deployedseed>893751378</deployedseed>
    <deployedseed>729232173</deployedseed>
    <deployedseed>1312881190</deployedseed>
    <deployedseed>1730593407</deployedseed>
    <deployedseed>53500684</deployedseed>
    <deployedseed>1336078135</deployedseed>
    <deployedseed>723605996</deployedseed>
    <deployedseed>49435596</deployedseed>
    <deployedseed>185481082</deployedseed>
    <deployedseed>457692060</deployedseed>
    <deployedseed>245482296</deployedseed>
    <deployedseed>460374165</deployedseed>
    <deployedseed>22249596</deployedseed>
    <deployedseed>849265734</deployedseed>
    <deployedseed>790521225</deployedseed>
    <deployedseed>693088290</deployedseed>
    <deployedseed>610182603</deployedseed>
    <deployedseed>559353338</deployedseed>
    <deployedseed>2049723852</deployedseed>
    <deployedseed>459861847</deployedseed>
    <deployedseed>792148248</deployedseed>
    <deployedseed>1527631215</deployedseed>
    <deployedseed>1263066543</deployedseed>
    <deployedseed>976886826</deployedseed>
    <deployedseed>1027292202</deployedseed>
    <deployedseed>1841339094</deployedseed>
    <deployedseed>1825317206</deployedseed>
    <deployedseed>1094388333</deployedseed>
    <deployedseed>69881763</deployedseed>
    <deployedseed>82819894</deployedseed>
    <deployedseed>1140978863</deployedseed>
    <deployedseed>1816336562</deployedseed>
    <deployedseed>1450327625</deployedseed>
    <deployedseed>1500434305</deployedseed>
    <deployedseed>364054914</deployedseed>
    <deployedseed>51768574</deployedseed>
    <deployedseed>2075313730</deployedseed>
    <deployedseed>220455392</deployedseed>
    <deployedseed>954047328</deployedseed>
    <deployedseed>854636186</deployedseed>
    <deployedseed>1696765990</deployedseed>
    <deployedseed>1818694345</deployedseed>
    <deployedseed>489511336</deployedseed>
    <deployedseed>1896204644</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
