<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>7.4.8 (Investment sim eqns. 1)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Suppose that Carmen invested £{@inv@}, part at {@r1@} and part at {@r2@}. Given that the total interest on the investment was £{@int@}. How much did she invest at each rate? Give your answer as a list of equations e.g. <code>[x=2,y=2]</code>.</p>

<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We can form simultaneous equations out of the total investment and the interest invested. Specifically if we let \(x\) be the amount invested at {@r1@} and \(y\) the amount invested at {@r2@}. Then the information from the question tells us that 
\[
\begin{align*}
&{@eqn1@}
\\
&{@eqn2@}
\end{align*}
\].
At this point it is probably easiest to make the second equation involve integers by multiplying it by \(100\) to get
\[
\begin{align*}
&{@eqn1@}
\\
&{@eqn2a@}
\end{align*}
\] 
We can then solve these equations simultaneously, to get our answers to be {@ta1@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>inv: rand([10000,20000,30000,40000,50000,60000,70000, 80000, 90000]);
int: 2000+rand([100,200,300,400,500,600,700,800,900,1000,1100,1200,1300,1400,1500]);
r1: 0.05+(1/100)*rand(5);
r2: 0.08+(1/100)*rand(20);

eqn1: x  + y = inv;
eqn2: r1*x +r2*y = int;
eqn2a: expand(100*eqn2); 

ta1: solve([eqn1,eqn2],[x,y])[1];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@eqn1@},{@eqn2a@} \(\implies\) {@ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[x=?,y=?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>setify(ans1)</sans>
        <tans>setify(ta1)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1673445458</deployedseed>
    <deployedseed>1418798248</deployedseed>
    <deployedseed>1424050949</deployedseed>
    <deployedseed>1563335966</deployedseed>
    <deployedseed>1846544381</deployedseed>
    <deployedseed>1657898663</deployedseed>
    <deployedseed>1869076517</deployedseed>
    <deployedseed>875300347</deployedseed>
    <deployedseed>1516689633</deployedseed>
    <deployedseed>549075569</deployedseed>
    <deployedseed>1977373410</deployedseed>
    <deployedseed>1522711336</deployedseed>
    <deployedseed>1266853670</deployedseed>
    <deployedseed>2004056343</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
