<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 17345  -->
  <question type="stack">
    <name>
      <text>Cubic spline</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<br>
<div class="stack-comment">
<h2>Specific feedback and multiple properties</h2>
<p>In this, more advanced, question the student is required to find a cubic polynomial which matches up with the given function \(f(x)\).  In particular it should <em>connect smoothly</em> with both given ends.</p>
<p>There are are <em>five separate properties</em> that characterize a correct answer</p>
<ol>
<li>\(p(x)\)is a cubic (the highest power of \(x\) is 3).</li>
<li>\(p(-2)=1\)(connects at the left)</li>
<li>\(p(0)=0\)(connects at the right)</li>
<li>\(p'(-2)=0\)(slope is correct at the left)</li>
<li>\(p'(0)=1\) (slope is correct at the right)</li>
</ol>
<p>In this case there is, up to algebraic equivalence, only one correct answer.  In fact, it is more helpful to students to check each property separately and give more specific feedback.  </p>
<p>Try <code>x^2*(x+3)/4</code>which is only partially correct.</p>
<p>(The correct answer is <code>x^3/2+7*x^2/4+x</code>.)</p>
</div>

<p>Consider the real function</p>
<p>\[ f(x) :=\left\{ \begin{array}{ll} {@f0@} &amp; \mbox{for } x\leq {@x0@}, \\ p(x) &amp; \mbox{for } {@x0@} &lt; x &lt; {@x1@}, \\ {@f1@} &amp; \mbox{for } {@x1@}\leq x.\end{array} \right. \]</p>
<p>{@plot([pg],[x,(x0-3),(x1+5)],[y,-3,3])@}</p>
<p>Find the cubic polynomial \(p(x)\) which makes \(f(x)\) continuously differentiable. </p>
<p>\(p(x)=\)[[input:ans1]]</p>
<div>[[validation:ans1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>First we need to find the information which \(p(x)\) needs to satisfy.  This is</p>
<p>\[ p({@x0@})={@subst(x=x0,f0)@}.\]</p>
<p>\[ p({@x1@})={@subst(x=x1,f1)@}.\]</p>
<p>\[ p'({@x0@})={@subst(x=x0,diff(f0,x))@}.\]</p>
<p>\[ p'({@x1@})={@subst(x=x1,diff(f1,x))@}.\]</p>
<p>If we define \(p(x)\) to be the cubic </p>
<p>\[ p(x) = a_3 x^3 + a_2 x^2 + a_1 x + a_0,\]</p>
<p>We then set up the matrix equation</p>
<p>\[ {@CS@} {@CV@} = {@CT@} \]</p>
<p>Solving this gives the coefficients, from which we get the polynomial.</p>
<p>{@ta@}.</p>
<p>We can view this as {@plot([pg,tap],[x,(x0-3),(x1+5)],[y,-3,3])@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025030700</text>
    </stackversion>
    <questionvariables>
      <text>x0 : -2;
f0 : 1;
x1 : 0;
f1 : sin(x);
pg: f0*(1-unit_step(x-x0)) + f1*unit_step(x-x1) +10000*(unit_step(x-x0)-unit_step(x-x1))+und*kron_delta(x,x0)+und*kron_delta(x,x1);
CS:matrix([1,x0,x0^2,x0^3],[0,1,2*x0,3*x0^2],[1,x1,x1^2,x1^3],[0,1,2*x1,3*x1^2]);
CT:transpose(matrix([ev(f0,x=x0),subst(x=x0,diff(f0,x)),ev(f1,x=x1),subst(x=x1,diff(f1,x))]));
CV:transpose(matrix([a0,a1,a2,a3]));
CS2:invert(CS).CT;
ta : matrix([1,x,x^2,x^3]).CS2;
tap: 1000*(1-unit_step(x-x0)) + 10000*unit_step(x-x1) +ta*(unit_step(x-x0)-unit_step(x-x1))+und*kron_delta(x,x0)+und*kron_delta(x,x1);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>anspt : 10000*(1-unit_step(x-x0)) + 10000*unit_step(x-x1) +ans1*(unit_step(x-x0)-unit_step(x-x1))+und*kron_delta(x,x0)+und*kron_delta(x,x1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check value of ans1 at x=x0.</description>
        <answertest>AlgEquiv</answertest>
        <sans>subst(x=x0,ans1)</sans>
        <tans>subst(x=x0,f0)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer does not satisfy \({@p(x0)=subst(x=x0,f0)@}\).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check value of ans1 at x=x1.</description>
        <answertest>AlgEquiv</answertest>
        <sans>subst(x=x1,ans1)</sans>
        <tans>subst(x=x1,f1)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer does not satisfy \({@p(x1)=subst(x=x1,f1)@}\).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check value of derivative of ans1 at x=x0.</description>
        <answertest>AlgEquiv</answertest>
        <sans>subst(x=x0,diff(ans1,x))</sans>
        <tans>subst(x=x0,diff(f0,x))</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer does not satisfy \(p'({@x0@})={@subst(x=x0,diff(f0,x))@}\).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Check value of derivative of ans1 at x=x1.</description>
        <answertest>AlgEquiv</answertest>
        <sans>subst(x=x1,diff(ans1,x))</sans>
        <tans>subst(x=x1,diff(f1,x))</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer does not satisfy \(p'({@x1@})={@subst(x=x1,diff(f1,x))@}\).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Check ans1 is a cubic.</description>
        <answertest>AlgEquiv</answertest>
        <sans>degree(ans1,x)</sans>
        <tans>3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer is not a cubic!</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is plotted below.</p>
<p>{@plot([pg,anspt],[x,(x0-3),(x1+5)],[y,-3,3])@}</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer is plotted below, although part of your graph might appear out of range of the plot!</p>
<p>{@plot([pg,anspt],[x,(x0-3),(x1+5)],[y,-3,3])@}</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-6-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>x^2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.2500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-6-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>x^2*(x+3)/4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.7500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-6-F</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>