<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-15-partialfrac-7</text>
    </name>
    <questiontext format="html">
      <text>Find \[ \int {@p@} d{@v@}. \]

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[In order to calculate \[\int {@p@} d{@v@}\] we need to express \[{@p@}\] as the sum of its partial fractions <hint>alg_partial_fractions</hint>.

\[\begin{array}{rcl}\displaystyle 
 {@p@} &\displaystyle =&\displaystyle  {@factor(p)@} \\
   &\displaystyle =&\displaystyle  {@A@}{@ppp1@} + {@B@}{@ppp2@} + {@C@}{@ppp3@}
 \end{array}\] 

By multiplying both sides by {@1/ppp1/ppp2/ppp3@} we get

\[\begin{array}{rcl}\displaystyle 
 {@simplify(p/ppp1/ppp2/ppp3)@} &\displaystyle =&\displaystyle  {@A/ppp2/ppp3@} + {@B/ppp1/ppp3@} + {@C/ppp1/ppp2@} \\
   &\displaystyle =&\displaystyle  {@expand(A/ppp2/ppp3)+expand(B/ppp1/ppp3)+expand(C/ppp1/ppp2)@}
 \end{array}\] 

Equating coefficients of {@v@} we get a system of equations that can be solved to get \(A = {@rdm4@}\), \(B = {@rdm5@}\) and \(C = {@rdm6@}\).

So, \[{@p@} = {@q@}\]

Hence,

\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int {@q@} d{@v@} \\
   &\displaystyle =&\displaystyle  {@rdm4@} \int {@ppp1@} d{@v@} + {@rdm5@} \int {@ppp2@} d{@v@} + {@rdm6@} \int {@ppp3@} d{@v@} \\
   &\displaystyle =&\displaystyle  {@rdm4*int(ppp1,v)@} + {@rdm5*int(ppp2,v)@} + {@rdm6*int(ppp3,v)@} +c \\
   &\displaystyle =&\displaystyle  {@rdm4*int(ppp1,v)+rdm5*int(ppp2,v)+rdm6*int(ppp3,v)@} +c
 \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm1:1+2*rand(3);
rdm2:1+rand(3);
rdm3:-1-rand(5);
rdm4:-1-rand(2);
rdm5:1+rand(3);
rdm6:-1-rand(3);
ppp1:1/(2*v+rdm1);
ppp2:1/(v+rdm2);
ppp3:1/(v+rdm3);
pp1:rdm4*ppp1;
pp2:rdm5*ppp2;
pp3:rdm6*ppp3;
q:pp1+pp2+pp3;
p:factor(simplify(q));
dumvar15:A;
dumvar16:B;
dumvar17:C;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int {@p@}d{@v@} = {@int(p,v)@}+c\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v)+c</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>int(p,v)+c</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1809810465</deployedseed>
    <deployedseed>204224421</deployedseed>
    <deployedseed>1336056803</deployedseed>
    <deployedseed>1346788195</deployedseed>
    <deployedseed>817748374</deployedseed>
    <deployedseed>1192500312</deployedseed>
    <deployedseed>1725374172</deployedseed>
    <deployedseed>1586940241</deployedseed>
    <deployedseed>746101730</deployedseed>
    <deployedseed>1686950002</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v),simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
