<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Diagnostic_ALG-IPF_2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Express \[{@f@}\] as a sum of a linear function and a rational function with a linear numerator and denominator {@den@}. [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Below are the rules of partial fractions.[[facts:alg_partial_fractions]]<br></p><p><br></p><p>We note that the numerator {@num@} is cubic and that the denominator {@den@} is quadratic. Dividing {@num@} by {@den@} we obtain a quotient with a linear numerator of the form {@p@} and a linear term of the form {@q@}, that is \[{@f@}=\frac{@p@}{@den@}+{@q@}.\] Multiplying through by {@den@} \[{@num@} = ({@den@})({@q@})+{@p@}.\] Expanding out the brackets we find \[{@num@} = {@denq@}+{@p@}.\] Matching coefficients in \(x^3\) we see \(m={@mval@}\). Matching coefficients in lower order terms and substituting we find that \(n={@nval@}\), \(g={@gval@}\) and \(h={@hval@}\). Therefore \[{@f@}={@answer@}.\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text>aval:rand(8) + 2;
bval:2*rand(2)-1;
cval:rand(8)+2;
mag_d:rand(4)+2;
sign_d:2*rand(2)-1;
dval:mag_d*sign_d;
gval:rand(8) + 2;
mag_h:rand(4)+2;
sign_h:2*rand(2)-1;
hval:mag_h*sign_h;
mval:rand(8) + 2;
mag_n:rand(4)+2;
sign_n:2*rand(2)-1;
nval:mag_n*sign_n;
den:aval * x^2 + bval * x + cval;
num:expand(den * (gval*x + hval)) + mval*x + nval;
f:num/den;
p:g*x+h;
q:m*x+n;
denq:expand(den*q);
answer:partfrac(f,x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[{@f@}={@answer@}.\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>answer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>divide</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>answer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ALG-IPF-TRUE</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ALG-IPF-FALSE</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>140638837</deployedseed>
    <deployedseed>1546222404</deployedseed>
    <deployedseed>1156854490</deployedseed>
    <deployedseed>1682760833</deployedseed>
    <deployedseed>640465359</deployedseed>
    <deployedseed>460795982</deployedseed>
    <deployedseed>461414454</deployedseed>
    <deployedseed>1472505969</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>answer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ALG-IPF-TRUE</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>f</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>ALG-IPF-FALSE</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
