<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>AlgMap-13.2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Factorise the following: 
\[ {@expand((cf1*v+a)*(cf2*v+b))@} .\] <br>
[[input:ans]] [[validation:ans]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text>Factorising a quadratic with an \(x^2\) coefficient equal to 1 is a relatively simple task. As can be seen by multiplying out a general, factorised quadratic:
\[ (x+a)(x+b) = x^2+(a+b)x+ab \] 
the task is to find values, \(a\) and \(b\), such that their sum is equal to the \(x\) coefficient and their product is equal to the constant term. In this example:
\[ {@expand((cf1*v+a)*(cf2*v+b))@}. \]
Equating with the general quadratic above gives:
\[{a+b = {@a+b@}}.\]
\[{ab = {@a*b@}}.\]
With practice you should be able to solve these simultaneous equations very quickly by inspection. We know that \(a\) and \(b\) are factors of {@a*b@} so there are just a few cases to check. The correct values are a = {@a@} and b = {@b@} (or vice versa). By substituting these back into the general quadratic we obtain the factored form: \[{@ta@}.\]</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>v:rand([x,y,z,t]);
cf1:1;
cf2:1;
a:(rand(9)+1);
b:(rand(9)+1);
ta:factor((cf1*v+a)*(cf2*v+b));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ {@expand((cf1*v+a)*(cf2*v+b))@} = {@ta@}.\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>FacForm</answertest>
        <sans>ans</sans>
        <tans>ta</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>1-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>335900966</deployedseed>
    <deployedseed>1929030640</deployedseed>
    <deployedseed>88359485</deployedseed>
    <deployedseed>408309564</deployedseed>
    <deployedseed>562239590</deployedseed>
    <deployedseed>89956134</deployedseed>
    <deployedseed>472550737</deployedseed>
    <deployedseed>299973407</deployedseed>
    <deployedseed>397577619</deployedseed>
    <deployedseed>2037020133</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans</name>
        <value>ev(expand((cf1*v+a)*(cf2*v+b)),simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>1-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
