<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>16.4.5.4 Differentiate to obtain power series</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>This exercise assumes knowledge of differentiation, this can be ignored if you have not seen differentiation before.<br><br>By differentiating the series for {@ser1@} find the power series for {@ser2@} and give a quadratic approximation. State its radius of convergence of the power series for {@ser2@}.</p>
<p>\({@ser2@} \approx\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]<br></p>
<p>\(R = \) [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">Using the known result from the binomial series \[(1+x)^p = 1 + px + \frac{p(p-1)}{2!}x^2 + \dots + \frac{p(p-1)(p-2)\cdots(p-n+1)}{n!}x^n \qquad |x|&lt;1,\] choose \(p={@po1@}\) and by differentiating obtain the power series expression for \({@ser2@}.\)</p>
<p dir="ltr" style="text-align: left;">\[ {@ser1@} = {@taylor(ser1, x,0,3)@} \]</p>
<p dir="ltr" style="text-align: left;">Differentiating both sides: \( \quad \displaystyle {@po1*ser2@} = {@diff(taylor(ser1, x,0,3), x)@}\)&nbsp;</p>
<p dir="ltr" style="text-align: left;">Multiplying through by {@a1@}: \( \quad \displaystyle {@ser2@} = {@ta1@} \)</p>
<p dir="ltr" style="text-align: left;">Using the sum notation, the binomial series is: \[(1+x)^p = \sum_{n=0}^{\infty} \frac{p(p-1)(p-2)\cdots(p-n-1)}{n!}x^n\]</p>
<p dir="ltr" style="text-align: left;">Here we have: \[b_p = \dfrac{p(p+1)\cdots(p-n+1)}{n!} \quad b_{p+1}=\dfrac{p(p-1)\cdots(p-n+1)(p-n)}{(n+1)!}\]</p>
<p dir="ltr" style="text-align: left;">In this case, \[R= \lim_{p \to \infty}\left| \dfrac{b_p}{b_{p+1}}\right| =\lim_{p \to \infty} \left| \dfrac{p(p-1)\cdots(p-n+1)(p-n)}{(n+1)!} \dfrac{n!}{p(p+1)\cdots(p-n+1)}\right|= \]<br></p>
<p dir="ltr" style="text-align: left;">\[ \lim_{p \to \infty} \frac{|p-n|}{n+1} = \lim_{p \to \infty} \frac{\left| 1- \frac{p}{n}\right|}{1+\frac{1}{n}} =1\]</p><p dir="ltr" style="text-align: left;">As we learned in the "Properties of power series" section, if \(P_1\) is a power series with radius of convergence \(R_1\) then \[ \dfrac{\mathrm{d}}{\mathrm{d}x}(P_1) \quad \mathrm{and} \quad \int(P_1)\mathrm{d}x\] are each power series with radius of convergence \(R_1.\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>simp:false;
ser_gen: (1+x)^p;
a1: 3;
po1: 1/a1;
po2: ev(po1-1,simp);
ser1: (1+x)^po1;
ser2: (1+x)^po2;
ta1: taylor(ser2, x, 0,2);
ta2: 1</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
