<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>43.1.2.6 predict elasticity</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Suppose we wish to be able to predict the modulus of rupture of a piece of timber from its modulus of elasticity. We wish to do this because to find the modulus of rupture of a piece of timber we need to destroy it, whereas if we can predict it reliably from its modulus of elasticity we do not need to destroy it. We have 15 observations as follows:&nbsp;</p><table style="text-align: center; width: 100%;">
<caption></caption>
<thead>
<tr>
<th scope="col"></th>
<th scope="col">Modulus of Elasticity (N/mm<sup>2</sup>×10<sup>3</sup>)</th>
<th scope="col">Modulus of Rupture (N/mm<sup>2</sup>)</th><th scope="col">&nbsp; &nbsp; &nbsp;MoE<sup>2</sup>&nbsp; &nbsp; &nbsp;</th><th scope="col">&nbsp; &nbsp; &nbsp;MoE×MoR&nbsp; &nbsp; &nbsp;</th>
<th scope="col">&nbsp; &nbsp; &nbsp;MoR<sup>2</sup>&nbsp; &nbsp; &nbsp;</th>
</tr>
</thead>
<tbody>
<tr>
<td></td>
<td>{@moe[1]@}</td>
<td>{@mor[1]@}</td>
<td>{@moe[1]*moe[1]@}</td>
<td>{@moe[1]*mor[1]@}</td>
<td>{@mor[1]^2@}<br></td>
</tr>
<tr>
<td></td>
<td>{@moe[2]@}</td>
<td>{@mor[2]@}</td>
<td>{@moe[2]*moe[2]@}</td>
<td>{@moe[2]*mor[2]@}</td>
<td>{@mor[2]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[3]@}</td>
<td>{@mor[3]@}</td>
<td>{@moe[3]*moe[3]@}</td>
<td>{@moe[3]*mor[3]@}</td>
<td>{@mor[3]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[4]@}</td>
<td>{@mor[4]@}</td>
<td>{@moe[4]*moe[4]@}</td>
<td>{@moe[4]*mor[4]@}</td>
<td>{@mor[4]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[5]@}</td>
<td>{@mor[5]@}</td>
<td>{@moe[5]*moe[5]@}</td>
<td>{@moe[5]*mor[5]@}</td>
<td>{@mor[5]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[6]@}</td>
<td>{@mor[6]@}</td>
<td>{@moe[6]*moe[6]@}</td>
<td>{@moe[6]*mor[6]@}</td>
<td>{@mor[6]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[7]@}</td>
<td>{@mor[7]@}</td>
<td>{@moe[7]*moe[7]@}</td>
<td>{@moe[7]*mor[7]@}</td>
<td>{@mor[7]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[8]@}</td>
<td>{@mor[8]@}</td>
<td>{@moe[8]*moe[8]@}</td>
<td>{@moe[8]*mor[8]@}</td>
<td>{@mor[8]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[9]@}</td>
<td>{@mor[9]@}</td>
<td>{@moe[9]*moe[9]@}</td>
<td>{@moe[9]*mor[9]@}</td>
<td>{@mor[9]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[10]@}</td>
<td>{@mor[10]@}</td>
<td>{@moe[10]*moe[10]@}</td>
<td>{@moe[10]*mor[10]@}</td>
<td>{@mor[10]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[11]@}</td>
<td>{@mor[11]@}</td>
<td>{@moe[11]*moe[11]@}</td>
<td>{@moe[11]*mor[11]@}</td>
<td>{@mor[11]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[12]@}</td>
<td>{@mor[12]@}</td>
<td>{@moe[12]*moe[12]@}</td>
<td>{@moe[12]*mor[12]@}</td>
<td>{@mor[12]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[13]@}</td>
<td>{@mor[13]@}</td>
<td>{@moe[13]*moe[13]@}</td>
<td>{@moe[13]*mor[13]@}</td>
<td>{@mor[13]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[14]@}</td>
<td>{@mor[14]@}</td>
<td>{@moe[14]*moe[14]@}</td>
<td>{@moe[14]*mor[14]@}</td>
<td>{@mor[14]^2@}</td>
</tr>
<tr>
<td></td>
<td>{@moe[15]@}</td>
<td>{@mor[15]@}</td>
<td>[[input:ans1]]<br></td>
<td>[[input:ans2]]<br></td>
<td>[[input:ans3]]</td>
</tr>
<tr>
<td><b>TOTAL</b></td>
<td>[[input:ans4]]</td>
<td>[[input:ans5]]</td>
<td>[[input:ans6]]</td>
<td>[[input:ans7]]</td>
<td>[[input:ans8]]</td></tr></tbody></table><p><br>Since we wish to predict the modulus of rupture, this is the dependent variable and we would plot this on the \(y\)-axis and the modulus of elasticity as the independent variable would go on the \(x\)-axis. The data suggests a straight line relationship would be reasonable.
</p><ol class="HELM_parts">

<li> Complete the table and use it to find the intercept and slope of the linear regression line. (Do not round the answers for the table, but give the intercept and slope to 3 significant figures).<p></p><p>slope = [[input:ans9]] [[validation:ans9]] [[feedback:prt9]]<br></p><p>intercept = [[input:ans10]] [[validation:ans10]] [[feedback:prt10]]</p>

</li><li>Use your model to predict what we should expect the modulus of rupture to be for a piece of timber that has a modulus of elasticity of {@var@} (<b>N/mm</b><sup style="color: inherit;">2</sup><span style="color: inherit; font-size: 10pt;">×10</span><sup style="color: inherit;">3</sup><span style="color: inherit; font-size: 10pt;">). Give your answer to 3 significant figures. [[input:ans11]] [[validation:ans11]] [[feedback:prt11]]

<p><span style="color: inherit; font-size: 10pt;">[[validation:ans8]] [[feedback:prt8]]</span><br></p><p></p><p>[[validation:ans1]] [[feedback:prt1]]&nbsp;<br></p><p>[[validation:ans2]]&nbsp;[[feedback:prt2]]&nbsp;</p><p>[[validation:ans3]] [[feedback:prt3]]&nbsp;</p><p>[[validation:ans4]] [[feedback:prt4]]&nbsp;</p><p>[[validation:ans5]] [[feedback:prt5]]&nbsp;</p><p>[[validation:ans6]] [[feedback:prt6]]&nbsp;</p><p>[[validation:ans7]] [[feedback:prt7]]</p></li></ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<b><span></span></b><b>FEEDBACK&nbsp;</b><p><b>a)</b>&nbsp;n = 15, \( \Sigma x_i = {@totalmoe@}\), \( \Sigma y_i = {@totalmor@}\)&nbsp;</p><p>&nbsp; &nbsp; \( \Sigma x_i^2 = {@totalmoe2@}\), \( \Sigma x_i y_i = {@totalmoemor@}\)</p><p></p><p>Hence we can calculate the slope, \(b\), using:&nbsp;</p><p>\[ b = \dfrac{n\times\sum_{i=1}^nx_iy_i-\sum_{i=1}^nx_i\times\sum_{i=1}^ny_i}{n\times\sum_{i=1}^nx_i^2-\left(\sum_{i=1}^nx_i\right)^2} \]&nbsp;</p><p></p><p>&nbsp; &nbsp; \(b\) = \( \dfrac{15\times {@totalmoemor@} - {@totalmoe@} \times {@totalmor@} }{ 15\times {@totalmoe2@} - {@totalmoe@}^2 } \)</p><p>&nbsp; &nbsp; &nbsp; &nbsp;= \( \dfrac{ {@15*totalmoemor@} - {@totalmoe*totalmor@} }{ {@15*totalmoe2@} - {@totalmoe^2@} } \)</p><p>&nbsp; &nbsp; &nbsp; &nbsp;= \( \dfrac{ {@15*totalmoemor-totalmoe*totalmor@} }{ {@15*totalmoe2-totalmoe^2@} } \)</p><p>&nbsp; &nbsp; &nbsp; &nbsp;= {@slope@} = {@ta9@} (3 s.f)&nbsp;</p><p>&nbsp; &nbsp; Using \( a = \bar{y}-b\bar{x} \) we first calculate</p><p>&nbsp; &nbsp; &nbsp; &nbsp; \( \bar{y} = \dfrac{ {@totalmor@} }{ 15 } = {@float(totalmor/15)@} \) and \( \bar{x} = \dfrac{ {@totalmoe@} }{ 15 } = {@float(totalmoe/15)@} \)</p><p>&nbsp; &nbsp; then we find that</p><p>&nbsp; &nbsp; &nbsp; &nbsp; \(a\) = \({@totalmor/15@} - {@slope@} \times {@float(totalmoe/15)@} \)</p><p>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp; = {@intercept@} = {@ta10@} (3 s.f)</p><p>&nbsp; &nbsp; Hence our fitted straight line is</p><p>&nbsp; &nbsp; &nbsp; &nbsp; MoR = {@intercept@} + {@slope@}MoE</p><p><b>b)</b>&nbsp;If x = {@var@}, MoR = {@intercept@} + ({@slope@})({@var@}) = {@intercept+var*slope@}</p><p>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp;= {@ta11@} (3 s.f)</p><p><br></p><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>12</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* this contains the pearson coefficient question (ex06q2) */
moe : [9.5,9,11,11.5,4.5,6,13.5,6.5,8,5.5,15,12.5,15.5,8.5,10.5]
mor : [47.5,53,59,66.5,34.5,43,72.5,41.5,51,31.5,79,63.5,88.5,44.5,64.5]

/*moe : makelist(float((10+rand(22))/2),15)
mor : makelist(float(ceiling(2*(10.339+4.659*moe[i])-3+(rand(60)/10))/2),i,15)*/

ta1 : moe[15]^2
ta2 : moe[15]*mor[15]
ta3 : mor[15]^2
totalmoe : 0 ; totalmor : 0 ; totalmoe2 : 0 ; totalmoemor : 0 ; totalmor2 : 0
dum1 : for i:1 thru 15 do totalmoe : totalmoe+moe[i]
dum2 : for i:1 thru 15 do totalmor : totalmor +mor[i]
dum3 : for i:1 thru 15 do totalmoe2 : totalmoe2 + moe[i]*moe[i]
dum4 : for i:1 thru 15 do totalmoemor : totalmoemor + moe[i]*mor[i]
dum5 : for i:1 thru 15 do totalmor2 : totalmor2 + mor[i]*mor[i]
ta4 : totalmoe
ta5 : totalmor
ta6 : totalmoe2
ta7 : totalmoemor
ta8 : totalmor2
slope : ((15*totalmoemor)-(totalmoe*totalmor))/((15*totalmoe2)-totalmoe^2)
ta9 : significantfigures(slope,3)
intercept : (totalmor/15)-slope*(totalmoe/15)
ta10 : significantfigures(intercept, 3)
var : 6+rand(6)
ta11 : significantfigures(intercept+var*slope,3)
alternateta11 : significantfigures(ta9+var*ta10,3)

/* I made it so that if they use the intercept and slope rounded to 3 sf to calculate their answer for part (b) they still get full marks. This can be changed by deleting the second node in prt11. They still need to have entered their answer to 3 significant figures though. */

unroundedta12 : (15*totalmoemor-totalmoe*totalmor)/sqrt((15*totalmoe2-totalmoe^2)*(15*totalmor2-totalmor^2)) 
ta12 : significantfigures(unroundedta12,3)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>ta1-ta8 are supposed to be entered unrounded.
slope (ta9) = {@ta9@} intercept (ta10) = {@ta10@}, ta11 = {@ta11@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>numerical</type>
      <tans>ta10</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>minsf:3, maxsf:3</options>
    </input>
    <input>
      <name>ans11</name>
      <type>numerical</type>
      <tans>ta11</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>minsf:3, maxsf:3</options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>numerical</type>
      <tans>ta9</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>minsf:3, maxsf:3</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>ta10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>ta11</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>alternateta11</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1174544768</deployedseed>
    <deployedseed>936439243</deployedseed>
    <deployedseed>1544943467</deployedseed>
    <deployedseed>1160848281</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ta10</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ta11</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta9</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ta10+1</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ta11+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4+1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5+1</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6+1</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7+1</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8+1</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta9+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
