<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>43.1.4.4 fake the data</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A corrupt scientist is determined to to have everyone believe their theory that the relationship between the length of an octopus and the number of fish it eats in a \(24\) hour period is a simple linear model.</p><p>Unfortunately the data collected show that this is not the case.</p> <p>Help the scientist manipulate the data in order to fool everyone that it does satisfy a linear model.</p>

<table border="1"><caption>&nbsp;</caption><thead><tr><th scope="col">Length of octopus (cm)</th><th scope="col">Number of fish eaten in a \(24\) hour period</th></tr></thead><tbody><tr><td style="text-align: center;">[[input:ans1]] [[validation:ans1]]</td><td style="text-align: center;">[[input:bns1]] [[validation:bns1]]</td></tr><tr><td style="text-align: center;">[[input:ans2]] [[validation:ans2]]</td><td style="text-align: center;">[[input:bns2]] [[validation:bns2]]</td></tr><tr><td style="text-align: center;">[[input:ans3]] [[validation:ans3]]</td><td style="text-align: center;">[[input:bns3]] [[validation:bns3]]</td></tr><tr><td style="text-align: center;">[[input:ans4]] [[validation:ans4]]</td><td style="text-align: center;">[[input:bns4]] [[validation:bns4]]</td></tr><tr><td style="text-align: center;">[[input:ans5]] [[validation:ans5]]</td><td style="text-align: center;">[[input:bns5]] [[validation:bns5]]</td></tr><tr><td style="text-align: center;">[[input:ans6]] [[validation:ans6]]</td><td style="text-align: center;">[[input:bns6]] [[validation:bns6]]</td></tr><tr><td style="text-align: center;">[[input:ans7]] [[validation:ans7]]</td><td style="text-align: center;">[[input:bns7]] [[validation:bns7]]</td></tr><tr><td style="text-align: center;">[[input:ans8]] [[validation:ans8]]<br></td><td style="text-align: center;">[[input:bns8]] [[validation:bns8]]</td></tr><tr><td style="text-align: center;">[[input:ans9]] [[validation:ans9]]<br></td><td style="text-align: center;">[[input:bns9]] [[validation:bns9]]</td></tr><tr><td style="text-align: center;">[[input:ans10]] [[validation:ans10]]<br></td><td style="text-align: center;">[[input:bns10]] [[validation:bns10]]</td></tr><tr><td style="text-align: center;">[[input:ans11]] [[validation:ans11]]<br></td><td style="text-align: center;">[[input:bns11]] [[validation:bns11]]</td></tr><tr><td style="text-align: center;">[[input:ans12]] [[validation:ans12]]<br></td><td style="text-align: center;">[[input:bns12]] [[validation:bns12]]</td></tr></tbody></table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2024072400</text>
    </stackversion>
    <questionvariables>
      <text>rows: 12;
criticalF: 241.9; /* from table for sig.05, df 1,10 */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[[[feedback:bns12]]<br>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>no rand</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>8</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>numerical</type>
      <tans>ta10</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>20</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>numerical</type>
      <tans>ta11</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>21</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>numerical</type>
      <tans>ta12</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>35</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>9</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>11</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>11</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>13</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>numerical</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>15</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>numerical</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>15</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>numerical</type>
      <tans>ta8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>17</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>numerical</type>
      <tans>ta9</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>20</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns1</name>
      <type>numerical</type>
      <tans>tb1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>2</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns10</name>
      <type>numerical</type>
      <tans>tb10</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>4</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns11</name>
      <type>numerical</type>
      <tans>tb11</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>4</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns12</name>
      <type>numerical</type>
      <tans>tb12</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>6</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns2</name>
      <type>numerical</type>
      <tans>tb2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>3</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns3</name>
      <type>numerical</type>
      <tans>tb3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>2</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns4</name>
      <type>numerical</type>
      <tans>tb4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>5</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns5</name>
      <type>numerical</type>
      <tans>tb5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>3</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns6</name>
      <type>numerical</type>
      <tans>tb6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>1</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns7</name>
      <type>numerical</type>
      <tans>tb7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>3</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns8</name>
      <type>numerical</type>
      <tans>tb8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>1</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>bns9</name>
      <type>numerical</type>
      <tans>tb9</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>2</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>bns12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/**calculate SSt**/
sumysquared: bns1^2+bns2^2+bns3^2+bns4^2+bns5^2+bns6^2+bns7^2+bns8^2+bns9^2+bns10^2+bns11^2+bns12^2;
sumy: bns1+bns2+bns3+bns4+bns5+bns6+bns7+bns8+bns9+bns10+bns11+bns12;
meany: sumy/rows;
SSt: sumysquared - rows*meany^2;

/**calculate a and b in y=a+bx**/
sumx: ans1+ans2+ans3+ans4+ans5+ans6+ans7+ans8+ans9+ans10+ans11+ans12;
meanx: sumx/rows;
sumxsquared: ans1^2+ans2^2+ans3^2+ans4^2+ans5^2+ans6^2+ans7^2+ans8^2+ans9^2+ans10^2+ans11^2+ans12^2;
sumxy: ans1*bns1+ans2*bns2+ans3*bns3+ans4*bns4+ans5*bns5+ans6*bns6+ans7*bns7+ans8*bns8+ans9*bns9+ans10*bns10+ans11*bns11+ans12*bns12;
b: ((sumxy/rows)-(sumx/rows)*(sumy/rows))/((sumxsquared/rows)-(sumx/rows)^2);
a: meany-b*meanx;

/**get predicted values**/
prd1: a+b*ans1;
prd2: a+b*ans2;
prd3: a+b*ans3;
prd4: a+b*ans4;
prd5: a+b*ans5;
prd6: a+b*ans6;
prd7: a+b*ans7;
prd8: a+b*ans8;
prd9: a+b*ans9;
prd10: a+b*ans10;
prd11: a+b*ans11;
prd12: a+b*ans12;

/** get error squares**/
SSe: (bns1-prd1)^2+(bns2-prd2)^2+(bns3-prd3)^2+(bns4-prd4)^2+(bns5-prd5)^2+(bns6-prd6)^2+(bns7-prd7)^2+(bns8-prd8)^2+(bns9-prd9)^2+(bns10-prd10)^2+(bns11-prd11)^2+(bns12-prd12)^2;

/**finally the values we actually want**/
SSr:  SSt - SSe;
Ftest: decimalplaces(float(SSr/(SSe/(rows-2))),1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>GT</answertest>
        <sans>Ftest</sans>
        <tans>criticalF</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>bns12-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The test statistic for the data you have faked is 
</p>
<p>\(F_{test} = {@Ftest@}\) which, looking it up in a table, is above the critical value \(f_{0.05,1,10} = 241.9.\)</p>
<p>So well done, the simple linear model is appropriate, and you have helped the scientist contribute to fake news.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>bns12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The test statistic for the data you have faked is 
</p>
<p>\(F_{test} = {@Ftest@}\) which, looking it up in a table, is not above the critical value \(f_{0.05,1,10} = 241.9.\)</p>
<p>Therefore, the simple linear model is not appropriate for these data. Please try again</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ta10</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ta11</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>ta12</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta9</value>
      </testinput>
      <testinput>
        <name>bns1</name>
        <value>tb1</value>
      </testinput>
      <testinput>
        <name>bns10</name>
        <value>tb10</value>
      </testinput>
      <testinput>
        <name>bns11</name>
        <value>tb11</value>
      </testinput>
      <testinput>
        <name>bns12</name>
        <value>tb12</value>
      </testinput>
      <testinput>
        <name>bns2</name>
        <value>tb2</value>
      </testinput>
      <testinput>
        <name>bns3</name>
        <value>tb3</value>
      </testinput>
      <testinput>
        <name>bns4</name>
        <value>tb4</value>
      </testinput>
      <testinput>
        <name>bns5</name>
        <value>tb5</value>
      </testinput>
      <testinput>
        <name>bns6</name>
        <value>tb6</value>
      </testinput>
      <testinput>
        <name>bns7</name>
        <value>tb7</value>
      </testinput>
      <testinput>
        <name>bns8</name>
        <value>tb8</value>
      </testinput>
      <testinput>
        <name>bns9</name>
        <value>tb9</value>
      </testinput>
      <expected>
        <name>bns12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>bns12-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
