<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>4.5.4.8 Exercise 8 (Amp Phase x2) (NOT IN QUIZ)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Given that if \( A\cos\omega t+B\sin\omega t= C\sin (\omega t+\alpha)\)&nbsp; then \[\quad C=\sqrt{A^2+B^2},\qquad \ \cos\alpha=\dfrac{B}{C}, \qquad \sin\alpha=\dfrac{A}{C}.\]</p><p></p><p>Express the following in the amplitude/phase form \(&nbsp;C\sin (\omega t+\alpha). \) Approximate all floating point numbers up to \(4\) d.p.</p>

<ol class="HELM_parts">
<li><p>\(y= {@na3a@}\cos({@na3c@}t)- {@na3b@}\sin({@na3c@}t), \)</p>
<p>\( y = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p></li>

<li><p>\(y= {@na3a@}\cos({@na3c@}t)+ {@na3b@}\sin({@na3c@}t), \)</p>
<p>\( y = \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Since \[\sin(\omega 
t+\alpha)=\sin\omega t\cos\alpha+\cos\omega t\sin \alpha\] then 
\(A=C\sin\alpha\) (coefficients of \(\cos\omega t\)) and \(B=C\cos\alpha\) 
(coefficients of \(\sin\omega t\),) from which \[C^2=A^2+B^2,\quad  \sin\alpha 
=\dfrac{A}{C},\quad \cos \alpha=\dfrac{B}{C}.\]</p>

<ol class="HELM_parts">
<li>
<p>\( C=\sqrt{{@na3a@}^2+ {@na3b@}^2}, \) &nbsp; \( 
\cos\alpha=-\dfrac{{@na3b@}}{{@ta3c@}}\) and \( \sin 
\alpha-\dfrac{{@na3a@}}{{@ta3c@}},\)&nbsp;</p>
<p>so \(\alpha\) is in the second 
quadrant and \(\alpha={@ta3a@}.\)</p>
<p>Therefore \(y={@ta3f@} \).</p>
</li>
<li><p>Similarly, \(y={@ta3f2@} \).</p>
</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021071900</text>
    </stackversion>
    <questionvariables>
      <text>na3a: sqrt(rand(4)+1);
na3b: rand(5)+2;
na3c: rand(7)+1;
ta3c: sqrt(na3a^2+ na3b^2);
ta3a2: decimalplaces(float(atan(na3a/na3b)),4);
ta3a: pi- ta3a2; 
ta3f: ta3c*sin(na3c*t+ta3a);
ta3f2: ta3c*sin(na3c*t+ta3a2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[ y={@ta3f@}\]
\[ y={@ta3f2@}\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta3f</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta3f2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta3f</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans2</sans>
        <tans>ta3f2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1061784043</deployedseed>
    <deployedseed>823854388</deployedseed>
    <deployedseed>1707673676</deployedseed>
    <deployedseed>1728599275</deployedseed>
    <deployedseed>1090253770</deployedseed>
    <deployedseed>2065395685</deployedseed>
    <deployedseed>1358698760</deployedseed>
    <deployedseed>1928678592</deployedseed>
    <deployedseed>1796296040</deployedseed>
    <deployedseed>1065080259</deployedseed>
    <deployedseed>40325521</deployedseed>
    <deployedseed>910928509</deployedseed>
    <deployedseed>359610798</deployedseed>
    <deployedseed>1125292434</deployedseed>
    <deployedseed>392249549</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta3f</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta3f2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta3c</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta3c</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
