<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Differential_Equation_6</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve \[{@n1@} \frac{d^2y}{dx^2} - {@bc@} \frac{dy}{dx} + {@cc@} y = 0.\] Use A and B for the arbitrary constants. Warning: ensure you write "A*e^()", "A*exp()", "A*cos()" or "A*sin()" to avoid misinterpretation. <br /> y(x) = [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To solve the differential equation \[{@n1@} \frac{d^2y}{dx^2} - {@bc@} \frac{dy}{dx} + {@cc@} y = 0\] Let \(y = e^{mx}\) and hence \(\frac{dy}{dx}=me^{mx}\) and \(\frac{d^2y}{dx^2} = m^2e^{mx}\) <br />Substituting for \(y\) gives \[{@n1@}m^2e^{mx} - {@bc@}me^{mx} + {@cc@}e^{mx} = 0\] and dividing throughout by \(e^{mx}\) and {@n1@} gives \[m^2 - {@bc/n1@}m + {@cc/n1@} = 0\] which can be solved using the quadratic formula to give {@m@} <br />For roots of \(m\) in the form \(a±bi\), the solution to the differential equation is of the form \(y =e^{ax}(Acos(bx)+Bsin(bx))\) <br /> Therefore \(y =  {@e^(avar*x)@}(Acos({@bvar@}x)+Bsin({@bvar@}x))\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>n1:rand(8)+2;
   n2:rand(6)+1;
   n3:rand(6)+1;
   bc:n1*2*n2;
   cc:n1*(n2^2+n3^2);
  
   m:solve(n1*m^2 - bc*m + cc = 0,m);
   m1:rhs(m[1]);
   m2:rhs(m[2]);
   avar:realpart(m1);
   bvar:abs(imagpart(m1));
  
   correctanswer:e^(avar*x)*(A*cos(bvar*x)+B*sin(bvar*x));
   altanswer:e^(avar*x)*(B*cos(bvar*x)+A*sin(bvar*x));
   wronganswer1:e^(avar*x)*(Q*cos(bvar*x)+P*sin(bvar*x));
   wronganswer2:e^(avar*x)*B*cos(bvar*x);
   wronganswer3:e^(avar*x)*A*sin(bvar*x);
   wronganswer4:B*cos(bvar*x)+A*sin(bvar*x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@n1@} \frac{d^2y}{dx^2} - {@bc@} \frac{dy}{dx} + {@cc@} y = 0\) gives {@e^(avar*x)@}(Acos({@bvar@}x)+Bsin({@bvar@}x))</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>stu_rhs:ev(n1*diff(ans1,x,2)-bc*diff(ans1,x)+cc*ans1,simp);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>altanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>stu_rhs</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-2-T </trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer solves the differential equation, but is either not the most general solution or does not use the prescribed constants "A" and "B".</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-2-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>326758595</deployedseed>
    <deployedseed>1843097710</deployedseed>
    <deployedseed>1154517065</deployedseed>
    <deployedseed>1019539824</deployedseed>
    <deployedseed>1643015905</deployedseed>
    <deployedseed>436689476</deployedseed>
    <deployedseed>1627007879</deployedseed>
    <deployedseed>134268899</deployedseed>
    <deployedseed>1955854519</deployedseed>
    <deployedseed>806337565</deployedseed>
    <deployedseed>454366487</deployedseed>
    <deployedseed>911689889</deployedseed>
    <deployedseed>359566684</deployedseed>
    <deployedseed>347757847</deployedseed>
    <deployedseed>1939343800</deployedseed>
    <deployedseed>1461356209</deployedseed>
    <deployedseed>2003686058</deployedseed>
    <deployedseed>1582343755</deployedseed>
    <deployedseed>1197736067</deployedseed>
    <deployedseed>104242900</deployedseed>
    <deployedseed>1319754704</deployedseed>
    <deployedseed>270092979</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer-c</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>altanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer3</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer4</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
