<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>3.2.13 (Give a counterexample quadratic)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Give a counterexample to disprove the statement  "all roots of the quadratic {@expand(quad)@} are positive."</p>
<p>\(x=\)[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Note that we can factorise the quadratic as {@quad@}. Therefore the roots are
\[
x = {@ta1@} \qquad \mbox{and} \qquad x = {@bb@}.
\]
Taking \(x = {@ta1@}\) is a counter example to the statement.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020081700</text>
    </stackversion>
    <questionvariables>
      <text>aa:1+rand(10);
bb:1+rand(10);

quad: (x+aa)*(x-bb);

ta1:-aa;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@quad@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>255171298</deployedseed>
    <deployedseed>1275294073</deployedseed>
    <deployedseed>736006444</deployedseed>
    <deployedseed>1645842747</deployedseed>
    <deployedseed>810979837</deployedseed>
    <deployedseed>501179493</deployedseed>
    <deployedseed>2085806798</deployedseed>
    <deployedseed>1428908592</deployedseed>
    <deployedseed>1347010026</deployedseed>
    <deployedseed>429300237</deployedseed>
    <deployedseed>1208602350</deployedseed>
    <deployedseed>1933931210</deployedseed>
    <deployedseed>1439402811</deployedseed>
    <deployedseed>62474582</deployedseed>
    <deployedseed>1947358395</deployedseed>
    <deployedseed>939505020</deployedseed>
    <deployedseed>542193395</deployedseed>
    <deployedseed>433217633</deployedseed>
    <deployedseed>1301525381</deployedseed>
    <deployedseed>1020493883</deployedseed>
    <deployedseed>1727950270</deployedseed>
    <deployedseed>970816854</deployedseed>
    <deployedseed>208152873</deployedseed>
    <deployedseed>2001789918</deployedseed>
    <deployedseed>1207123922</deployedseed>
    <deployedseed>1545740956</deployedseed>
    <deployedseed>255961907</deployedseed>
    <deployedseed>2098284523</deployedseed>
    <deployedseed>631161808</deployedseed>
    <deployedseed>246293161</deployedseed>
    <deployedseed>1660879243</deployedseed>
    <deployedseed>2070147303</deployedseed>
    <deployedseed>195814343</deployedseed>
    <deployedseed>294889853</deployedseed>
    <deployedseed>2073622133</deployedseed>
    <deployedseed>894651232</deployedseed>
    <deployedseed>49636071</deployedseed>
    <deployedseed>310148927</deployedseed>
    <deployedseed>977359026</deployedseed>
    <deployedseed>206930677</deployedseed>
    <deployedseed>392799759</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>aa</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
