<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Matrices_Solve_Equation_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>If \[{\bf A}={@matrixa@},  {\bf B} = {@matrixb@},\]</p>
<p>find a 2 by 2 matrix \({\bf X}\) such that</p>
<p> \[{\bf A}\cdot{\bf X} = {\bf B}.\]</p>
<p> [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>For matrix \[{\bf A} = {@matrixa@}\]</p>
<p> find the inverse \[{\bf A^{-1}} = {@inva@}\]</p>
<p>Then \[{\bf X} = {\bf A^{-1}}\cdot{\bf B} ={@inva@} \cdot {@matrixb@}={@correctanswer@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text>matrixa:matrix([1,rand(6)+1],[0,1]);
 n1:rand(4)+2;
 matrixx:matrix([n1,-n1*(rand(9)+1)],[0,n1]);
 matrixb:matrixa . matrixx;
 correctanswer:matrixx;
 wronganswer1:matrixb.matrixa;
 inva:invert(matrixa);
 matrixx:inva.matrixb;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({\bf A} = {@matrixa@}, {\bf A} \cdot {\bf X} = {\bf B} = {@matrixb@}, {\bf X} ={@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>determinant,invert</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wronganswer1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>976080074</deployedseed>
    <deployedseed>2106134711</deployedseed>
    <deployedseed>1751612605</deployedseed>
    <deployedseed>1863697726</deployedseed>
    <deployedseed>1566840520</deployedseed>
    <deployedseed>911977444</deployedseed>
    <deployedseed>838012139</deployedseed>
    <deployedseed>735490211</deployedseed>
    <deployedseed>1707961224</deployedseed>
    <deployedseed>1367456570</deployedseed>
    <deployedseed>45556478</deployedseed>
    <deployedseed>2027515029</deployedseed>
    <deployedseed>502791254</deployedseed>
    <deployedseed>2109866972</deployedseed>
    <deployedseed>1661328744</deployedseed>
    <deployedseed>79724673</deployedseed>
    <deployedseed>973690359</deployedseed>
    <deployedseed>1941144962</deployedseed>
    <deployedseed>895514976</deployedseed>
    <deployedseed>1343575767</deployedseed>
    <deployedseed>1782613434</deployedseed>
    <deployedseed>819819563</deployedseed>
    <deployedseed>1178299975</deployedseed>
    <deployedseed>52041360</deployedseed>
    <deployedseed>972613035</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
