<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Matrices_Add_1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Let \({\bf A}\) and \({\bf B}\) be the following matrices: \[{\bf A} = {@matrixa@}\] <br />\[{\bf B} = {@matrixb@}\]</p>
<p>Find the sum \({\bf A + B:}\)</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\[{\bf A} = {@matrixa@}\] <br />and<br />\[{\bf B} = {@matrixb@}\] <br /> Therefore \[\begin{array}{rcc} {\bf A} + {\bf B} &amp; =&amp; \begin{bmatrix} {@matrixa[1,1]@} +{@matrixb[1,1]@} &amp; {@matrixa[1,2]@} +{@matrixb[1,2]@} \\ {@matrixa[2,1]@} +{@matrixb[2,1]@} &amp; {@matrixa[2,2]@} +{@matrixb[2,2]@} \end{bmatrix} \end{array}\] \[={@correctanswer@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text>matrixa:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
 matrixb:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
 correctanswer:matrixa + matrixb;
wronganswer:matrixa-matrixb;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({\bf A} = {@matrixa@}, {\bf B} = {@matrixb@}, {\bf A} + {\bf B}  = {@correctanswer@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>correctanswer</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>641496332</deployedseed>
    <deployedseed>1070002371</deployedseed>
    <deployedseed>1415945220</deployedseed>
    <deployedseed>1496950566</deployedseed>
    <deployedseed>1245121717</deployedseed>
    <deployedseed>848744367</deployedseed>
    <deployedseed>1739925625</deployedseed>
    <deployedseed>1556712773</deployedseed>
    <deployedseed>2052076846</deployedseed>
    <deployedseed>1014580985</deployedseed>
    <deployedseed>32695809</deployedseed>
    <deployedseed>881246372</deployedseed>
    <deployedseed>843841335</deployedseed>
    <deployedseed>193675110</deployedseed>
    <deployedseed>1907139196</deployedseed>
    <deployedseed>99870475</deployedseed>
    <deployedseed>1782060220</deployedseed>
    <deployedseed>431721198</deployedseed>
    <deployedseed>397590908</deployedseed>
    <deployedseed>730311019</deployedseed>
    <deployedseed>791671438</deployedseed>
    <deployedseed>1510005573</deployedseed>
    <deployedseed>1985211190</deployedseed>
    <deployedseed>37406481</deployedseed>
    <deployedseed>1847899318</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>wronganswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
