<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CR-Int-14-substitution-1.f</text>
    </name>
    <questiontext format="html">
      <text>Integrate {@p@} with respect to {@v@} using the substitution \(u={@ppp@}\).

[[input:ans1]] [[validation:ans1]]</text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[We need to calculate \[\int {@p@} d{@v@}\] by means of the substitution \(u={@ppp@}\) [[facts:calc_int_methods_substitution]].

So, if \(u = {@ppp@}\) then \({@v@} = {@q1@}\) and \(d{@u@} = {@diff(ppp,v)@} d{@v@}\), so, \(d{@v@} = d{@u@}/{@diff(ppp,v)@}\).

Hence, 

\[\begin{array}{rcl}\displaystyle  
\int {@p@} d{@v@} &\displaystyle =&\displaystyle  \int {@q1@}{@q2@} \frac{d{@u@}}{@diff(ppp,v)@} \\
   &\displaystyle =&\displaystyle  \int {@q3/diff(ppp,v)@} d{@u@} \\
   &\displaystyle =&\displaystyle  \int {@expand(q3/diff(ppp,v))@} d{@u@} \\
   &\displaystyle =&\displaystyle  {@int(expand(q3/diff(ppp,v)),u)+c@} \\
   &\displaystyle =&\displaystyle  {@int(p,v)+c@}
 \end{array}\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2019121800</text>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm1:2+rand(5);
rdm2:-1-rand(4);
ppp:rdm1*v+rdm2;
pp:ppp^(1/2);
p:v*pp;
dumvar6:u;
q1:rhs(solve(u=rdm1*v+rdm2,v)[1]);
q2:u^(1/2);
q3:q1*q2;
ta1:int(p,v)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:Result]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\[\int {@p@}d{@v@} = {@int(p,v)@}+c\]</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int(p,v)+c</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>331928844</deployedseed>
    <deployedseed>526945568</deployedseed>
    <deployedseed>659497263</deployedseed>
    <deployedseed>42601172</deployedseed>
    <deployedseed>1970146387</deployedseed>
    <deployedseed>1855493014</deployedseed>
    <deployedseed>784254259</deployedseed>
    <deployedseed>574808047</deployedseed>
    <deployedseed>741506496</deployedseed>
    <deployedseed>1449076209</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v)+c,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
