<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>23.4.2.1 FS expression 1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
Obtain the Fourier series of \[ f(t) = \lvert t \rvert \hspace{1cm} \pi \leq t \leq \pi \hspace{1cm} f(t + 2\pi) = f(t) \] and use it to calculate the summation \( \displaystyle{\sum _{n = 1}^\infty }\dfrac{1}{(2n - 1)^2} \) by putting \( t = 0\).<p>
</p><p></p>
\( \displaystyle{ \sum _{n = 1}^\infty } \dfrac{1}{(2n - 1)^2} = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]<p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[As \( f(t) \) is an even function, \( b_n = 0 \). For \( a_0\) we have
\[ 
\begin{align}
a_0 &amp;= \dfrac{1}{\pi} \displaystyle{ \int _{-\pi}^\pi } \lvert t \rvert \hspace{0.1cm} dt \\
&amp;= \dfrac{1}{\pi}  \displaystyle{ \int _{-\pi}^0 } -t \hspace{0.1cm} dt +  \dfrac{1}{\pi}  \displaystyle{ \int _{0}^\pi } t \hspace{0.1cm} dt \\
&amp;= \pi
\end{align}
\]
By using integration by parts, we then find
\[
\begin{align}
a_n &amp;= \dfrac{1}{\pi}  \displaystyle{ \int _{-\pi}^\pi } \lvert t \rvert \cos nt \hspace{0.1cm} dt \\
&amp;= \dfrac{1}{\pi}  \displaystyle{ \int _{-\pi}^0 } - t \cos nt \hspace{0.1cm} dt + \dfrac{1}{\pi}  \displaystyle{ \int _{0}^\pi }  t \cos nt \hspace{0.1cm} dt \\
&amp;= \dfrac{1}{\pi} \left( \left[ -\dfrac{t}{n} \sin nt \right]_{-\pi}^0 + \dfrac{1}{n} \displaystyle{ \int _{-\pi}^0 }  \sin nt \hspace{0.1cm} dt + \left[ \dfrac{t}{n} \sin nt \right]_{0}^\pi -  \dfrac{1}{n} \displaystyle{ \int _{0}^\pi }  \sin nt \hspace{0.1cm} dt \right) \\
&amp;= \dfrac{1}{\pi} \left( \left[ - \dfrac{1}{n^2} \cos nt  \right]_{-\pi}^0 + \left[  \dfrac{1}{n^2} \cos nt \right]_{0}^\pi \right) \\
&amp;= \dfrac{2}{n^2 \pi} \left( \cos n\pi - 1 \right)
\end{align}
\]
\[ =
\begin{cases}
0 \hspace{1cm} n = 2, 4, 6, \dots \\
-\dfrac{4}{n^2 \pi} \hspace{1cm} n = 1, 3, 5, \dots
\end{cases}
\]
Now rewriting \(n\) to take the form of an odd number as \( 2n - 1\) (where \(n = 1, 2, 3, 4, \dots\)), we have the Fourier series
\[ f(t) = \dfrac{\pi}{2} + \sum _{n=1}^\infty \dfrac{-4}{(2n-1)^2\pi} \cos \left[ (2n - 1)t \right] \]
Now setting \(t = 0\) we have \(f(t) = 0\), and therefore
\[ 0 = \dfrac{\pi}{2} +  \sum _{n=1}^\infty \dfrac{-4}{(2n-1)^2\pi} \]
which can be rearranged to give
\[  \sum _{n=1}^\infty \dfrac{1}{(2n-1)^2} = \dfrac{\pi^2}{8} \]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2021120900</text>
    </stackversion>
    <questionvariables>
      <text>ta1: pi^2/8</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
